```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab_3_plotting.ipynb")
```

## Lab 3: Plotting data

Read in the wrist torque data and plot it for each of the x, y, and z channels. I recommend doing this in the following order:

- Copy in your lecture activity code and do that plot first - make sure it's working and the titles/labels etc are correct 
- Once that's working, add a **for** loop to loop over x, y, z
- If doing a **for** loop is confusing, then start with just copying the code over, once, and changing the dimension you're using in the copy. Then go back and put the for loop in

[optional] Do nicer colors than the default ones.

For setting the colors and the data labels, you can either do it by hand, or setup a list of color names and use that
Note: Plot the data for the first pick (the first row of the spread sheet)

```python
# Install a pip package in the current Jupyter kernel
# Note: this is a safety check in case your JN install was not in the same place as your libraries
import sys
!{sys.executable} -m pip install matplotlib
```

```python
# Libraries that we need to import - numpy and json (for loading the description file)
import numpy as np
import json as json
import matplotlib.pyplot as plt
%matplotlib inline
```

```python
wrist_torque_data = ...
ts = ...

```

```python
nrows = 1
ncols = 1
_, axs_wf = plt.subplots(nrows, ncols, figsize=(6, 4))

# Useful for labeling the channels with x,y,z instead of 0, 1, 2
channel_labels = ['x', 'y', 'z']

...
```

```python
grader.check("plot_torque")
```

## Part 2 - plot both the wrist force and the wrist torque in the same figure
Force on the left, Torque on the right. Don't forget labels and titles.

You can do this one of two ways; note that in the hwk you'll need to do the 2nd option (with some help), but if you want to try it now, feel free.
- Option 1: Copy and paste the code from above (twice) and change the relevant data/lines of code (axs[], names/titles, start index) in the second copy
- Option 2: Encapsulate the plotting part in a function then call that function twice
  - pass in the axes to plot in, along with t values (or create them within the function)
  - do the x,y,z loop inside of the plot

```python
# Reading in the wrist force data for you
wrist_force_data = np.loadtxt("Data/Wrist_force_data.csv", dtype="float", delimiter=",")
```

```python

def plot_function(axs, ts, data, channel_info):
    # See below for possible format for channel_info
    # optional do plot as a function
```

```python
# Make my own channel_info - use this to pass information into the plot about what units/dimensions to use
channel_info_torque = {"name": "Wrist torque", "units":"n.m", "dimensions":3}
channel_info_force = {"name": "Wrist force", "units":"n.m", "dimensions":3}


nrows = ...
ncols = ...
fig, axs = ...

# Call plot function here, once with the wrist force data, once with the wrist torque data
```

```python
grader.check("plot_force_and_torque")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS TAF4 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

- Make sure your plots are visible and there aren't pages of print outs the TAs have to look through to find them
- Save the file
- Submit just this .ipynb file through gradescope, Lab 3, plotting
- You do NOT need to submit the data files - we will supply those
- Where there are given variable/file names (eg, foo = ...) DON'T change those, or the autograder will fail

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Most likely failure for this assignment is not naming the data directory and files correctly; capitalization matters for the Gradescope grader. 

Second most likely cause for failure is reading in any files but ./Data/Wrist_torque_data.csv and Data/Wrist_force_data.csv (make sure when you copied your code over you didn't accidentally include code that reads in other files)

Several of these problems are manually graded by looking at the plots. If we can't see the plots, we can't grade them.

Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```


Tests for plot_force_and_torque (points: 5)
```
>>> assert axs.size == 2
```
```
>>> assert 'rist force' in axs[0].get_title()
```
```
>>> assert 'rist torque' in axs[1].get_title()
```
```
>>> assert len(axs[0].lines) >= 3
```
```
>>> assert len(axs[1].lines) >= 3
```


Tests for plot_torque (points: 5)
```
>>> assert 'Wrist torque' in axs_wf.get_title()
```
```
>>> assert len(axs_wf.lines) == 3
```
