```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab_3_plotting.ipynb")
```

## Lab 3: Plotting data

Read in the right hand accelerometer data for subject 1's first clap and plot it for each of the x, y, and z channels. I recommend doing this in the following order:

- Write plotting code for a single dimension. Make sure it's working and the titles/labels etc are correct 
- Once that's working, add a **for** loop to loop over x, y, z
- If doing a **for** loop is confusing, then start with just copying the code over, once, and changing the dimension you're using in the copy. Then go back and put the for loop in

[optional] Do nicer colors than the default ones.

For setting the colors and the data labels, you can either do it by hand, or setup a list of color names and use that.

Slides: https://docs.google.com/presentation/d/1IiGGUNet-4Nj07x2cTXU6IOYXy9TSdAF5OUWCCKIYEM/edit?usp=sharing


```python
# Libraries that we need to import - numpy and json (for loading the description file)
import numpy as np
import json as json
import matplotlib.pyplot as plt
%matplotlib inline
```

```python
# Numeric ids to indicate hand motion type from Lab 1
clap_id = 1
high_five_id = 2
snap_id = 3
```

```python
# TODO: Copy get_descriptor, get_file_info, get_data, get_data_from_files, get_channel_data from lab 2.
#  Put each function in its own cell

```

```python
# Read in one data file and the data description.
all_data = get_data_from_files([
    get_file_info("Data/S01C01.csv"),
])

try:
    with open("Data/data_description.json", "r") as fp:
        data_description = json.load(fp)
except FileNotFoundError:
    print(f"The file was not found; check that the data directory is in the current one and the file is in that directory")

```

```python
# TODO Get the descriptor and the data
rh_accelerometer_descriptor = ...
rh_accelerometer_data = ...

ts_descriptor = ...
ts = ...
```

```python
nrows = 1
ncols = 1
_, axs_wf = plt.subplots(nrows, ncols, figsize=(6, 4))

# Useful for labeling the channels with x,y,z instead of 0, 1, 2
channel_labels = ['x', 'y', 'z']

# TODO: Graph the right hand accelerometer data.

```

```python
grader.check("plot_one")
```

## Part 2 - plot both the right hand gyroscope and right hand accelerometer in the same figure
Gyroscope on the left, Accelerometer on the right. Don't forget labels and titles.

You can do this one of two ways; note that in the homework you'll need to do the 2nd option (with some help), but if you want to try it now, feel free.
- Option 1: Copy and paste the code from above (twice) and change the relevant data/lines of code (axs[], names/titles, start index) in the second copy
- Option 2: Encapsulate the plotting part in a function then call that function twice
  - pass in the axes to plot in, along with t values (or create them within the function)
  - do the x,y,z loop inside of the plot

```python
# Reading the right hand gyroscope data for you
rh_gyroscope_descriptor = get_descriptor(data_description, "Right hand gyroscope")
rh_gyroscope_data = get_channel_data(all_data, index_offset=rh_gyroscope_descriptor["index_offset"], n_dims=rh_gyroscope_descriptor["dimensions"])
```

```python
# TODO: Fill in the plot function
#  Reminder to use axs to plot, not plt
def plot_function(axs, ts, data, channel_info, your_choice):
    """ Plot all of the dimensions for a given channel
    @param axs - the place to plot in
    @param ts - the time steps to use
    @param data - the numpy array from the csv file
    @param channel_info - a dictionary with the channel info, e.g. "name" and "units"
    @param your_choice - either total number of columns or names for the dimensions
    @returns - nothing"""
    pass
```

```python
# TODO: Call subplots here (Do not change axs_two). One row, two columns
nrows = ...
ncols = ...
fig, axs_two = ...

# TODO Call plot function here, once with the gyroscope data, once with the accelerometer data

fig.tight_layout()
plt.show()
```

```python
grader.check("plot_gyroscope_and_accelerometer")
```

### Plotting all data channels

In the previous step, you made plots for the right hand accelerometer and gyroscope. Now, let's make plots for every data channel (except for "Timestamp").

The first column should be the left hand data channels, and the second column should be the right hand data channels. The first row should be the accelerometers, and the second row should be the gyroscopes.

```python
# TODO: Get the descriptor and the data
lh_accelerometer_descriptor = ...
lh_accelerometer_data = ...

lh_gyroscope_descriptor = ...
lh_gyroscope_data = ...

# TODO: Create the figure (2 rows, 2 columns)
nrows = ...
ncols = ...
_, axs_all = plt.subplots(nrows, ncols, figsize=(10, 8))

# TODO: Call plot_function for each data channel.
fig.tight_layout()
```

```python
grader.check("plot_all")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS I25 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

- Make sure your plots are visible and there aren't pages of print outs the TAs have to look through to find them
- Save the file
- Submit just this .ipynb file through gradescope, Lab 3, plotting
- You do NOT need to submit the data files - we will supply those
- Where there are given variable/file names (eg, foo = ...) DON'T change those, or the autograder will fail

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Most likely failure for this assignment is not naming the data directory and files correctly; capitalization matters for the Gradescope grader. 

Second most likely cause for failure is reading in any files but `data/S01C01.csv` (make sure when you copied your code over you didn't accidentally include code that reads in other files)

Several of these problems are manually graded by looking at the plots. If we can't see the plots, we can't grade them.

Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```


Tests for plot_all (points: 5)
```
>>> assert 'eft' in axs_all[0, 0].get_title() and 'ccel' in axs_all[0, 0].get_title()
```
```
>>> assert 'eft' in axs_all[1, 0].get_title() and 'yroscope' in axs_all[1, 0].get_title()
```
```
>>> assert 'ight' in axs_all[0, 1].get_title() and 'ccel' in axs_all[0, 1].get_title()
```
```
>>> assert 'ight' in axs_all[1, 1].get_title() and 'yroscope' in axs_all[1, 1].get_title()
```
```
>>> assert len(axs_all[0, 0].lines) == 3
```
```
>>> assert len(axs_all[0, 1].lines) == 3
```
```
>>> assert len(axs_all[1, 0].lines) == 3
```
```
>>> assert len(axs_all[1, 1].lines) == 3
```


Tests for plot_gyroscope_and_accelerometer (points: 5)
```
>>> assert axs_two.size == 2
```
```
>>> assert 'ight hand gyroscope' in axs_two[0].get_title()
```
```
>>> assert 'ight hand accelerometer' in axs_two[1].get_title()
```
```
>>> assert len(axs_two[0].lines) >= 3
```
```
>>> assert len(axs_two[1].lines) >= 3
```


Tests for plot_one (points: 5)
```
>>> assert 'ight' in axs_wf.get_title() and 'ccelerometer' in axs_wf.get_title()
```
```
>>> assert len(axs_wf.lines) == 3
```
