```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_3_plotting.ipynb")
```

# Lecture goals

1. Understand the difference between a figure/window and a plotting area/subplot/axes
2. Be able to create a figure with subplots in it
3. Plot something in the subplot and label it
4. Pass parameters to functions by both position and name

Note: To see what the correct result should look like, see slides https://docs.google.com/presentation/d/1IiGGUNet-4Nj07x2cTXU6IOYXy9TSdAF5OUWCCKIYEM/edit?usp=sharing


```python
# Install a pip package in the current Jupyter kernel
# Note: this is a safety check in case your JN install was not in the same place as your libraries
import sys
!{sys.executable} -m pip install matplotlib
```

```python
# Access all numpy functions as np.
import numpy as np

# Everything matplotlib-y will be plt.
import matplotlib.pyplot as plt
%matplotlib inline

```

# Practice plotting data

## Do a circle
Plot a cosine, sine, and then a circle, but also practice creating and putting data in a numpy array and getting it back out again. And yet more practice with functions.

## Add a few markers to the plot

Make the plot a bit fancier with some markers

## Make a shape

Plot a "fun" shape

## Make data

TODO: Make a 3 x n matrix that has t values in the first row, cos(t) values in the second row, and sin(t) values in the third row

```python
# TODO Write a function that takes in how many samples to create, and outputs a 3xn_samples matrix
#  with t-values in the first row, cos(t) in the second row, and sin(t) in the third
#  The t samples should go from 0 to 2 pi
#
# If you're stuck
#   1a) Look at the tutorial for making t and cos(t) values
#   1b) Do it first NOT in a function
#   2a) Make a 3xn_samples array using np.zeros
#   2b) Fill in each row in turn



```

```python
# TODO Now call that function to create the data. 

n_samples = 4

# TODO Call your function here. Make sure to use n_samples and put the data in my_sortofa_circle_data
my_sortofa_circle_data = ...   

```

```python
# TODO: Write a function that creates the three plots
def plot_one_row(axs, data):
    """ Plot cosine, sine, circle. See slides for what this should look like
    @param axs the axes for the row (created by matplotlib)
    @param data - the data you created in the make_circle_data method
    @returns: nothing"""

    ...
```

```python
# TODO 1) Create two sets of data, one with a small number of samples (< 10) and the other with more samples (> 20)
#   Do this by calling your function for making data


# TODO Call your function here; use the names given
my_lowdim_data = ...
my_highdim_data = ...

# TODO: Use subplots to make a 2 x 3 sized set of plots
#   Plot the low dimensional data in the top row, the high dimensional data in the bottom row
#   you must use your plot function for full credit


# TODO Call subplots to make fig and axs, then call your plot function. 
#    axs should have size 2x3 (2 rows, 3 columns)
fig, axs = ...

```

```python
grader.check("plot_circle")
```

## Plot something fun

TODO: Create a plot of a shape. Use a couple colors. Add text in the plot somewhere using text (see https://matplotlib.org/stable/api/_as_gen/matplotlib.axes.Axes.text.html#matplotlib.axes.Axes.text)


```python
fig2, axs2 = plt.subplots(1, 1)

```

```python
grader.check("fun_shape")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS I25 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

- Do a restart then run all to make sure everything runs ok
- Save the file
- Submit just this .ipynb file through gradescope, Lecture activity 3, plotting
- You do NOT need to submit the data files - we will supply those
- Where there are given variable/file names (eg, foo = ...) DON'T change those, or the autograder will fail

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Most likely failure for this assignment is not naming the data directory and files correctly; capitalization matters for the Gradescope grader. 

Second most likely cause for failure is reading in any file but ./Data/Wrist_torque_data.csv 

RUN YOUR FILE AND SAVE SO THE PLOTS SHOW

Tests for fun_shape (points: 1)
```
>>> assert not axs2.get_title() == ''
```
```
>>> assert axs2.lines[0]
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```


Tests for plot_circle (points: 3)
```
>>> assert my_sortofa_circle_data.shape == (3, 4)
```
```
>>> assert my_lowdim_data.shape[0] == 3
>>> assert my_lowdim_data.shape[1] <= 10
```
```
>>> assert my_highdim_data.shape[0] == 3
>>> assert my_highdim_data.shape[1] >= 20
```
```
>>> assert axs.shape == (2, 3)
```
```
>>> for ax in axs:
...     for a in ax:
...         assert not a.get_title() == ''
```
```
>>> assert axs[0, 2].get_aspect() == 1.0
>>> assert axs[1, 2].get_aspect() == 1.0
```
