```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_3_plotting.ipynb")
```

# Lecture goals

1. Understand the difference between a figure/window and a plotting area/subplot/axes
2. Be able to create a figure with subplots in it
3. Plot something in the subplot and label it
4. Pass parameters to functions by both position and name

Note: To see what the correct result should look like, see slides https://docs.google.com/presentation/d/1IiGGUNet-4Nj07x2cTXU6IOYXy9TSdAF5OUWCCKIYEM/edit?usp=sharing


# Install a pip package in the current Jupyter kernel
# Note: this is a safety check in case your JN install was not in the same place as your libraries
import sys
!{sys.executable} -m pip install matplotlib

```python
# Access all numpy functions as np.
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
import json as json
```

## Read data

TODO: Read in the Wrist_torque_data.csv file. This is the data we'll be plotting; note that this data is the result of calling **get_channel_data** for the wrist torque data - it is **n_picks** by **n_timesteps * 3**.

```python
# Load in the Wrist_torque_data.csv file
#  Note: Use the filename "Data/Wrist_torque_data.csv" - do not put in a full path name/a different file name, or the
#   autograder on Gradescope will fail and you'll get a zero.


wrist_torque_data = ...
```

```python
grader.check("Read data")
```

## Get data

TODO: Get the Wrist torque **z** channel for the first row of pick data (these will be the y values you plot with). Create the **t** (time) values (these will be the x values you plot with). Remember that the z values start at index 2 and are every 3rd value

For the **t** values, assume the data is sampled at 30 Hz, i.e., the time sampling is 1/30th of a second

- Step 1: How big does the t array have to be? (hint: How many z values are there for the wrist torque data? What was n_time_steps?)
- Step 2: How do you make an array of that size with that spacing? (hint: **np.arange**)


```python
# TODO: Create a numpy array that starts at 0, ends at number of time steps * time_step, and has step size time_step
#   See np.arange
# Step 1: Calculate the number of time steps in the data (assume a data spacing of 3)
# Step 2: Create a numpy array with that many elements with a spacing of 1.0/30.0 (30 Hz) (this is ts)
# Step 3: Get out just the z values from the wrist torque data (this is z_values) for the first row


# TODO Set up the time (x axis) data 
time_step = ...   # This is 1/30th of a second
n_time_steps = ... # This is the number of time steps/data samples. You should calculate it from wrist_torque_data
# Use the above, with np.arrange, to create the t values
ts = ...
z_values = ...

# Use this to check that your data is the right size - both should be of size (40,)
#  Note: it says (40,) instead of (40,1). That's just numpy's way of saying that it's a 1 dimensional array,
#   or a tuple with 1 value
print(f"ts is size {ts.shape}")
print(f"z_values is size {z_values.shape}")

```

```python
grader.check("Get data")
```

## Plot data

TODO: Plot the **ts** versus **z_values** data. Make sure to add x and y labels, along with a title.

x-axis label: time - make sure to use units (seconds) in the label

y-axis label: what you're plotting - make sure you use the units (N.m)

Title: Should be what data channel this is

```python
# Create the plotting window. Use subplots.
nrows = 1
ncols = 1

fig, axs = ...
```

```python
grader.check("Plot data")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS TAF4 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

- Do a restart then run all to make sure everything runs ok
- Save the file
- Submit just this .ipynb file through gradescope, Lecture activity 3, plotting
- You do NOT need to submit the data files - we will supply those
- Where there are given variable/file names (eg, foo = ...) DON'T change those, or the autograder will fail

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Most likely failure for this assignment is not naming the data directory and files correctly; capitalization matters for the Gradescope grader. 

Second most likely cause for failure is reading in any file but ./Data/Wrist_torque_data.csv 

RUN YOUR FILE AND SAVE SO THE PLOTS SHOW

Tests for Get data (points: 2)
```
>>> assert ts.size == 40
```
```
>>> assert z_values.size == 40
```
```
>>> assert np.isclose(ts[0], 0.0)
```
```
>>> assert np.isclose(ts[-1], 1.3)
```
```
>>> assert np.isclose(z_values[0], 0.12849379, atol=0.001)
```
```
>>> assert np.isclose(z_values[-1], -0.077284, atol=0.001)
```


Tests for Plot data (points: 2)
```
>>> assert ts.shape == z_values.shape
```
```
>>> assert 's' in axs.get_xlabel() or 'S' in axs.get_xlabel()
```
```
>>> assert 'N.m' in axs.get_ylabel()
```
```
>>> assert 'rist' in axs.get_title() and 'orque' in axs.get_title()
```


Tests for Read data (points: 1)
```
>>> assert wrist_torque_data.shape == (660, 40 * 3)
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```
