```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_7_fmin.ipynb")
```

```python
# Usual imports
from scipy.optimize import fmin
import matplotlib.pyplot as plt
import numpy as np
```

# Change the starting location to get the other local minima

Slides: https://docs.google.com/presentation/d/1xKhWyquiP-_wM9AnJJ57BBtF0eitEqSdODq34ttMZgk/edit?usp=sharing

TODO: Set up fmin so that you get the minima at x = 0

```python
# This function has several local minima - how to find the best one?
def my_multiple_minima_func(x):
    """A function with multiple minima
    @param x the input x value
    @returns ax^2 + bx + c"""
    return (3.0 * x - 4.0) * (0.5 * x + 0.2) * (x - 3.0) * (x - 3.0) + 2.0 * x + 0.5
```

```python
# Change this so you get the global minima near x = 0.15
x_at_min1, f_at_x_min1, _, _, _ = fmin(my_multiple_minima_func, 4.0, full_output=True)
print(f"Minimum of f is {f_at_x_min1}, happens at x={x_at_min1[0]}")
```

```python
# Plot the function and the minima
fig, axs = plt.subplots(1, 1, figsize=(4, 4))
ts = np.linspace(-2.0, 5.0)
axs.plot(ts, my_multiple_minima_func(ts), linestyle='-', label="Local minima")
axs.plot(x_at_min1, f_at_x_min1, 'Xr')

axs.set_xlabel('x')
axs.set_ylabel('y')

axs.set_title("Minimum of quartic: Local minima")
```

```python
grader.check("fmin_start_location")
```

# Be less accurate (but faster)

TODO: Change xtol and/or ftol until the total number of iterations is under 10

```python
# TODO: Change xtol and ftol until the number of iterations is under 10
# (Also play around with the values to see how LONG you can make it run...
x_not_as_accurate_min, f_not_as_accurate_min, iters, _, _ = fmin(my_multiple_minima_func, 4.0, xtol=0.0001, ftol=0.0001, full_output=True)
...

```

```python
fig2, axs2 = plt.subplots(1, 1, figsize=(4, 4))
ts = np.linspace(2.6, 2.9)
axs2.plot(ts, my_multiple_minima_func(ts), linestyle='-', label="Zoom in on curve")

x_at_min2, f_at_x_min2, _, _, _ = fmin(my_multiple_minima_func, 4.0, full_output=True)
axs2.plot(x_at_min2, f_at_x_min2, 'Xr', label="Accurate minima")
print(f"Minimum of f is {f_at_x_min2}, happens at x={x_at_min2[0]}")

axs2.plot(x_not_as_accurate_min, f_not_as_accurate_min, 'Xg', label="Not Accurate minima")

axs2.set_xlabel('x')
axs2.set_ylabel('y')
axs2.legend()

axs2.set_title("Minimum of quartic: Zoom in on Local minima")
```

```python
grader.check("fmin_iter")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS TAF4 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

All the usual - rerun, remove print statements, save. Make sure plots are visible.

- Submit this .ipynb file to lecture activity 7 (fmin)

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Hope your term is going well!

Tests for fmin_iter (points: 2)
```
>>> assert iters < 10
```
```
>>> assert np.isclose(f_at_x_min2, f_not_as_accurate_min, atol=0.1)
```


Tests for fmin_start_location (points: 1)
```
>>> assert np.isclose(x_at_min1, 0.1545625, atol=0.0001)
```
```
>>> assert np.isclose(f_at_x_min1, -7.1299378, atol=0.01)
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```
