```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab_8_calculus.ipynb")
```

# Area under an elliptic paraboloid

Use the monte carlo technique in the tutorial to estimate the area under the 2D elliptic parabaloid

Your code should work even if the center and scale values shift (provided the parabaloid is still above the z = 0 plane)

Slides: https://docs.google.com/presentation/d/1sq3ERLBET1ourJZJzXMtJe6-cN8Tlt-GLEzSDOxrd0M/edit?usp=sharing

```python
import numpy as np
import matplotlib.pyplot as plt
from functools import partial
```

# Create the function

TODO: Create a function with the center/scale parameters bound to the given values.

You might find it useful to go back to **a_tutorial_fmin** to remember how to bind parameters to create a simple function to use using **partial**.

Also get the plotting code from that file and add that to the plot question below; you can then add plotting the samples as you work through this lab.

```python
# A fancier quadratic-y type function (but in 2D)
def elliptic_paraboloid(x, y, x0=0, y0=0, sclx=1, scly=1):
    """ x,y -> f(x) in the shape of a 'bowl'
    @param x - x value in the plane
    @param y - y value in the plane
    @param x0 - amount to shift the bottom of the bowl by in x
    @param y0 - amount to shift the bottom of the bowl by in y
    @param sclx - scale the bowl's side in x
    @param scly - scale the bowl's side in y
    @returns - f(x,y) the bowl's height over the point x,y
    """
    return (x-x0)**2 / sclx**2 + (y-y0)**2 / scly**2
```

```python
# Paramaters to bind
sclx = 1.5
scly = 0.75
x0 = 1.0
y0 = 2.0

# TODO: Use partial to create a function with the parameters bound to the above values
my_func = ...
# END PROMPT"""
```

```python
print(f"Func at x0, y0 is {my_func(x0, y0)}, expected 0.0")
print(f"Func at 2.0, 2.0 is {my_func(2.0, 2.0)}, expected 0.4444")
```

```python
grader.check("create_func")
```

# Generating samples (see tutorial)

TODO: Generate samples in x, y, and z inside the "box" that contains the function
- samples need to be in the box
- Find the max z value using fminbounds
- Generate samples using np.random.uniform (you can use something fancier if you want)
- Calculate the area of the box
- You need to generate x, y, and z samples - 3D points in the box
- I suggest generating the points, drawing them, then do the area calculation (next question)

Use x0 +-2 and y- +-2 as the integration bounds in the x,y directions.

```python
# Make it easy to change the number of samples.
#. Reminder: You need the same number of samples in all directions (x,y, and z)
n_samples = 1000

# Doing the meshgrid for you - notice the +- 2.0 from x0 and y0, which is the bounds we'll use for this problem
# These are NOT the random samples you want - this is a grid of samples to use for creating the surface
# ... you can also use it to find the min/max values of the function
xs_grid, ys_grid = np.meshgrid(np.linspace(x0 - 2.0, x0 + 2.0), np.linspace(y0 - 2.0, y0 + 2.0))

# TODO: Generate samples in x, y, and z inside the "box" using np.random.uniform


xs_samples = ...
ys_samples = ...
zs_samples = ...
area_box = ...
# END PROMPT"""

print(f"Box area is {area_box}, expected 4 x 4 x height = 142.2222")
```

```python
grader.check("samples")
```

# Points above/below

TODO: Are the samples above or below the surface? (see tutorial)

Note that you can assume the function is above the z=0 plane, so you only need to check if the sample is below **my_func(x,y)**

Note: The number points that are below should be around 30%, but it can vary quite a lot.

You might also want to create a boolean array so that you can re-use it in the print:

**b_under = samples < function**

```python
area_under = ...
# END PROMPT"""
print(f"Area is {area_under}, expected between 43 and 44")
```

```python
grader.check("area")
```

# Plot the function and the points above and below

TODO Grab the plotting code from the fmin tutorial then add in plotting the samples, then separate them by whether or not they're under the surface

```python
...
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, week 8 lab (calculus)

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for area (points: 4)
```
>>> assert(np.isclose(area_under, 43, atol=1.1))
```


Tests for create_func (points: 2)
```
>>> assert(np.isclose(my_func(x0, y0), 0.0, atol=0.001))
```
```
>>> assert(np.isclose(my_func(2.0, 2.0), 0.4444, atol=0.001))
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for samples (points: 3)
```
>>> assert(len(xs_samples) == n_samples and len(ys_samples) == n_samples and len(zs_samples) == n_samples)
```
```
>>> assert(np.isclose(np.max(zs_samples), np.max(my_func(xs_grid, ys_grid)), atol=0.1))
```
```
>>> assert(np.all(np.abs(xs_samples - x0) <= 2.0))
```
```
>>> assert(np.all(np.abs(ys_samples - y0) <= 2.0))
```
```
>>> assert(np.isclose(area_box, 4 * 4 * np.max(my_func(xs_grid, ys_grid)), atol=0.1))
```
