```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_8_calculus.ipynb")
```

# Lecture activity: calculus 

Practice calculus with this function:

f(x) = (3/x) ln(x) - x^2 / 10.0 - x / 3.0 + (x-3) * x * x

# Create the function

Write the Python function for the above equation

```python
# Do the imports
```

```python
# Use my_func for the name of the Python function
```

```python
print(f"Func at 1.0 is {my_func(1.0)}, expected -2.433333")
print(f"Func at 2.0 is {my_func(2.0)}, expected -4.02694589")
```

```python
grader.check("create_func")
```

# Plot the function, along with the bounds and an x at minima and maxima

Use 0.5 and 3.0 as the integration bounds

```python
# Integration limits
a = 0.5
b = 3.0
```

```python
# Find the x locations of the two bends
#
first_max = ...
first_min = ...
# END PROMPT """

print(f"Max {first_max}, expected 0.94813, Min {first_min}, expected 2.091114")
```

```python
...
```

```python
grader.check("plot_fn")
```

# Derivative
TODO Numerically alculate the derivatives at the first_max/first_min points (what should the derivative be?)

Pick a reasonable **h** for the derivative calculation. You can do forward, backward, or central

```python
deriv_max = ...
deriv_min = ...
# END PROMPT"""
```

```python
grader.check("derivative")
```

# Integration

TODO Calculate the area under the curve using trapz or quad, whichever you perfer.

```python
...

area_under_curve = ...
# END PROMPT"""
```

```python
grader.check("area")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, Lecture activity 8, Calculus

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for area (points: 1)
```
>>> assert(np.isclose(area_under_curve[0], -7.90504, atol=0.001))
```


Tests for create_func (points: 1)
```
>>> assert(np.isclose(my_func(1.0), -2.433333, atol=0.001))
```
```
>>> assert(np.isclose(my_func(2.0), -4.02694589, atol=0.001))
```


Tests for derivative (points: 1)
```
>>> assert(np.isclose(deriv_max, 0.0, atol=0.0001))
```
```
>>> assert(np.isclose(deriv_min, 0.0, atol=0.0001))
```
```
>>> assert(deriv_min is not 0.0)
```
```
>>> assert(deriv_max is not 0.0)
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for plot_fn (points: 1)
```
>>> assert(np.isclose(first_max, 0.94813, atol=0.1))
```
```
>>> assert(np.isclose(first_min, 2.091114, atol=0.1))
```
