```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_8_calculus.ipynb")
```

# Lecture activity: calculus 

Slides: https://docs.google.com/presentation/d/1sq3ERLBET1ourJZJzXMtJe6-cN8Tlt-GLEzSDOxrd0M/edit?usp=sharing

Practice calculus with this function:

f(x) = (3/x) ln(x) - x^2 / 10.0 - x / 3.0 + (x-3) * x * x

# Create the function

TODO: Write the above equation as a function named "my_func"

```python
# Do the imports - you'll need fminbound from scipy.iptimize and quad from scipy.integrate in addition to the usual numpy/matplotlib
```

```python
# Use my_func for the name of the Python function
```

```python
# Check my_func is returning the correct answer for a couple x values
print(f"Func at 1.0 is {my_func(1.0)}, expected -2.433333")
print(f"Func at 2.0 is {my_func(2.0)}, expected -4.02694589")
```

```python
grader.check("create_func")
```

# Minima/maxima and Plot

Find the minima/maxima within the bounds **a=0.5** and **b=3.0**

Plot the function, along with the minima/maxima

TODO: Use a new function, **fminbound**, that is essentially **fmin** but you can put bounds on the **x** values to limit the search to a part of the function.


```python
# Integration limits and the boundaries of the search for minima and maxima
a = 0.5
b = 3.0
```

```python
# TODO Find the x locations of the two bends (minima and maxima)
# Use fminbound for both of these
#  Reminder that the maximum can be found by negating
#  You might find a lambda function useful
#  Use the ability to set bounds to keep the search to just the "bump" around the minima or maxima (you can hardwire these bounds)
#     I'd suggest plotting the function and/or looking at the plot in the slides to find reasonable bounds
first_max = ...
first_min = ...

print(f"Max {first_max}, expected 0.94813, Min {first_min}, expected 2.091114")
```

```python
# TODO Plot the function from a to b
# TODO Put a green x at the first maxima, and a red x at the first minima
```

```python
grader.check("plot_fn")
```

# Derivative
TODO Numerically alculate the derivatives at the first_max/first_min points (what should the derivative be?)

Pick a reasonable **h** for the derivative calculation. You can do forward, backward, or central

```python
# TODO: Numerically calculate the derivative at the minima/maxima
deriv_max = ...
deriv_min = ...
```

```python
grader.check("derivative")
```

# Integration

TODO Calculate the area under the curve using **trapz** or **quad**, whichever you perfer.

```python
...

area_under_curve = ...
```

```python
grader.check("area")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS I25 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

All the usual - rerun, remove print statements, save.

- Submit this .ipynb file to lecture activity 8 (calculus)

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

This submit rarely fails. 

Tests for area (points: 1)
```
>>> assert np.isclose(area_under_curve[0], -7.90504, atol=0.001)
```


Tests for create_func (points: 1)
```
>>> assert np.isclose(my_func(1.0), -2.433333, atol=0.001)
```
```
>>> assert np.isclose(my_func(2.0), -4.02694589, atol=0.001)
```


Tests for derivative (points: 1)
```
>>> assert np.isclose(deriv_max, 0.0, atol=0.0001)
```
```
>>> assert np.isclose(deriv_min, 0.0, atol=0.0001)
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```


Tests for plot_fn (points: 1)
```
>>> assert np.isclose(first_max, 0.94813, atol=0.1)
```
```
>>> assert np.isclose(first_min, 2.091114, atol=0.1)
```
