```python
# Initialize Otter
import otter
grader = otter.Notebook("HW_4_pinball.ipynb")
```

# HWK 4 pinball simulation

The full pinball simulation: slides https://docs.google.com/presentation/d/1fNTOc_4YXfBkQ_ko9X8eQqAgK3Opnn9c02i72sVDqKU/edit?usp=sharing

Full pinball 
 - position and velocity with Euler simulation (lecture activity, compute_next_step)
 - stop when ball passes below y=0 (lab, simulate_pinball version 1)
 - "Bouncing" off of the top and side walls (lab, outside_xx_wall, handle bounce, simulate_pinball version 1)
 - "Bouncing" off of bumpers (homework)

Week 1: we'll add two things for week 1:
- Circular "bumpers"
- Some bumpers will impart an acceleration to the pinball


```python
# The usual imports
import numpy as np
import matplotlib.pyplot as plt
```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

<!-- BEGIN QUESTION -->

# Circular bumper

TODO: Copy your lab simulation code, set up your imports, and add a circular bumper
- decide which functions/what the format of the functions should be to handle a bumper
- write test code for that function
- add handling bumpers to the simulation


```python
# Imports from week 5
# This bit of code tells Jupyter notebooks to also look in the week_6_optimization folder for files
import os
import sys

# If you want them...
sys.path.insert(0, os.path.abspath('../Week_5_matrices'))
```

```python
# Pinball_routines imports
...
```

```python
# Simulate pinball routine
#. TODO Decide how to set up your parameters to support an arbitrary number of bumpers in addition to the walls
...
```

```python
# Time step
delta_t = 0.01

# Test 1 - do you stop when you hit the floor?
starting_state = np.zeros([3, 2])  # location, velocity, acceleration
starting_state[0, :] = [0, 0] # Start at zero, zero
# Velocity - mostly up with a bit of x noise
# Check top wall only
# starting_state[1, :] = [-0.2, 10.0]
# Check all walls
# starting_state[1, :] = [5.2, 5.0]
# Check top left bumper
starting_state[1, :] = [2.4, 8.5]
starting_state[2, :] = [0.0, acceleration_due_to_gravity()]

# Define a dictinary that holds the walls and bumpers
# Default is top wall at y=5, left/right walls at +- 3.0
# Bumpers of radius around 1 at each corner
...

```

```python
# TODO: your tests
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

# Plot bumpers and walls and path

TODO: Start with the plotting code from the lab
 - add plotting the bumpers
 

```python
# TODO: Plot the track, walls, and bumpers
#  Start with the lab plot
...

```

```python
# TODO Call the plot function
...
```

<!-- END QUESTION -->

# Use integration to calculate length of path

```python

```

```python
# TODO: Add in checks for outside_left_wall and outside_right_wall
# Add in check for poses not going outside of left/right walls

```

```python
grader.check("integrate_path")
```

<!-- BEGIN QUESTION -->

# Plot

I've provided the plot routine; you can run this at any point to check your current ret_poses as you're debugging

```python
# Plot the walls and the pinball path
def plot_pinball_lab(ret_poses, walls, total_time):
    """ plot the results of running the system AND the "correct" closed form result
    @param ret_poses - x y position values in a 2xn numpy array
    @param walls - The walls and ceiling locations (top, left, right)
    @param total_time - the total time the system ran (for closed form solution, delta_t * n time steps)
    @return Nothing
    """
    nrows = 1
    ncols = 1
    fig, axs = plt.subplots(nrows, ncols, figsize=(4, 4))

    # The values we calculated in calculate_n_time_steps
    axs.plot([walls[1], walls[2]], [walls[0], walls[0]], '-m', label=f"Top wall {top_wall}")
    axs.plot([walls[1], walls[2]], [0, 0], '-k', label="Floor wall")
    axs.plot([walls[1], walls[1]], [0, walls[0]], '-g', label="Left wall")
    axs.plot([walls[2], walls[2]], [0, walls[0]], '-g', label="Right wall")

    axs.plot(ret_poses[0, 0], ret_poses[1, 0], 'xr', label="Start")
    axs.plot(ret_poses[0, :], ret_poses[1, :], '.-k', label="Poses")

    axs.axis('equal')
    axs.set_title(f"Boring pinball, 0-{total_time} s")
    axs.legend()
```

```python
# Pass the walls into the plot as a list
total_time = delta_t * ret_poses.shape[1]
plot_pinball_lab(ret_poses, [top_wall, left_wall, right_wall], total_time)
```

```python
grader.check("plotting")
```

<!-- END QUESTION -->

# Question 5 instructions


```python
# code question 3
```

```python
grader.check("Question 3")
```

# Question 6 instructions


```python
# code question 3
```

```python
grader.check("Question 3")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, homework 4 pinball. Don't forget pinball_routines.py

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for Question 3 (points: 2)
```
>>> # Code test 3
```
```
>>> # Code test 3
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for integrate_path (points: 4)
```
>>> assert(outside_left_wall([left_wall + 0.5, 0.0], x_wall=left_wall) == False)
```
```
>>> assert(outside_left_wall([left_wall - 0.5, 0.0], x_wall=left_wall) == True)
```
```
>>> assert(outside_right_wall([right_wall - 0.5, 0.0], x_wall=right_wall) == False)
```
```
>>> assert(outside_right_wall([right_wall + 0.5, 0.0], x_wall=right_wall) == True)
```
```
>>> assert(np.all(ret_poses[0, :] < right_wall + 0.01))
```
```
>>> assert(np.all(ret_poses[0, :] > left_wall - 0.01))
```


Tests for plotting (points: 2)
```
>>> # Code test 3
```
