```python
# Initialize Otter
import otter
grader = otter.Notebook("HW_4_pinball.ipynb")
```

# HWK 4 pinball simulation

The full pinball simulation: slides https://docs.google.com/presentation/d/1fNTOc_4YXfBkQ_ko9X8eQqAgK3Opnn9c02i72sVDqKU/edit?usp=sharing

Full pinball 
 - position and velocity with Euler simulation (lecture activity, compute_next_step)
 - stop when ball passes below y=0 (lab, simulate_pinball version 1)
 - "Bouncing" off of the top and side walls (lab, outside_xx_wall, handle bounce, simulate_pinball version 1)
 - "Bouncing" off of bumpers (homework)

Week 1: we'll add two things for week 1:
- Circular "bumpers"
- [optional] Some bumpers will impart an acceleration to the pinball
- [optional] Add damping

Week 2: Some simple calculations
- Length of path pinball travels
- Integral of vx and vy


```python
# The usual imports
import numpy as np
import matplotlib.pyplot as plt
```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

<!-- BEGIN QUESTION -->

## Circular bumper: How to implement?

TODO: Write code to handle a circular bumper

- decide which functions/what the format of the functions should be to handle a bumper

Note: You'll add tests and simulation and plotting in the following questions

```python
# Imports from week 5
# This bit of code tells Jupyter notebooks to also look in the week_6_optimization folder for files
import os
import sys

# If you want them...
sys.path.insert(0, os.path.abspath('../Week_5_matrices'))
sys.path.insert(0, os.path.abspath('../Week_7_simulation'))
```

```python
# TODO: Write circular bumper routines
# You can write them here or put them in pinball_routines.py and import them here
# Feel free to add more cells
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Circular bumper test routines

TODO: Create tests to check your bumper routines. 

```python
# Feel free to add more cells
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Simulation routine

TODO: Copy over the simulation routine from your lab and add circular bumpers
- How will you represent/store the bumpers?
- If you implement acceleration or drag then describe it here 

```python
# Pinball_routines imports
...
```

```python
# Simulate pinball routine
#. TODO Decide how to set up your parameters to support an arbitrary number of bumpers in addition to the walls
...
```

```python
# Time step
delta_t = 0.01

# Test 1 - do you stop when you hit the floor?
starting_state = np.zeros([3, 2])  # location, velocity, acceleration
starting_state[0, :] = [0, 0] # Start at zero, zero
# Velocity - mostly up with a bit of x noise
# Check top wall only
# starting_state[1, :] = [-0.2, 10.0]
# Check all walls
# starting_state[1, :] = [5.2, 5.0]
# Check top left bumper
starting_state[1, :] = [2.4, 8.5]
starting_state[2, :] = [0.0, acceleration_due_to_gravity()]

# Define a dictinary that holds the walls and bumpers
# Default is top wall at y=5, left/right walls at +- 3.0
# Bumpers of radius around 1 at each corner
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Plot bumpers and walls and path

TODO: Start with the plotting code from the lab
 - add plotting the bumpers
 - pass the data structure you use to define the bumpers and walls to the plot routine
 
TODO: Generate at least 3 plots showing the pinball doing different behaviors (must include reflection off of wall, bumper[s])

```python
# TODO: Create a function to plot the pinball track, walls, and bumpers
#  Start with the lab plot
...
```

```python
# TODO: set up different starting positions and/or pinball setups and plot them
delta_t = 0.01

starting_state = np.zeros([3, 2])  # location, velocity, acceleration
starting_state[0, :] = [0, 0] # Start at zero, zero
# Velocity - mostly up with a bit of x noise
starting_state[1, :] = [2.4, 8.5]
starting_state[2, :] = [0.0, acceleration_due_to_gravity()]

...
```

<!-- END QUESTION -->

## Week 2 - Use integration to calculate length of path

TODO: Calculate the length of the pinball path using euclidian distances

```python
def length_of_pinball_path(poses):
    """ Calculate the length of the pinball path using euclidian distances
    @param poses - the x,y poses
    @return length"""
    # TODO: Sum up the distance between all of the pose points
    # Reminder: Distance is sqrt(x^2 + y^2)
    return 0.0
    # END PROMPT """
```

```python
# TODO For these tests use this starting state
#  - do NOT place a bumper/wall in the path - it should just go up and down
#  - Use the given starting position and velocity
#  - Use a delta t of 0.01
#  - Do NOT do damping
starting_state = np.zeros([3, 2])  # location, velocity, acceleration
starting_state[0, :] = [0, 0] # Start at zero, zero
# Velocity - mostly up with a bit of x noise
starting_state[1, :] = [0.5, 3.5]
starting_state[2, :] = [0.0, acceleration_due_to_gravity()]

# TODO: call your simulate routine with  delta_t = 0.01 and the above starting state

ret_poses, ret_velocities = simulate_pinball(starting_state, ..., delta_t=delta_t)
# END PROMPT """

```

```python
len_path = (length_of_pinball_path(ret_poses))
print(f"Distance traveled {len_path}, expected 1.386298")
```

```python
grader.check("integrate_path")
```

## Integrate x velocity and y velocity

TODO: Calculate the total relative distance the pinball traveled in x and y by integrating the velocity in x and y
- Change your simulation code so it returns velocities in addition to poses
- Call trapz with the vx values, using a fixed dx (delta_t) (this is the integral)

```python
def integral_of_path_velocity(velocities, delta_t):
    """ Calculate the x and y distance traveled of the pinball path by integrating the velocities
    @param velocities - the vx,vy velocities at each time step
    @param delta_t - the sampling used to generate the velocities
    @return integral of vx, integral of vy"""
    # TODO: Calculate the integral of vx and vy
    # Note: You'll want trapz because these are samples of the velocity function
    return integral_x, integral_y
    # END PROMPT """
```

```python
# Note - make sure you're using the starting state given in the path integration
integral_x, integral_y = integral_of_path_velocity(ret_velocities, delta_t)
print(f"Integral of x velocity {integral_x}, expected 0.365")
print(f"Integral of y velocity {integral_y}, expected -0.056209")
```

```python
grader.check("integrate_velocity")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, homework 4 pinball. Don't forget pinball_routines.py

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for integrate_path (points: 4)
```
>>> assert(np.isclose(len_path, 1.386298, atol=0.001))
```


Tests for integrate_velocity (points: 4)
```
>>> assert(np.isclose(integral_x, 0.365, atol=0.001))
```
```
>>> assert(np.isclose(integral_y, -0.056209, atol=0.001))
```
