```python
# Initialize Otter
import otter
grader = otter.Notebook("HW_2_system_analysis.ipynb")
```

# HW 2: Analysis of a mass-damped system

Resources: Lecture slides describing the homework: https://docs.google.com/presentation/d/1MwF34w30rjzLS5-IYj0ZWT5CjsLeFc7qlG_Eq9rFdZ8/edit?usp=sharing All of the equations for this homework are in those slides.

- Part 1/week 3 does not build on anything specific in lab 3, but uses code from labs 1-3
- Part 2/week 4: Do lab 4 before hwk 4

In the first part of this homework the focus is on doing an analysis "from scratch", in the sense that we will not be
providing a lot of shell code. The second half is on how to translate a general equation of the form

  *state at time t+delta_t = f( state at time t )*
  
into code (also known as simulation).

Programming practice: The focus in this assignment is on deciding when to create a function and what parameters to pass in and out of the function.

Following on from lab 4, the second part of the assignment will be an example of writing an iterative function - one where the intention is to call the function multiple times, each time passing in the values returned from the previous function call. This can be a bit difficult to wrap your head around, so if you get stuck go back to the simpler examples and/or do it on paper yourself a few iterations.

We've provided results (as a json file) for two of the included data sets. We will, however, be testing on different data sets, so make sure your code works for any data file of the correct format.

TODOS: For each helper function
- Write the function
- Check that it works

You have some flexibility as to how the write the helper functions, but you must use the names/input/output format given or the autograder won't work...

Helper function is not a technical term - it's just what I call a function that does some part of the functionality of the entire task.

## Helper function 1 - load the data and check that it's valid
 - Input: File name (with path)
 - Output: Two numpy arrays, floats
   -- 1st numpy array: The time values
   -- 2nd numpy array: The function values that correspond to the time values
   
You MUST use this format and the file name given, or the autograder will fail.

Lists should be of the same length

Some expectations on the data files
 - Always has at least three rows
 --First row is a header row
- The first column is the time values (will be monotonically increasing)
- The second column is the values that correspond to the time values (final value of this will be greater than initial)

A good habit is to check that these are true, and throw an error if it's not

```python

# You must name your function this
# You might also want to write some test code...
def load_data_from_file(fname):
    pass

```

```python
grader.check("load_data")
```

## Helper function 2 - get the index of the first number greater than or equal to the input number x

Look at the equations in the slides - why is this a useful helper function?

- Input (2): A numpy array, a single number x
- Output: The index of the first number greater than or equal to x (or None if no such number)

```python
# You must name your function this
def greater_than_index(in_list, in_x):
    pass
# Tests
# list [1.0, 3.0, 4.0, 7.0, 10.0], x 6.0, should give index 3
# list [-2.5, 1.0, 4.0, 8.0, 4.0, 1.0, -2.5], x 4.0, should give index 2
# list [1.1, 2.2, 3.3, 4.4], x 5.5, should return None
#

```

```python
grader.check("get_index")
```

## Analysis code

TODO: Write a function that takes in a file name and outputs a dictionary with the values listed below
 - You must use the names given (see json file/example in homework slide)
 - See homework slides for further definitions/equations

Check the corresponding json file and plot for the correct values 

### System values (Figure 4.14 in slides)
 - c_initial - the initial position of the system (the first value)
 - c_max - the largest position of the system (the maximum value, see Figure 4.14)
 - c_final - the final, steady-state position of the system (the last value, see Figure 4.14)

Note: For this assignment we'll just use the last value, but you could also average the last few values

###  Estimate characteristics
 - rise_time - The rise time, Tr this is the time that it takes to go from 10% to 90% of the way from c_initial to c_final (item 1 on Figure 4.14 - refer to plot)
 -- The “10% time” is defined as the first time at which the system obtains a value greater than or equal to the
value which is 10% between c_initial and c_final. Same with the 90% time.
 - peak_time - The peak time, Tp. This is the time at which the position has the maximal value (item 2 on Figure 4.14)
 - perc_overshoot - percentage over shoot (% OS). This is the amount that the system overshoots c_final, expressed as a percentage of the range c_initial to c_final.
 - settling_time - The settling time, Ts. Estimate the 2% Settling Time, T_s.  This is the earliest time when the current and all subsequent positions of the system are within a certain threshold of c_final. This threshold is
defined by 2% of the range between c_initial and c_final.

For example, if c_initial = -1 and c_final = 1, the 2% threshold would be (0.96, 1.04) (non-inclusive of endpoints).

###    Estimate model values
 - system_mass - assume the mass is 1 (this will not be true in part II)
 - system_spring - this is omega_n^2 (see slides)
 - system_damping - the damping term (the linear coefficient, see slides)

You may write some additional helper functions here, if you wish. Actually, I recommend writing this in pieces/multiple functions. One breakdown is the one given above - calculate the values in turn.

I'd suggest writing a test function for each of the group of equations above, where you test against the answers in the data1.json file (with the data1.csv file as input). 

I have written a general-purpose test function for you, but it tests everything all at once.

```python

# Open the filename and create a dictionary with all of the parameters. Return the dictionary
def analyze_data(fname):
    pass

```

```python
def compare_dictionaries(fname_dictionary, my_dict):
    """ Open up the json file in fname and compare to my_dict
    @param fname_dictionary one of dataX.json
    @param my_dict - your dictionary with parameters
    @returns True or False if the same (within epsilon)"""
    from json import load
    with open(fname_dictionary, 'r') as fp:
        answ_dict = load(fp)

    b_ret = True
    for k, v in answ_dict.items():
        try:
            if not np.isclose(v, my_dict[k], atol=0.1):
                print(f"Key {k} is not close, correct value {v}, incorrect {my_dict[k]}")
                b_ret = False
        except KeyError:
            print(f"Key {k} not found in your dictionary")
            b_ret = False
    return b_ret

```

```python
grader.check("analysis")
```

<!-- BEGIN QUESTION -->

## Plot the data and the system parameters

TODO Plot the data with the parameters, similar to the figure 4.14

Refer to figure in slides for what the figure should look like

You should at least have **c final**, **Ts**, and **Tp** plotted on the figure

Optional: Plot all of the data on the left hand side, then plot a "clipped" version of the data (clip to where it stops vibrating) with the parameters on them.

Create one plot each for each data file.

- matplot lib has a text function for placing text

When deciding what functions to write: I've left four blocks here. The first block should be function(s) that do the plotting/reading/analysis, etc. Blocks 2-4 should just call those functions with the different data files.

```python
# An example of creating a "helper" function that does a plotting task 
def draw_corner(axs, x_values, y_values, ls, col):
    axs.plot([x_values[0], x_values[1]], [y_values[1], y_values[1]], linestyle=ls, color=col)
    axs.plot([x_values[1], x_values[1]], [y_values[0], y_values[1]], linestyle=ls, color=col)


# more functions here

```

```python
# Read, analyze, and plot data 1
```

```python
# Read, analyze, and plot data 1
```

```python
# Read, analyze, and plot data 1
```

<!-- END QUESTION -->

# Part 2 - ode simulation

In this part of the assignment you'll simulate the spring-mass system yourself. We'll follow the second half of Lab 4 - using ode to do the simulation (you just have to write the derivative function and the set-up code).

This should be nearly identical to the predator_prey_derivative ode example (from a coding syntax standpoint). The equations and the meanings of the variables will be different, however

- Like the c_tutorial_iterative_systems example, you will need to calculate two values, but this time dx/dt and dx^2/dt^2 
- Like the predator_prey_derivative example, you will
-- create a derivative function that takes in a 2 dimensional current state (x, dx/dt), t (not used), and params (c, k, and m) and returns dx/dt and dx^2/d^2t
-- params will be a dictionary that you create in the calling code
-- create a set of time values to evaluate the system at
-- use ode to do the actual forward simulation

Functions you'll need

- a derivative function (like predator_prey_derivative)
- a function/cell to create the initial data to simulate and write out/save the results
- a function that calls your code from part 1 to analyze and plot the data (do NOT write new plotting/analysis code). Two options:
-- Write the data out to a file then read it back in using the code you just wrote in the previous part
-- If you split up your plot/analysis functions the right way, you can just pass the data directly to the plot/analysis functions

Parameters:

- c: damping term (system_damping in json file)
- k: spring term (system_spring in json file)
- m: mass (don't forget to put this value in the analysis dictionary)

Equation:

-  dx/dt = dx/dt
-  d^2x/d^2t = (-c * dx/dt - k * x) / m


```python
# Functions to run the ode and then call the analysis code
```

```python
# Run the ode with the following parameters and store the answer in dict_answ
# Make sure you run the simulation long enough for it to stabilize
# c = 1.1, m = 2.0, k = 50.0
dict_answ = ...
```

```python
grader.check("ode_functions")
```

<!-- BEGIN QUESTION -->

# Part b: Recreate data 1
Using your simulation code, re-create the plot for data 1

- Get the parameters from your initial analysis
- Pick a reasonable stop time
- Plot the result

```python
# Run your simulation, analyze it, and plot the results
# This SHOULD just be calling existing functions
```

<!-- END QUESTION -->

# Part c - increasing mass

Increase the mass from 1 to 10 (but use the same c and k parameters from the previous problem).

Plot the system. Has it stabilized? Adjust the time to run so that it stabilizes. Then plot the result.

Set dict_bigger_analysis to be the analysis with mass 

```python

# Run and plot the simulation with the mass increased to 10.
dict_bigger_mass_analysis = ...
```

```python
# Check your analysis (set dict_bigger_mass_analysis)
assert(compare_dictionaries('Data/sim_and_plot_answer_c.json', dict_bigger_mass_analysis))
```

```python
grader.check("bigger_mass")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit just the .ipynb/zip file to Gradescope, HWK2 System Analysis.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for analysis (points: 7)
```
>>> assert(compare_dictionaries('Data/data1.json', analyze_data('Data/data1.csv')))
```
```
>>> assert(compare_dictionaries('Data/data2.json', analyze_data('Data/data2.csv')))
```
```
>>> assert(compare_dictionaries('Data/data3.json', analyze_data('Data/data3.csv')))
```


Tests for bigger_mass (points: 2)
```
>>> assert(np.isclose(dict_bigger_mass_analysis["settling_time"], 74.89, atol=0.1))
```


Tests for get_index (points: 3)
```
>>> assert(greater_than_index(np.array([1.0, 3.0, 4.0, 7.0, 10.0]), 6.0) == 3)
```
```
>>> assert(greater_than_index(np.array([-2.5, 1.0, 4.0, 8.0, 4.0, 1.0, -2.5]), 4.0) == 2)
```
```
>>> assert(greater_than_index(np.array([1.1, 2.2, 3.3, 4.4, 5.5]), 100.5) == None)
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for load_data (points: 2)
```
>>> ts, vs = load_data_from_file("Data/data1.csv")
>>> assert(len(ts) == len(vs))
```
```
>>> assert(len(ts) == 20000)
```
```
>>> assert(vs[-1] > vs[0])
```


Tests for ode_functions (points: 6)
```
>>> assert(compare_dictionaries('Data/sim_and_plot_answer_a.json', dict_answ))
```
