```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab_3_data_fitting.ipynb")
```

# Lab week 3: Data fitting to see if we can detect a pick success y/n

Can you use the finger 1 stop time and max value to tell the difference between successful and not?
Lab slides: https://docs.google.com/presentation/d/1NtoXYbl2nq1dkIU0KQE8ogjSCfb22wgLigY3JJpHHGI/edit?usp=sharing


```python
# Doing the imports for you
import numpy as np
import json as json
from scipy.stats import linregress
import matplotlib.pyplot as plt
import csv
from numpy.polynomial import polynomial
import math
```

## Read the data in 

```python
# Read the same data in as for pre_lecture_3.py data in and put it in the data_ variables

# Read in week3_Motor position f1_success and failed
data_successful = ...  # week3_Motor position f1_successful.csv
data_failed = ...  # week3_Motor position f1_failed.csv
# END PROMPT"""

```

```python
grader.check("Read data")
```

## Box plot of fitted lines for all data

In the pre-lecture you fitted the line to just one row each of the successful/failed picks. In this lab you'll fit the line to ALL of the picks. For each pick, save the end point (time and y_max value of the motor). You'll be making a box plot of this data to see if it looks different for successful versus failed picks.

Why a box plot? Box plots are better for visually showing multiple distributions in the same plot, where they can be compared side-by-side

See lab slides for what this should look like when you're done

_Type your answer here, replacing this text._

```python
# from pre_lecture_3_fitting import fit_line_to_middle_bit

...

# nx2 arrays of t, y values of end point of fitted line
end_times_successful = ...  
end_times_failed = ...  
# END PROMPT"""

# Should print Mean successful [3.98791579e-01 4.07402569e+02] and failed [  0.40810473 392.76044707]
print(f"Mean successful {np.mean(end_times_successful, 0)} and failed {np.mean(end_times_failed, 0)}")

```

```python
grader.check("Fitted line code")
```

_Type your answer here, replacing this text._

```python
from pre_lecture_3_fitting import fit_line_to_middle_bit

n_rows = 1
n_cols = 2
fig, axs = plt.subplots(n_rows, n_cols)

...

# Do the box plots...
# END PROMPT"""
```

<!-- END QUESTION -->

# Part 2: Fitting a curve to the summed wrist force data

## First step: calculate the data

Data to plot: sqrt (sum x,y,z (wrist force data)^2)

Read in wrist force data, sum it, and (optional) write it out to two csv files (one for successful, one for failed)
  Feel free to swipe code from Lecture_3_data_analysis.ipynb to do this

The data files are also in the directory, so you can read them in if you're having trouble calculating the data yourself.


```python

data_wrist_force_successful = ...
data_wrist_force_failed = ,,,
# END PROMPT"""

```

```python
grader.check("Fit curve")
```

<!-- BEGIN QUESTION -->

## Second step: Fit a cubic to the first rows of the successful and failed plots

I've copied over the code from Lecture 3 - you just need to fit the polynomial and plot that, too

plt.close(fig)

```python
n_rows = 6
n_cols = 6
fig, axs = plt.subplots(n_rows, n_rows, figsize=(10, 10))
for p in range(0, (n_rows // 2) * n_cols):
    r, c = p // n_cols, p % n_cols
    axs[r, c].plot(ts, data_wrist_force_successful[p])
    ...
    axs[r, c].set_title(f"{ch_name} summed" + f" suc {p}", fontsize=10)

for p in range(0, (n_rows // 2) * n_cols):
    r, c = 3 + p // n_cols, p % n_cols
    axs[r, c].plot(ts, data_failed[p])
    ...
    axs[r, c].set_title(f"{ch_name} summed" + f" fail {p}", fontsize=10)

plt.tight_layout()
```

<!-- END QUESTION -->

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, week 3 lab

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for Fit curve (points: 2)
```
>>> assert(data_wrist_force_successful.shape[0] == 355)
```
```
>>> assert(data_wrist_force_failed.shape[0] == 305)
```
```
>>> assert(math.isclose(np.max(data_wrist_force_successful), 20.46056888454263, rel_tol=0.001))
```
```
>>> assert(math.isclose(np.max(data_wrist_force_failed), 18.862565694678903, rel_tol=0.001))
```


Tests for Fitted line code (points: 2)
```
>>> assert(math.isclose(np.mean(end_times_successful, 0)[0], 3.98791579e-01, rel_tol=0.001))
```
```
>>> assert(math.isclose(np.mean(end_times_successful, 0)[1], 4.07402569e+02, rel_tol=0.001))
```
```
>>> assert(math.isclose(np.mean(end_times_failed, 0)[0], 0.40810473, rel_tol=0.001))
```
```
>>> assert(math.isclose(np.mean(end_times_failed, 0)[1], 392.76044707, rel_tol=0.001))
```


Tests for Read data (points: 1)
```
>>> assert(data_successful.shape == (355, 40) and data_failed.shape == (305, 40))
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```
