```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab_3_data_fitting.ipynb")
```

# Lab week 3: Data fitting to see if we can detect a pick success y/n

Can you use the finger 1 stop time and max value to tell the difference between a successful and a failed pick?
Lab slides: https://docs.google.com/presentation/d/1NtoXYbl2nq1dkIU0KQE8ogjSCfb22wgLigY3JJpHHGI/edit?usp=sharing


```python
# Doing the imports for you
import numpy as np
import json as json
from scipy.stats import linregress
import matplotlib.pyplot as plt
import csv
from numpy.polynomial import polynomial
import math
```

## Read the data in 

```python
# Read the same data in as for pre_lecture_3.py data in and put it in the data_ variables

# Read in week3_Motor position f1_success and failed
data_successful = ...  # week3_Motor position f1_successful.csv
data_failed = ...  # week3_Motor position f1_failed.csv
# END PROMPT"""

```

```python
grader.check("Read data")
```

## Box plot of fitted lines for all data

In the pre-lecture you fit a line to just one row (each) of the successful versus failed picks. In this lab you'll fit the line to ALL of the picks. For each pick, save the end point (time at the y_max value and the y_max value). You'll be making a box plot of this data to see if it looks different for successful versus failed picks.

Why a box plot? Box plots are better for visually showing multiple distributions in the same plot, where they can be compared side-by-side

See lab slides for what this should look like when you're done.

```python
...

# TODO: get the fit_line_to_middle_bit function, either by copying it over here, or importing it from a .py file
#. Optional (import):
# from pre_lecture_3_fitting import fit_line_to_middle_bit
# Otherwise, copy function here
#
# TODO: Check that it works
#. - Create t values
#  - Call fit_line_to_middle_bit with the first row of the successful data and see if your answers
#. match the answers from pre-lecture
# END PROMPT"""

```

```python

# TODO:
#  Create a function that computes the 
# nx2 arrays of t, y values of end point of fitted line
end_times_successful = ...  
end_times_failed = ...  
# END PROMPT"""

# Should print Mean successful [3.98791579e-01 4.07402569e+02] and failed [  0.40810473 392.76044707]
print(f"Mean successful {np.mean(end_times_successful, 0)} and failed {np.mean(end_times_failed, 0)}")

```

```python

# TODO:
#  Create a function that loops over all of the rows of the data and accumulates the start and end times
#. in an nx2 array of t (first row) and y (second row)
# I'd recommend a function that takes both the t values and the data, so you don't have to keep recomputing
#.  the t values
#. def fit_middle_all(ts, data)
#.    make numpy array to put data in (nx2)
#.    for each row in data
#.       start_pt, end_pt = fit_line_to_middle_bit(ts, current row)
#        copy end_pt information into appropriate row in nx2 output
#. return nx2 numpy array
# END PROMPT"""
```

```python
# Check that it works properly

# TODO:
#  Create the t values
#. Call your function twice, once with the successful data, once with the failed
end_times_successful = ...
end_times_failed = ...
# END PROMPT"""
```

```python
grader.check("Fitted line code")
```

```python
# Two plots
n_rows = 1
n_cols = 2
fig, axs = plt.subplots(n_rows, n_cols, figsize=(8, 4))

...

# Do the box plots...
# END PROMPT"""
```

<!-- END QUESTION -->

# Part 2: Fitting a curve to the wrist force data

## First step: Getting the data


The wrist force data has three channels (x,y,z). For this assignment, we want to get out the total summed force at the wrist, rather than each channel individually:

Equation: **sqrt (sum x,y,z (wrist force data)^2)**

Just like the last problem, the data has already been extracted from the proxy_pick_data.csv file and output for you (week3_Wrist force_summed_*.csv). 

Optional: Do this extraction yourself. Read in wrist force data, sum it, and write it out to two csv files (one for successful, one for failed). Feel free to swipe code from Lecture_3_data_analysis.ipynb to do this.

```python

# TODO: Read in the data from Data/week3_Wrist_force summed*.csv (successful and failed)
# TODO: Optional: Read in the pick data and its description file, extract out the rows and columns you
#. care about, and write that data out (should be same as the .csv files)
data_wrist_force_successful = ...
data_wrist_force_failed = ...
# END PROMPT"""

```

```python
grader.check("Fit curve")
```

<!-- BEGIN QUESTION -->

## Second step: Fit a cubic to the wrist force data (18 rows each successful/failed)

Split by successful and failed.

TODO: See slides for what the final plots should look like. I've handled the for loop and the plot indexing for you.

At this point, you should have the data in the form ts, ys for both the successful and the failed data. 

Where you're headed (pseudo code)

for 18 rows of data

    Fit a cubic polynomial to the data
    
    Plot the original data plus the fitted polynomial
    
       First three rows: successful, last three rows: failed

Implementation steps: You could write a function for doing the polynomial fit, but since there already is one (polynomial.polyval) you probably don't need one.

Steps
- Fit polynomial for one row of data
- Plot it and make sure it looks correct
- Copy that code into the for loop, adjusting for which row (p) and which subplot (axs[r, c])
- Repeat for the failed data

Note: It can take a few seconds to plot.

```python
n_rows = 6
n_cols = 6
fig, axs = plt.subplots(n_rows, n_rows, figsize=(16, 16))
ch_name = "Wrist force"
for p in range(0, (n_rows // 2) * n_cols):
    r, c = p // n_cols, p % n_cols
    axs[r, c].plot(ts, data_wrist_force_successful[p])
    ...
    axs[r, c].set_title(f"{ch_name} summed" + f" suc {p}", fontsize=10)

for p in range(0, (n_rows // 2) * n_cols):
    r, c = 3 + p // n_cols, p % n_cols
    axs[r, c].plot(ts, data_wrist_force_failed[p])
    ...
    axs[r, c].set_title(f"{ch_name} summed" + f" fail {p}", fontsize=10)

plt.tight_layout()
```

<!-- END QUESTION -->

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, lab 3 data analysis. Turn in just the .ipynb file - we will supply the data files.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for Fit curve (points: 2)
```
>>> assert(data_wrist_force_successful.shape[0] == 355)
```
```
>>> assert(data_wrist_force_failed.shape[0] == 305)
```
```
>>> assert(math.isclose(np.max(data_wrist_force_successful), 20.46056888454263, rel_tol=0.001))
```
```
>>> assert(math.isclose(np.max(data_wrist_force_failed), 18.862565694678903, rel_tol=0.001))
```


Tests for Fitted line code (points: 2)
```
>>> assert(math.isclose(np.mean(end_times_successful, 0)[0], 3.98791579e-01, rel_tol=0.001))
```
```
>>> assert(math.isclose(np.mean(end_times_successful, 0)[1], 4.07402569e+02, rel_tol=0.001))
```
```
>>> assert(math.isclose(np.mean(end_times_failed, 0)[0], 0.40810473, rel_tol=0.001))
```
```
>>> assert(math.isclose(np.mean(end_times_failed, 0)[1], 392.76044707, rel_tol=0.001))
```


Tests for Read data (points: 1)
```
>>> assert(data_successful.shape == (355, 40) and data_failed.shape == (305, 40))
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```
