```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab_2_plotting.ipynb")
```

## Lab 2: Plotting data

Read in the wrist force data and plot it for each of the x, y, and z channels. I recommend doing this in the following order:

- Copy in your pre_lecture code and do that plot first - make sure it's working and the titles/labels etc are correct 
- Once that's working, add a for loop to loop over x, y, z
- If doing a for loop is confusing, then start with just copying the code over, once, and changing the dimension you're using in the copy. Then go back and put the for loop in

[optional] Do nicer colors than the default ones.

For setting the colors and the data labels, you can either do it by hand, or setup a list of color names and use that

```python
# Libraries that we need to import - numpy and json (for loading the description file)
import numpy as np
import json as json
import matplotlib.pyplot as plt
%matplotlib inline
import otter
grader = otter.Notebook()
```

_Type your answer here, replacing this text._

```python
pick_data = ...
```

```python
...
```

```python
nrows = 1
ncols = 1
fig, axs = plt.subplots(nrows, ncols, figsize=(6, 4))

...
```

```python
grader.check("plot_force")
```

```python

nrows = ...
ncols = ...
fig, axs = ...
```

```python
grader.check("plot_force_and_torque")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit just the .ipynb file to Gradescope. Don't change the provided variable names or autograding will fail.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for plot_force (points: 3)
```
>>> axs.get_title() == "Wrist torque"
True
```
```
>>> len(axs.lines) == 3
True
```


Tests for plot_force_and_torque (points: 2)
```
>>> axs.size == 2
True
```
```
>>> axs[0].get_title() == 'Wrist force'
True
```
```
>>> axs[1].get_title() == 'Wrist torque'
True
```
```
>>> len(axs[0].lines) == 3
True
```
```
>>> len(axs[1].lines) == 3
True
```
