```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_2_plotting.ipynb")
```

# Lecture goals

1. Understand the difference between a figure/window and a plotting area/subplot/axes
2. Be able to create a figure with subplots in it
3. Plot something in the subplot and label it
4. Pass parameters to functions by both position and name

Note: To see what the correct result should look like, see Lab lecture https://docs.google.com/presentation/d/1IiGGUNet-4Nj07x2cTXU6IOYXy9TSdAF5OUWCCKIYEM/edit?usp=sharing

```python
# Access all numpy functions as np.
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
import json as json
```

## Read data

TODO: Read in the data from lab 1/hwk 1. Note that this folder also has a Data folder in it with the same data files, so you should be able to just copy your code from week 1

```python
# Note: your repository is set up to have a copy of the Data files in this folder as well as the Week_1_arrays
#. folder - this is so the autograder works correctly.  
pick_data = np.loadtxt("Data/proxy_pick_data.csv", dtype="float", delimiter=",")

# Load in the proxy_pick_data_description.json file

pick_data_description = ...
```

```python
grader.check("Read data")
```

## Get data

TODO: Get the Wrist **torque** z channel for the first row of pick data (y values). Create the t values (x values). Refer to lab 1 for how get out the data (it's the same data).

For the **t** values, assume the data is sampled at 30 Hz, i.e., the time sampling is 1/30th of a second

- Step 1: How big does the t array have to be? (hint: How many data samples are there for the wrist force data?)
- Step 2: How do you make an array of that size with that spacing? (hint: **np.arange**)


```python
# TODO: Create a numpy array that starts at 0, ends at number of time steps * time_step, and has step size time_step
#   See np.arange
#.  See lab 1 for how to calculate the number of time steps

# TODO Set up the time (x axis) data 
time_step = ...   # This is 1/30th of a second
n_time_steps = ... # This is the number of time steps/data samples
# Use the above, with np.arrange, to create the t values
ts = ...


# TODO: Use slicing to get the y data out of pick_data
y_data = ... 

# Use this to check that your data is the right size - both should be of size 40
print(ts.size)
print(y_data.size)

```

```python
grader.check("Get data")
```

## Plot data

TODO: Plot the ts and ys values. Make sure to add x and y labels, along with a title.

```python
# Create the plotting window. Use subplots.
nrows = 1
ncols = 1

fig, axs = ...
```

```python
grader.check("Plot data")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, Pre-lecture 2 plotting. Submit just the .ipynb file

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for Get data (points: 2)
```
>>> assert(ts.size == 40)
```
```
>>> assert(y_data.size == 40)
```
```
>>> assert(np.isclose(ts[0], 0.0))
```
```
>>> assert(np.isclose(ts[-1], 1.3))
```
```
>>> assert(np.isclose(y_data[0], -0.26818306))
```
```
>>> assert(np.isclose(y_data[-1], -0.43396691))
```


Tests for Plot data (points: 2)
```
>>> assert(ts.shape == y_data.shape)
```
```
>>> assert("s" in axs.get_xlabel())
```
```
>>> assert("N.m" in axs.get_ylabel())
```
```
>>> assert("Wrist torque" in axs.get_title())
```


Tests for Read data (points: 1)
```
>>> assert(pick_data.size == 871860)
```
```
>>> assert(len(pick_data_description["Data channels"]) == 17)
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```
