```python
# Initialize Otter
import otter
grader = otter.Notebook("pre_lecture_2_plotting.ipynb")
```

# Lecture goals

1. Understand the difference between a figure/window and a plotting area/subplot/axes
2. Be able to create a figure with subplots in it
3. Plot something in the subplot and label it
4. Pass parameters to functions by both position and name

Note: To see what the correct result should look like, see Lab lecture https://docs.google.com/presentation/d/1IiGGUNet-4Nj07x2cTXU6IOYXy9TSdAF5OUWCCKIYEM/edit?usp=sharing

```python
# Access all numpy functions as np.
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
import json as json
import otter
grader = otter.Notebook()
```

## Read data

Read in the data from lab 1/hwk 1

```python
# Note: your repository is set up to have a copy of the Data files in this folder as well as the Week_1_arrays
#. folder - this is so the autograder works correctly.  
pick_data = np.loadtxt("Data/proxy_pick_data.csv", dtype="float", delimiter=",")

# Now load in the description data - remember to change the path name like the line above

pick_data_description = ...
```

```python
grader.check("Read data")
```

## Plot data

Plot the Wrist **torque** z channel for the first row of pick data. Look through the following "steps" before you start. (The torque channel is the 2nd one in the data channels list.)

For the **t** values, assume the data is sampled at 30 Hz, i.e., the time sampling is 1/30th of a second

- Step 1: How big does the t array have to be? (hint: How many data samples are there for the wrist force data?)
- Step 2: How do you make an array of that size with that spacing? (hint: np.arange)

Create the data to plot first, then do all the plot commands.

_Type your answer here, replacing this text._

```python
# Create a numpy array that starts at 0, ends at number of time steps * time_step, and has step size time_step
#   See np.arange

ts = ...


y_data = ... 
```

```python
# Use this to check that your data is the right size
print(ts.size)
print(y_data.size)
```

```python
# Create the plotting window. Use subplots.
nrows = 1
ncols = 1

fig, axs = ...
```

```python
grader.check("Plot data")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, week 2 pre-lecture

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for Plot data (points: 2)
```
>>> ts.shape == y_data.shape
True
```


Tests for Read data (points: 1)
```
>>> pick_data.size == 871860
True
```
```
>>> len(pick_data_description["Data channels"]) == 17
True
```
