```python
# Initialize Otter
import otter
grader = otter.Notebook("HW_2_system_analysis.ipynb")
```

# HW 2: Analysis of a mass-damped system

Resources: Lecture slides describing the homework: https://docs.google.com/presentation/d/1MwF34w30rjzLS5-IYj0ZWT5CjsLeFc7qlG_Eq9rFdZ8/edit?usp=sharing All of the equations for this homework are in those slides.

- Part 1/week 4 does not build on anything specific in lab 4, but uses functions and numpy
- Part 2/week 5: Do lab 5 before homework 

In the first part of this homework the focus is on doing an analysis "from scratch", in the sense that we will not be
providing a lot of shell code. The second half is on how to translate a general equation of the form

  *derivative of state at time t = f( state at time t )*
  
into code (also known as simulation).

Programming practice: The focus in this assignment is on deciding when to create a function and what parameters to pass in and out of the function.

Following on from lab 5, the second part of the assignment will be an example of writing an iterative function - one where the intention is to call the function multiple times, each time passing in the values returned from the previous function call. This can be a bit difficult to wrap your head around, so if you get stuck go back to the simpler examples and/or do it on paper yourself a few iterations.

We've provided results (as a json file) for three of the included data sets. We will, however, be testing on different data sets, so make sure your code works for any data file of the correct format.

TODOS: For each helper function
- Write the function
- Check that it works

You have some flexibility as to how the write the helper functions, but you must use the names/input/output format given or the autograder won't work...

Helper function is not a technical term - it's just what I call a function that does some part of the functionality of the entire task.

## Helper function 1 - load the data and check that it's valid
 - Input: File name (with path)
 - Output: Two numpy arrays
   -- 1st numpy array: The time values
   -- 2nd numpy array: The function values that correspond to the time values
   
You MUST use this format and the file name given, or the autograder will fail.

The two arrays should be of the same length

Some expectations on the data files
 - Always has at least three rows
 --First row is a header row
- The first column is the time values (will be monotonically increasing)
- The second column is the values that correspond to the time values (final value of this will be greater than initial)

A good habit is to check that these are true, and throw an error if it's not

```python
# TODO: Put your imports here
```

```python
# TODO: Put your function definition for loading the data here

# You must name your function this with the filename passed in
#  It should return ts, vs
def load_data_from_file(fname):
    t_valuess = ...
    y_values = ...
    return t_values, y_values

```

```python
# TODO: Call your function here

# Make sure you assign the t and y values to the variable names ts, vs
ts, vs = ...
```

```python
grader.check("load_data")
```

## Helper function 2 - get the index of the local minima or local maxima

Look at the equations in the slides - why is this a useful helper function?

- Input (2): A numpy array, boolean (True: find local maxima, False: find local minima)
- Output: The index of every local maxima or local minima in the input array

Use the **argrelmin** and **argrelmax** functions to find the indices of the minimum and maximum values in a numpy array.
These functions are in scipy.signal module. They essentially bundle the find max, and where code we've been writing
- https://docs.scipy.org/doc/scipy/reference/generated/scipy.signal.argrelmin.html
- https://docs.scipy.org/doc/scipy/reference/generated/scipy.signal.argrelmax.html

```python


# TODO: Add importing argrelmin and argrelmax to your imports. They are in scipy.signal
 
# TODO: Write this function.
# This function takes in data and whether or not to search for the peak or the valley. It returns an index or None if no peak/valley
# If b_peak_or_valley is True, return the peak
def find_peak_or_valley(data, b_peak_or_valley):
    pass


```

```python
# Test your function with datas.
data_a = np.array([1.0, 3.0, 4.0, 7.0, 6.0, 8.0, 9.0, 10.0, 5.0])
data_b = np.array([-2.5, 1.0, 4.0, 8.0, 4.0, 1.0, -2.5])
data_c = np.array([1.1, 2.2, 3.3, 4.4])

# TODO Add tests - what should the correct values be? If you are unsure, try plotting the data
```

```python
grader.check("get_index")
```

## Analysis code

TODO: Write a function that takes in a file name and outputs a dictionary with the values listed below
 - You must use the names given (see json file/example in homework slide)
 - See homework slides for further definitions/equations

Check the corresponding json file and plot for the correct values 

### System values (Slide 8)
 - c_initial - the initial position of the system (the first value)
 - c_peak - the largest position of the system (or the first peak position, see Figure 4.14)
 - c_final - the final, steady-state position of the system (the last value, see Figure 4.14)

Note: For this assignment we'll just use the last value, but you could also average the last few values

###  Estimate peak and valley times and values
 - c_second_peak - the second peak position of the system (local maximum value)
 - c_valley - the first valley position of the system (local minimum value)
 - t_peak - the peak time of the system
 - t_second_peak - the second peak time of the system
 - t_valley - the first valley time of the system
 - peak_period - the period between two peaks, (t_second_peak - t_peak) or (t_valley - t_peak)*2
 - perc_overshoot - percentage over shoot (% OS). This is the amount that the system overshoots c_final, expressed as a percentage of the range c_initial to c_final.

###    Estimate model values
 - system_mass - assume the mass is 1 (this will not be true in part II)
 - system_spring - this is omega_n^2 * system_mass (see slides)
 - system_damping - the damping coeifficient (the linear coefficient, see slides)

You may write some additional helper functions here, if you wish. Actually, I recommend writing this in pieces/multiple functions. One breakdown is the one given above - calculate the values in turn.

I'd suggest writing a test function for each of the group of equations above, where you test against the answers in the data1.json file (with the data1.csv file as input). 

I have written a general-purpose test function for you, but it tests everything all at once.

```python
# Cell for the first function you write
```

```python
# Test cell for that function
```

```python
# Another cell for another function
```

```python
# Another cell for testing the above function...
```

```python
# Another cell for another function...
```

```python


# TODO: Based on the given ts and vs, create a dictionary with all of the parameters. Return the dictionary.
# If no mass m is given, assume it is 1
def analyze_data_from_values(ts, vs, m=1.0):
    pass    

```

```python


# TODO: Open the filename and call the analyze_data_from_values function. Return the dictionary.
# If no mass is given, assume it is 1
def analyze_data(fname, m=1.0):
    pass

```

```python
# TEST CELL

# Call the analyze_data function(s) with one of the data files 

```

```python
def compare_dictionaries(fname_dictionary, check_dict):
    """ Open up the json file in fname and compare to check_dict
    @param fname_dictionary one of dataX.json
    @param check_dict - your dictionary with parameters
    @returns True or False if the same (within epsilon)"""
    from json import load
    with open(fname_dictionary, 'r') as fp:
        answ_dict = load(fp)

    b_ret = True
    for k, v in answ_dict.items():
        try:
            if not np.isclose(v, check_dict[k], atol=0.1):
                print(f"Key {k} is not close, correct value {v}, incorrect {check_dict[k]}")
                b_ret = False
        except KeyError:
            print(f"Key {k} not found in your dictionary")
            b_ret = False
    return b_ret

```

```python
grader.check("analysis")
```

<!-- BEGIN QUESTION -->

## Plot the data and the system parameters

TODO Plot the data with the parameters, similar to the figures in the slides

You should have **the first and second peaks**, **the first valley**, and **peak period** plotted on the figure

Optional: Plot all of the data on the left hand side, then plot a "clipped" version of the data (clip to where it stops vibrating) with the parameters on them.

Create one plot each for each data file.

- matplotlib has a text function for placing text

When deciding what functions to write: I've left four blocks here. The first block should be function(s) that do the plotting/reading/analysis, etc. Blocks 2-4 should just call those functions with the different data files.

```python
# An example of creating a "helper" function that does a plotting task 
def draw_corner(axs, x_values, y_values, ls, col):
    axs.plot([x_values[0], x_values[1]], [y_values[1], y_values[1]], linestyle=ls, color=col)
    axs.plot([x_values[1], x_values[1]], [y_values[0], y_values[1]], linestyle=ls, color=col)

```

```python
# Cell for additional plot functions

```

```python
    
# TODO: 
#  Write a function that reads in the data, analyzes it, then plots it
def plot_system(axs, fname, m=1.0):
    pass

```

```python
nrows = 1
ncols = 2
fig, axs = plt.subplots(nrows, ncols, figsize=(12, 3))

# Read, analyze, and plot data 1

# Make it fit better
fig.tight_layout()
```

```python
nrows = 1
ncols = 2
fig, axs = plt.subplots(nrows, ncols, figsize=(12, 3))

# Read, analyze, and plot data 2

# Make it fit better
fig.tight_layout()
```

```python
nrows = 1
ncols = 2
fig, axs = plt.subplots(nrows, ncols, figsize=(12, 3))

# Read, analyze, and plot data 3

# Make it fit better
fig.tight_layout()
```

```python
# Manual grade
print("This is a manually-graded question; there is no grader.check() function. See rubric and slides for more information on expected output.")
```

<!-- END QUESTION -->

# Part 2 - ode simulation

In this part of the assignment you'll simulate the spring-mass system yourself. We'll follow the second half of Lab 5 - using ode to do the simulation (you just have to write the derivative function and the set-up code).

This should be nearly identical to the predator_prey_derivative ode example (from a coding syntax standpoint). The equations and the meanings of the variables will be different, however.

- Like the c_tutorial_iterative_systems example, you will need to calculate two values, but this time dx/dt and dx^2/dt^2 
- Like the predator_prey_derivative example, you will
-- create a derivative function that takes in a 2 dimensional current state (x, dx/dt), t (not used), and params (c, k, and m) and returns dx/dt and dx^2/d^2t
-- params will be a dictionary that you create in the calling code
-- create a set of time values to evaluate the system at
-- use ode to do the actual forward simulation

Functions you'll need

- a derivative function (like predator_prey_derivative)
- a function/cell to create the initial data to simulate and write out/save the results
- a function that calls your code from part 1 to analyze and plot the data (do NOT write new plotting/analysis code). Two options:
-- Write the data out to a file then read it back in using the code you just wrote in the previous part
-- If you split up your plot/analysis functions the right way, you can just pass the data directly to the plot/analysis functions

Parameters:

- c: damping term (system_damping in json file)
- k: spring term (system_spring in json file)
- m: mass (don't forget to put this value in the analysis dictionary)

Equation:

-  dx/dt = dx/dt
-  d^2x/d^2t = (-c * dx/dt - k * x) / m


```python
# TODO: You'll need to import integrate
```

```python
# TODO: write a deriative function to pass to odeint.
#  Note that the input state will have x and the derivative of x in it
#   You will return the derivative of x and the derivative of the derivative of x
#   Note that the derivative never changes - you only have to calculate the second derivative

```

```python

```

```python
# TEST CELL
# Write one call to the derivative function to test it with known values
```

```python
# TODO write a function that takes in the starting conditions (x and the derivative of x), along with the parameters
#  you need to calculate the derivative (zeta, k, and m)
# Call integrate and return the result of that function, along with the t values
#   You might also want the step for the t and the total amount of time to integrate over as parameters
# This function should return 
```

```python
# Run the ode with the following parameters and store the answer in dict_answ
# Make sure you run the simulation long enough for it to stabilize
# x0 = -0.5, c = 1.1, m = 2.0, k = 50.0
dict_answ = ...
```

```python
grader.check("ode_functions")
```

<!-- BEGIN QUESTION -->

# Part b: Recreate data 1
Using your simulation code, re-create the plot for data 1

- Get the parameters from your initial analysis
- Pick a reasonable stop time
- Plot the result

```python
nrows = 1
ncols = 2
fig, axs = plt.subplots(nrows, ncols, figsize=(8, 4))

# Run your simulation, analyze it, and plot the results
# This SHOULD just be calling existing functions
# x0 = -1.0, c = 1.0, m = 1.0, k = 100.0
dict_answ = ...

# Make it fit better
fig.tight_layout()

```

```python
# Manual grade
print("This is a manually-graded question; there is no grader.check() function. See rubric and slides for more information on expected output.")
```

<!-- END QUESTION -->

# Part c - increasing mass

Increase the mass from 1 to 10 (but use the same c and k parameters from the previous problem).

Plot the system. Has it stabilized? Adjust the time to run so that it stabilizes. Then plot the result.

Set dict_bigger_analysis to be the analysis with mass 

```python
nrows = 1
ncols = 2
fig, axs = plt.subplots(nrows, ncols, figsize=(8, 4))


# Run and plot the simulation with the mass increased to 10.
# Other parameters are the same as in the part b.
dict_bigger_mass_analysis = ...

# Make it fit better
fig.tight_layout()
```

```python
# Check your analysis (set dict_bigger_mass_analysis)
assert(compare_dictionaries('Data/sim_and_plot_answer_c.json', dict_bigger_mass_analysis))
```

```python
grader.check("bigger_mass")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS TCW3 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit just the .ipynb file to Gradescope, HWK2 Spring Mass damp.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for analysis (points: 9)
```
>>> assert(compare_dictionaries('Data/data1.json', analyze_data('Data/data1.csv')))
```
```
>>> assert(compare_dictionaries('Data/data2.json', analyze_data('Data/data2.csv')))
```
```
>>> assert(compare_dictionaries('Data/data3.json', analyze_data('Data/data3.csv')))
```


Tests for bigger_mass (points: 2)
```
>>> assert(np.isclose(dict_bigger_mass_analysis["system_spring"], 100.11740135673081, atol=0.2))
```


Tests for get_index (points: 3)
```
>>> assert(all(find_peak_or_valley(data_a, b_peak_or_valley=True) == [3, 7]))
```
```
>>> assert(all(find_peak_or_valley(data_a, b_peak_or_valley=False) == 4))
```
```
>>> assert(all(find_peak_or_valley(data_b, b_peak_or_valley=True) == 3))
```
```
>>> assert(len(find_peak_or_valley(data_b, b_peak_or_valley= False)) == 0)
```
```
>>> assert(len(find_peak_or_valley(data_c, True)) == 0)
```
```
>>> assert(len(find_peak_or_valley(data_c, False)) == 0)
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for load_data (points: 2)
```
>>> ts, vs = load_data_from_file("Data/data1.csv")
>>> assert(len(ts) == len(vs))
```
```
>>> assert(len(ts) == 20000)
```
```
>>> assert(vs[-1] > vs[0])
```


Tests for ode_functions (points: 6)
```
>>> assert(compare_dictionaries('Data/sim_and_plot_answer_a.json', dict_answ))
```
