```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab_4_data_fitting.ipynb")
```

# Lab week 4: Data fitting to see if we can detect a pick success y/n

Can you use the finger 1 stop time and max value to tell the difference between a successful and a failed pick?
Lab slides: https://docs.google.com/presentation/d/1NtoXYbl2nq1dkIU0KQE8ogjSCfb22wgLigY3JJpHHGI/edit?usp=sharing


```python
# Install a pip package in the current Jupyter kernel
# Note: this is a safety check in case your JN install was not in the same place as your libraries
import sys
!{sys.executable} -m pip install scipy
```

```python
# Doing the imports for you
import numpy as np
import json as json
from scipy.stats import linregress
import matplotlib.pyplot as plt
from numpy.polynomial import polynomial
```

## Read the data in 

```python
# Read the same data in as for pre_lecture_3.py data in and put it in the data_ variables

# Read in week3_Motor position f1_success and failed
data_successful = ...  # Motor position f1_successful.csv
data_failed = ...  # Motor position f1_failed.csv
# END PROMPT"""

```

```python
grader.check("Read data")
```

## Box plot of fitted lines for all data

In the lecture activity you fit a line to just one row (each) of the successful versus failed picks. In this lab you'll fit the line to ALL of the picks. For each pick, save the end point (time at the y_max value and the y_max value). You'll be making a box plot of this data to see if it looks different for successful versus failed picks.

Why a box plot? Box plots are better for visually showing multiple distributions in the same plot, where they can be compared side-by-side

See lab slides for what this should look like when you're done.

```python

# TODO: Copy the fit_line_to_middle_bit function here
# END PROMPT"""

```

```python
# TEST CELL
# TODO: Check that it works
#. - Create t values
#  - Call fit_line_to_middle_bit with the first row of the successful data and see if your answers
#. match the answers from lecture activity

...
```

```python
# EXAMPLE CODE
# This bit of code shows how to do the nx2 array part

# Make the nx2 array - one row for each element
np_array_for_res = np.zeros((data_successful.shape[0], 2))
# For each row
for row in range(0, data_successful.shape[0]):
    # Put the min value in the first column, the max in the second
    np_array_for_res[row, 0] = np.min(data_successful[row, :])    
    np_array_for_res[row, 1] = np.max(data_successful[row, :])

    # This also works
    min_value = np.min(data_successful[row, :]) 
    max_value = np.max(data_successful[row, :])
    np_array_for_res[row, :] = [min_value, max_value]

print(np_array_for_res)
```

```python
def fit_middle_all(ts, data):
    """ Fit the middle for every row of data
    @param ts - the t values to use
    @param data - n picks X timesteps
    @returns all end times, max values as one numpy array (n picks X 2)"""


    # TODO:
    #.    make numpy array to put data in (nx2)
    #.    for each row in data
    #.       start_pt, end_pt = fit_line_to_middle_bit(ts, current row)
    #        copy end_pt information into appropriate row in nx2 output
    #. return nx2 numpy array
    ...
# END PROMPT """
```

```python
# TEST CELL
# Write the code to call the function here


# TODO:
#  Create the t values
#. Call your function twice, once with the successful data, once with the failed
end_times_successful = ...
end_times_failed = ...
# END PROMPT"""

# Note for both of these: Small changes in how you index can result in slightly different numbers,
#  but the overall magnitude should be correct

# Note: Doing np.mean(nx2 array, 0) tells numpy to do the mean over each column (instead of the mean
#  over the entire nx2 array). So it should return 2 numbers - the mean of the times and the mean
#  of the y intercepts

# Should print Mean successful [3.98791579e-01 4.07402569e+02] and failed [  0.40810473 392.76044707]
print(f"Mean successful {np.mean(end_times_successful, 0)} and failed {np.mean(end_times_failed, 0)}")

# Should print First successful [4.09512247e-01 4.29880005e+02] and first failed [4.10117354e-01 4.29791992e+02]
print(f"First successful {end_times_successful[0]} and first failed {end_times_failed[0]}")

```

```python
grader.check("Fitted_line_code")
```

<!-- BEGIN QUESTION -->

## Do the box plots

- Left window: t values of y_max point (the first column of end_times_*)

- Right window: y_max values (the sceond column of end_times_*)

Note that boxplot uses labels for labeling the x axis of the box plot.

Implementation order: for plotting, try calling box plot with just one column of data. Then add the second. 

```python
# Two plots
n_rows = 1
n_cols = 2
fig, axs = plt.subplots(n_rows, n_cols, figsize=(8, 4))

...

# Do the box plots...
```

```python
# Manual grade
print("This is a manually-graded question; there is no grader.check() function. See rubric and slides for more information on expected output.")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

Looking at the plots, do you think there is a difference between the successful versus failed data? Which would you use, time or max y value? Write your answer in the next cell.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Part 2: Fitting a curve to the wrist force data

## First step: Getting the data


The wrist force data has three channels (x,y,z). For this assignment, we want to get out the total summed force at the wrist, rather than each channel individually:

Equation: **sqrt (sum x,y,z (wrist force data)^2)**

Just like the last problem, the data has already been extracted from the proxy_pick_data.csv file and output for you (Wrist force_summed_*.csv). 

Optional: Do this extraction yourself. Read in wrist force data, sum it, and write it out to two csv files (one for successful, one for failed). Your homework code should be easy to change to do this.

```python

# TODO: Read in the data from Data/Wrist_force summed*.csv (successful and failed)
# TODO: Optional: Read in the pick data and its description file, extract out the rows and columns you
#. care about, and write that data out (should be same as the .csv files)
data_wrist_force_successful = ...
data_wrist_force_failed = ...
# END PROMPT"""

```

```python
# Optional: If you do the extraction and sum yourself, put that code here
```

```python
grader.check("Fit_curve")
```

<!-- BEGIN QUESTION -->

## Second step: Fit a cubic to the wrist force data (18 rows each successful/failed)

Split by successful and failed.

TODO: See slides for what the final plots should look like. I've handled the for loop and the plot indexing for you.

At this point, you should have the data in the form ts, ys for both the successful and the failed data. 

Where you're headed (pseudo code)

```python
- for 18/2 plots
-   for successful/failed
-      fit a cubic polynomial to the data
-      plot the data and the fitted polynomial
```

Implementation steps: You could write a function for doing the polynomial fit, but since there already is one (polynomial.polyval) you probably don't need one.

Steps (use the scratch cell to do the first part)
- Fit polynomial for one row of data
- Plot it and make sure it looks correct
- Copy that code into the for loop, adjusting for which row (p) and which subplot (axs[r, c])

Note: It can take a few seconds to plot.

Note 2: This is really only two lines of code; but you should spend some time understanding the provided code. 

```python
# EXAMPLE CODE - row column indexing, how to put n things into a grid

# If you have n things, and you want to put them in r rows, first figure out how many columns you'll need
#  Integer division: 
#    n // r  gives the number of times r can go into n
#    n % r gives the left over - it is always in the range 0 to n-1
n = 15
r = 4
c = n // r  # this gives the floor of n / r, which in this case will be 15/4 is 3
if n % r != 0:  # % is modulus - unless r goes exactly into n we need to add a bit (some of the last columns will be empty)
    c += 1   # Make c one bigger

# Go through all of n
for i in range(0, n):
    row = i // r   # Which row we're on
    col = i % r    # Which column we're on
    end_str = " "       # Use a space at the end of the pring
    if col == r-1:
        end_str = "\n"  # ... except at the end of the row
    print(f"n{i}r{row}c{col} ", end=end_str)  # No return

```

```python
# EXAMPLE CODE - zip
# These two pieces of code do the same thing
a_str = ("foo", "bar")
b_number = (3, 10)
for i in range(len(a_str)):
    a = a_str[i]
    b = b_number[i]
    print(f"Loop iteration {i} has {a} and {b}")

# Look, no i!
for a, b in zip(a_str, b_number):
    print(f"Loop iteration... um... has {a} and {b}")
```

```python
# SCRATCH CELL
# Practice the polynomial fit with one row of data
# Plot the data, plot the fitted polynomial
# You'll need ts to do both of these
```

```python
n_rows = 6
n_cols = 6
fig, axs = plt.subplots(n_rows, n_rows, figsize=(10, 10))

ch_name = "WF sum"
# This loops over the first half of the rows
for p in range(0, (n_rows // 2) * n_cols):
    # This is fancy python-ness to have the first time through the loop
    #   - first time, the row offset is 0, the data is data_wrist_force_successful, the str_title is "suc"
    #   - second time, the row offset is 3, the data is data_wrist_force_failed, the str_title is "fail"
    for row_offset, data, str_title in zip((0, 3), (data_wrist_force_successful, data_wrist_force_failed), ("suc", "fail")):
        # Yes, this works - it's the same as doing r = row_offset + p // ncols on one line,
        #   and c = p % n_cols on the second line
        r, c = row_offset + p // n_cols, p % n_cols
        axs[r, c].plot(ts, data_wrist_force_successful[p], label="Data")

        # Fit a polynomial of degree 3 to the variable data
        # Then plot the polynomial
        ...

        axs[r, c].set_xlabel("time (s)")
        axs[r, c].set_ylabel("N.m")
        # Build the title string from the channel name and the success/failed string
        axs[r, c].set_title(f"{ch_name} " + str_title, fontsize=10)

plt.tight_layout()
```

```python
# Manual grade
print("This is a manually-graded question; there is no grader.check() function. See rubric and slides for more information on expected output.")
```

<!-- END QUESTION -->

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS TCW3 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, lab 4 data analysis. Turn in just the .ipynb file - we will supply the data files.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for Fit_curve (points: 2)
```
>>> assert(data_wrist_force_successful.shape[0] == 355)
```
```
>>> assert(data_wrist_force_failed.shape[0] == 305)
```
```
>>> assert(np.isclose(np.max(data_wrist_force_successful), 20.46056888454263, atol=0.1))
```
```
>>> assert(np.isclose(np.max(data_wrist_force_failed), 18.862565694678903, atol=0.1))
```


Tests for Fitted_line_code (points: 2)
```
>>> assert(np.isclose(np.mean(end_times_successful, 0)[0], 3.98791579e-01, atol=0.01))
```
```
>>> assert(np.isclose(np.mean(end_times_successful, 0)[1], 4.07402569e+02, atol=10.0))
```
```
>>> assert(np.isclose(np.mean(end_times_failed, 0)[0], 0.40810473, atol=0.1))
```
```
>>> assert(np.isclose(np.mean(end_times_failed, 0)[1], 392.76044707, atol=10.0))
```


Tests for Read data (points: 1)
```
>>> assert(data_successful.shape == (355, 40) and data_failed.shape == (305, 40))
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```
