```python
# Initialize Otter
import otter
grader = otter.Notebook("HW_5_pinball_classes.ipynb")
```

# Pinball Classs 

One more class - the pinball class that does the simulation

In this homework you'll take the classes you created in the lecture activity and lab (**PinballWall** and **PinballBumper**) and use them in a class that wraps up the entire simulation (**Pinball**)

The pinball class is in **pinball.py**.

The homework is designed to have you copy your code in one step at a time, checking that it works before doing the next bit.

- Fill in **_compute_next_step**
- Get **_simulate_pinball** to handle *just* the ball going up and down (no reflections)
- Make sure plotting works
- Add in bouncing off of walls
- Add in bouncing off of bumpers
- Make a "fun" bounce

Slides: https://docs.google.com/presentation/d/1DCSykIMBSlzCNNnmUcxPqfGxJmeGYd-7LF_W_-_7fUU/edit?usp=sharing

```python
# The usual imports
import numpy as np
import json as json
```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

# Import PinballWall and PinballBumper and make sure they work

TODO import the class definitions and run the test cases for each of those classes

```python

```

```python
# Assumes you copied your wall class into lec_act_9_class_pinball_wall.py
from pinball_wall import PinballWall
from pinball_bumper import PinballBumper
```

```python
# TODO: Create at least one instance of each class and run their test functions to make sure they work
#. (you're free to create more than one, but at least create one wall and one bumper with the given names)

check_wall_instance = ...
check_bumper_instance = ...
"""
```

```python
grader.check("wall_and_bumper")
```

# Pinball class

TODO this question: Copy your **compute_next_step** code into the **_compute_next_step** method in **pinball.py*, and check that it works

```python
# Reminder that you will need to re-run this line when you make changes to the class
from pinball import Pinball
```

```python
# TODO Create the pinball class
#   Use: left -5.0, right 5.0, height 10.0, delta_t 0.1

pinball = ...
# END PROMPT"""
```

```python
# Check one step with the following start parameters
starting_state = np.zeros([3, 2])  # meters
starting_state[0, :] = [0, 0] # Start at zero, zero
# Velocity - mostly up with a bit of x
starting_state[1, :] = [-0.25, 5.0]
# Acceleration is really boring
starting_state[2, :] = [0.0, Pinball.acceleration_due_to_gravity()]

# TODO: Check first time step works
first_time_step = ...
print(f"Checking first time step {first_time_step}")
"""

# Notice my_pinball.delta_t - this is the delta t that you save and store in the pinball class
assert(np.all(np.isclose(first_time_step[0, :], starting_state[0, :] + pinball.delta_t * starting_state[1, :])))
assert(np.all(np.isclose(first_time_step[1, :], starting_state[1, :] + pinball.delta_t * starting_state[2, :])))
```

```python
grader.check("pinball_class_next_step")
```

# Simulate pinball method

TODO: Copy your simulate code into the **simulate_pinball** method in the class. For this question, just check that the simulation works without adding walls and bumpers. Adding walls/bumpers will happen in the next question.

Note: The plot code should work

```python
# Now run a simulation
#  setup
starting_state = np.zeros([3, 2])  # location, velocity, acceleration
starting_state[0, :] = [0, 0] # Start at zero, zero
# Velocity - mostly up with a bit of x noise
starting_state[1, :] = [2.4, 8.5]
starting_state[2, :] = [0.0, pinball.acceleration_due_to_gravity()]

# TODO: Run the simulation
...

# Check that last pose is below y = 0 line
print(pinball.simulate_pinball(starting_state)[0][1][-1])
```

```python
# ... and plot the results - note that the poses should have been saved in the class, so you don't need to pass them in
pinball.plot_pinball_hw()
```

```python
grader.check("add_simulation")
```

# Add walls to the simulation

Note: Walls are split out from bumpers here to enable partial credit; if you've done your bumpers correctly, the next question should just work.

- step 1: Use the code given in the todo section in **simulate_pinball** to add in looping over all of the obstacles, calling the **collided** method followed by the **reflect** method
- step 2: Check that calling **simulate_pinball** still works
- step 3: Add in a wall and make the ball bounce off of the wall
- step 4: Make sure plotting works
- step 5: Add in top, left, and right walls, and do the starting condition given (which should bounce off of all 3 walls)

```python
# TODO Now create and add three walls
# Required for this question: Side walls at -4.5, 4.5 and top at 9.5

# Remove all the obstacles in order to start "clean" when you execute this cell
pinball.remove_all_obstacles()

# Now create and add each wall instance

# This creates an instance of a wall and adds it to the pinball
my_wall = PinballWall("Vertical", -4.5)
pinball.add_obstacle(my_wall)
# END PROMPT"""
```

```python
# Check that simulate still works
# This velocity should make the ball bounce off of all walls
starting_state[1, :] = [10.4, 15.5]

...
```

```python
# ... and plot the results - this will call the wall draw code, so your walls should show up
pinball.plot_pinball_hw()
```

```python
grader.check("add_walls")
```

# Add bumpers to the simulation

TODO: Add at least one instance of a bumper to the simulation. Set up the starting state so that you hit the bumper. 

- Bumper center: 0.5, 5.0
- Bumper radius: 1.0

```python
# Remove all the obstacles in order to start "clean" when you execute this cell
pinball.remove_all_obstacles()

# TODO: Use add_obstacle to add at least one bumper
...
```

```python
# TODO Setup starting state, make sure you run into the bumper. Optional: Plot (but it will really help you to debug)

...
```

```python
grader.check("add_bumpers")
```

<!-- BEGIN QUESTION -->

# Create an interesting bounce

Set up a few obstacles so that you get multiple bounces. Plot the result.

```python
...
```

<!-- END QUESTION -->

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, hwk 5 class(ic) pinball

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for add_bumpers (points: 5)
```
>>> assert len(pinball.obstacles) > 0
```
```
>>> assert np.all(pinball.poses[1][0:-1] >= 0.0)
```
```
>>> assert pinball.obstacles[0].collided([0.45, 4.5])
```
```
>>> assert len(pinball.obstacles[0].reflect([0.45, 4.5], [0.5, 0.5])) == 2
```
```
>>> assert not pinball.obstacles[0].collided(pinball.obstacles[0].reflect([0.45, 4.5], [0.5, 0.5])[0])
```


Tests for add_simulation (points: 5)
```
>>> assert len(pinball.simulate_pinball(starting_state)) == 2
```
```
>>> assert pinball.poses.shape[1] > 2
```
```
>>> assert pinball.velocities.shape[1] > 2
```
```
>>> assert pinball.poses[1][-1] <= 0.0
```
```
>>> assert np.all(pinball.poses[1][0:-1] >= 0.0)
```


Tests for add_walls (points: 5)
```
>>> assert len(pinball.obstacles) == 3
```
```
>>> assert np.all(pinball.poses[1][0:-1] >= 0.0)
```
```
>>> assert pinball.poses[1][-1] <= 0.0
```
```
>>> assert np.all(pinball.poses[1][:] <= 9.75)
```
```
>>> assert np.all(pinball.poses[0][:] >= -4.75)
```
```
>>> assert np.all(pinball.poses[0][:] <= 4.75)
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```


Tests for pinball_class_next_step (points: 5)
```
>>> assert np.all(np.isclose(first_time_step[0, :], starting_state[0, :] + pinball.delta_t * starting_state[1, :]))
```
```
>>> assert np.all(np.isclose(first_time_step[1, :], starting_state[1, :] + pinball.delta_t * starting_state[2, :]))
```


Tests for wall_and_bumper (points: 5)
```
>>> assert check_wall_instance.inside([0.0, 0.0])
```
```
>>> assert not check_bumper_instance.inside([-100, -100])
```
