```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab_9_class_pinball_bumper.ipynb")
```

# Lab 9 classes 

TODO: Create a class for your bumper. Exactly how you create and store the bumper data is up to you, but you should define method(s) to determine if your pinball was inside/outside of the bumper, and a method to reflect the pinball if it is inside.

(Optional) add in your test code here as test methods on the class

Class code is in **pinball_bumper.py**

Slides: https://docs.google.com/presentation/d/172YYy0RNsGW4Gpruj3lejtBYuW0uTwZwjD-gMOrLYIc/edit?usp=sharing

```python
# The usual imports
import numpy as np
import matplotlib.pyplot as plt
```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

# Turn your bumper code into a class

The framework of the class is given. 

TODO: Copy the code from your last homework into **pinball_bumper.py**. Store the data for the bumper in the class.

Provide default values for all parameters.

Notice the method name changes to **inside** and **reflect** - we'll see why when you do the homework

You're welcome to add more methods to the class; just make sure not to rename the ones that are there or change their parameters

```python

```

```python
from pinball_bumper import PinballBumper
```

```python
grader.check("bumper_class")
```

## Testing

TODO: Create instance(s) of Bumper and check that they work

Make an instance of bumper called **my_bumper** that is centered at 0.5, 0.5, with radius 1

TODO Call your test functions

```python
my_bumper = ...
# END PROMPT"""
```

```python
# Put your test call(s) here
```

```python
grader.check("bumper_instance")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, lecture activity 9.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for bumper_class (points: 2)
```
>>> assert PinballBumper() is not None
```


Tests for bumper_instance (points: 2)
```
>>> assert my_bumper.inside([0.5, 0.25])
```
```
>>> assert not my_bumper.inside([1.5, 0.25])
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```
