```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_10_classes.ipynb")
```

# Creating a ground segment class

This is mostly filling in the methods in the class

You will have to decide what variables you want to keep in the class

```python
# Doing the imports for you
import numpy as np
import matplotlib.pyplot as plt
from numpy.polynomial import polynomial

```

# Ground segment class

Note, this is a simplified version of the one in ski_jump.py

In this problem, fill in **init** method and **plot** 

You're welcome to create more variables in the **__init__** function or add more methods to the class

TODO: 
- Fill in all of the methods in the class

```python
class SimpleGroundSection:
    def __init__(self, x_start, x_end, polynomial_coefs=(0, 0)):
        """
        @param x_start What x value to start at
        @param x_end What x value to end at
        @param polynomial coeficients for the section y(x) = a0 x^2 + a1 x + a2
        """
        # TODO - store the start and stop values along with the coefficients
        ...

    def is_between(self, x):
        """ Is the x value between the start and end x value?
        @param x the x value 
        @ return True or False
        """
        ...

    def height(self, x):
        """ Return the height, given the x (assumes over the ground)
        @param x - x value
        @returns height f(x)"""
        ...
        
    def plot(self, axs):
        """Plot the section
        @param axs the figure to plot into"""

        ...
```

```python
# Create instances and call test functions

# Create and test a single, horizontal linear segment at height 1
#   polynomial is p0 + p1 x + p2 x^2... etc
flat_bit = SimpleGroundSection(-0.5, 1.0, polynomial_coefs=(1, 0))
assert not flat_bit.is_between(-0.6)
assert flat_bit.is_between(-0.5)
assert flat_bit.is_between(0.0)    
assert flat_bit.is_between(1.0)    
assert not flat_bit.is_between(1.1)
assert flat_bit.is_over(0.1, 1.2)
assert np.isclose(flat_bit.height(0.0), 1.0)

fig, axs = plt.subplots(1, 1, figsize=(2, 2))
flat_bit.plot(axs)

```

```python

# Create and test a slightly descending straight line
#   Starts at 1 and drops by 0.25
slope_down_bit = SimpleGroundSection(1.0, 1.5, polynomial_coefs=(1.0 + 0.25, -0.25))
assert not slope_down_bit.is_between(0.9)
assert slope_down_bit.is_between(1.0)
assert slope_down_bit.is_between(1.25)    
assert slope_down_bit.is_between(1.5)    
assert not slope_down_bit.is_between(1.6)
assert slope_down_bit.is_over(1.25, 1.0)
assert np.isclose(slope_down_bit.height(1.0), 1.0)

fig, axs = plt.subplots(1, 1, figsize=(2, 2))
slope_down_bit.plot(axs)

```

```python

# Create and test a parabola
#   Interpolates the points (0, 2), (1, 1), (3, 2)
parabola = SimpleGroundSection(0.0, 2.0, polynomial_coefs=(2.0, -3.0/2.0, 1.0/2.0))
assert not parabola.is_between(-0.1)
assert parabola.is_between(1.0)
assert not slope_down_bit.is_between(2.1)
assert parabola.is_over(1.0, 1.5)
assert np.isclose(parabola.height(0.0), 2.0)
assert np.isclose(parabola.height(3.0), 2.0)


fig, axs = plt.subplots(1, 1, figsize=(2, 2))
parabola.plot(axs)

```

```python
grader.check("class_defn")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS 2U5 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

Make sure you include any .py files you are importing (pinball_wall.py).

- Submit this .ipynb file and all .py files you are using to Lecture activity 10: Classes

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

The primary problem with this submission is forgetting the .py files...

Tests for class_defn (points: 2)
```
>>> assert not flat_bit.is_between(-0.6)
>>> assert flat_bit.is_between(-0.5)
>>> assert flat_bit.is_between(0.0)
>>> assert flat_bit.is_between(1.0)
>>> assert not flat_bit.is_between(1.1)
>>> assert flat_bit.is_over(0.1, 1.2)
>>> assert np.isclose(flat_bit.height(0.0), 1.0)
```
```
>>> assert not slope_down_bit.is_between(0.9)
>>> assert slope_down_bit.is_between(1.0)
>>> assert slope_down_bit.is_between(1.25)
>>> assert slope_down_bit.is_between(1.5)
>>> assert not slope_down_bit.is_between(1.6)
>>> assert slope_down_bit.is_over(1.25, 1.0)
>>> assert np.isclose(slope_down_bit.height(1.0), 1.0)
```
```
>>> assert not parabola.is_between(-0.1)
>>> assert parabola.is_between(1.0)
>>> assert not slope_down_bit.is_between(2.1)
>>> assert parabola.is_over(1.0, 1.5)
>>> assert np.isclose(parabola.height(0.0), 2.0)
>>> assert np.isclose(parabola.height(3.0), 2.0)
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```
