```python
# Initialize Otter
import otter
grader = otter.Notebook("HW_4_robot_arm_class.ipynb")
```

# Robot arm using a class

Move homework 3 code (robot arm) into a class. We'll do this in the same order that you wrote functions in HW 3.


Slides: https://docs.google.com/presentation/d/1JnRTK3F3x1V8OUGUXh9-NRxkYniXa2oFEbUmi6SG8VY/edit?usp=sharing

```python
# The usual imports
import numpy as np
import matplotlib.pyplot as plt
from scipy.optimize import fmin
```

```python
# matrix_rouintes.py functions
#  If this doesn't work, make sure matrix_routines.py is in the same directory as this file
import matrix_routines as mt
```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

```python
# import the class
from arm_component import ArmComponent

# Note: Feel free to copy the class from arm_component.py rather than doing the
#  import while debugging this question. 
```

## Arm component

TODO: fill in the class methods for *ArmComponent*

This is setup so you can test one piece at a time.

```python
# Create an instance of the class
arm_component = ArmComponent("check_name", ArmComponent.points_in_a_wedge(), "check_color")

# Step 1: Check that variable names are created correctly
assert arm_component.name == "check_name"
assert arm_component.color == "check_color"
assert np.all(np.isclose(arm_component.matrix_shape, np.identity(3)))
assert np.all(np.isclose(arm_component.matrix_pose, np.identity(3)))

```

```python
# Step 2: add the points_in_square method
# Create an instance of the class, but with square points
arm_component = ArmComponent("check_name", ArmComponent.points_in_a_square(), "check_color")

assert np.isclose(arm_component.pts[0, 0], -1.0)
assert np.isclose(arm_component.pts[1, 0], -1.0)
assert np.all(np.isclose(arm_component.pts.shape, (3, 5)))
```

```python
# Step 3: Add the matrix_shape_X methods
#  Check 1, matrix_shape_base

base_component = ArmComponent("base", ArmComponent.points_in_a_wedge(), "Grey")
base_component.matrix_shape_base(base_width=0.5, base_height=0.25)

print(f"{base_component.matrix_shape}")
assert np.isclose(base_component.matrix_shape[0, 0], 0.25)
assert np.isclose(base_component.matrix_shape[1, 1], 0.125)

```

```python
#  Check 2, matrix_shape_link

link_component = ArmComponent("link", ArmComponent.points_in_a_square(), "Red")
link_component.matrix_shape_link(link_length=0.75, link_width=0.5)

print(f"{link_component.matrix_shape}")
assert np.isclose(link_component.matrix_shape[0, 0], 0.375)
assert np.isclose(link_component.matrix_shape[1, 1], 0.25)

```

```python
#  Check 3, matrix_shape_palm

palm_component = ArmComponent("link", ArmComponent.points_in_a_square(), "Red")
palm_component.matrix_shape_palm(palm_width=0.5)

print(f"{palm_component.matrix_shape}")
assert np.isclose(palm_component.matrix_shape[0, 0], 0.025)
assert np.isclose(palm_component.matrix_shape[1, 1], 0.25)

```

```python
#  Check 3, finger_component
#   No checks for this one, since it was optional in the original assignment. But you should still copy it in

finger_component = ArmComponent("link", ArmComponent.points_in_a_wedge(), "Green")
finger_component.matrix_shape_finger(palm_width=0.5, finger_size=(0.5, 0.2), b_is_top=True)
```

```python
grader.check("arm_component")
```

## Complete robot arm class

TODO: 
- Step 1: Edit the __init__  method in *robot_arm_class.py* 
- Step 2: Edit the set_angles method
- Step 3: Edit the  plot method

As before, you can either copy the class in or use the import method.

```python
# If the import is not working, feel free to copy the code for the class here
from robot_arm_class import RobotArm
```

```python
# TODO step 1 (init function)

# Create the arm geometry 
base_size_param = (0.5, 1.0)
link_sizes_param = [(0.5, 0.25), (0.3, 0.1), (0.2, 0.05)]
palm_width_param = 0.1
finger_size_param = (0.075, 0.025)

# The equivalent of create_arm_geometry
#   Fill in the init method
robot_arm = RobotArm(base_size_param, link_sizes_param, palm_width_param, finger_size_param)

# self.arm_geometry should be the list of arm components
assert len(robot_arm.arm_geometry) == 5
```

```python
## Check the combined link/gripper/finger rotations
# Several different angles to check your routines with 
# Pass the one you want to check into set_angles_of_arm_geometry below
#  Feel free to change these
angles_none = [0.0, 0.0, 0.0, [0.0, 0.0, 0.0]]
angles_check_link_0 = [np.pi/4, 0.0, 0.0, [0.0, 0.0, 0.0]]
angles_check_link_0_1 = [np.pi/4, -np.pi/4, 0.0, [0.0, 0.0, 0.0]]
angles_check_wrist = [np.pi/2, -np.pi/4, -3.0 * np.pi/4, [np.pi/3.0, 0.0, 0.0]]
angles_check_fingers = [np.pi/2, -np.pi/4, -3.0 * np.pi/4, [0.0, np.pi/4.0, -np.pi/4.0]]

# Don't change this one
angles_check = [np.pi/2, -np.pi/4, -3.0 * np.pi/4, [np.pi/3.0, np.pi/4.0, -np.pi/4.0]]

robot_arm.set_angles(angles_check)
```

```python
fig, axs = plt.subplots(1, 1, figsize=(6, 6))
robot_arm.plot_complete_arm(axs)
```

```python
mat_check_link_1 = np.array([[ -1.0,  0.0,  0.0], \
                             [  0.0, -1.0,  0.5], \
                             [  0.0,  0.0,  1.0]])
                             
print(robot_arm.arm_geometry[1].matrix_pose)
assert(np.all(np.isclose(robot_arm.arm_geometry[1].matrix_pose, mat_check_link_1, atol=0.01)))
```

```python
mat_check_link_2 = np.array([[ -0.7071, -0.7071, -0.5], \
                             [  0.7071, -0.7071,  0.5], \
                             [  0.0,     0.0,  1.0]])
print(robot_arm.arm_geometry[2].matrix_pose)
assert(np.all(np.isclose(robot_arm.arm_geometry[2].matrix_pose, mat_check_link_2, atol=0.01)))
```

```python
mat_check_link_3 = np.array([[ 1.0, 0.0, -0.71213], \
                             [ 0.0, 1.0,  0.71213], \
                             [ 0.0, 0.0,  1.0]])
print(robot_arm.arm_geometry[3].matrix_pose)
assert(np.all(np.isclose(robot_arm.arm_geometry[3].matrix_pose, mat_check_link_3, atol=0.01)))
```

```python
mat_check_wrist = np.array([[ 0.5, -0.8660,  -0.5121], \
                            [ 0.8660,  0.5,   0.7121], \
                            [ 0.0,  0.0,  1.0]])
print(robot_arm.arm_geometry[-1][0].matrix_pose)
assert(np.all(np.isclose(robot_arm.arm_geometry[-1][0].matrix_pose, mat_check_wrist, atol=0.01)))
```

```python
grader.check("robot_arm")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS TAF4 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

- Submit this .ipynb file AND (if you did not copy the classes in) arm_components.py and robot_arm_classes.py. 
- We will supply matrix_routines.py for you (it won't break anything if you do include it)

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Make sure you remove all the print statements you put in that print out lots of stuff.

Tests for arm_component (points: 15)
```
>>> assert arm_component.name == 'check_name'
```
```
>>> assert arm_component.color == 'check_color'
```
```
>>> assert np.all(np.isclose(arm_component.matrix_shape, np.identity(3)))
```
```
>>> assert np.all(np.isclose(arm_component.matrix_pose, np.identity(3)))
```
```
>>> assert np.isclose(arm_component.pts[0, 0], -1.0)
```
```
>>> assert np.isclose(arm_component.pts[1, 0], -1.0)
```
```
>>> assert np.all(np.isclose(arm_component.pts.shape, (3, 5)))
```
```
>>> assert np.isclose(base_component.matrix_shape[0, 0], 0.25)
```
```
>>> assert np.isclose(base_component.matrix_shape[1, 1], 0.125)
```
```
>>> assert np.isclose(link_component.matrix_shape[0, 0], 0.375)
```
```
>>> assert np.isclose(link_component.matrix_shape[1, 1], 0.25)
```
```
>>> assert np.isclose(palm_component.matrix_shape[0, 0], 0.025)
```
```
>>> assert np.isclose(palm_component.matrix_shape[1, 1], 0.25)
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```


Tests for robot_arm (points: 10)
```
>>> assert len(robot_arm.arm_geometry) == 5
```
```
>>> assert np.all(np.isclose(robot_arm.arm_geometry[1].matrix_pose, mat_check_link_1, atol=0.01))
```
```
>>> assert np.all(np.isclose(robot_arm.arm_geometry[2].matrix_pose, mat_check_link_2, atol=0.01))
```
```
>>> assert np.all(np.isclose(robot_arm.arm_geometry[3].matrix_pose, mat_check_link_3, atol=0.01))
```
```
>>> assert np.all(np.isclose(robot_arm.arm_geometry[-1][0].matrix_pose, mat_check_wrist, atol=0.01))
```
