```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_10_fibonacci_class.ipynb")
```

# Making it a class

We're going to take the instrumented, cached version of the fibonacci sequence and turn it into a class. This is mostly (more) practice with classes as data encapsulation.

# Move to a class

TODO: 
- add the cache variable
- copy the fib_cached_instrumented code into fib_sequence
-  Change what needs to be changed (see below)

```python
class FibonacciClass:
    # We haven't done this before - this is exactly the same as doing func.call_count
    call_count = 1
    # TODO: Add the cache variable initialization here
    ...

    def __init__(self):
        """ Initialize call_count and cache here"""
        return

    def fib_sequence(self, n):
        """ Calculate the given sequence
        @param n
        @return fibonacci sequence value of n"""
        # TODO: Copy in the fib_cached_instrumented code here. You'll need to make the following changes:
        #   Use ClassName.xx instead of FunctionName.xxx
        #   Use self.fib_sequence to call routine again
        #   Take out the AttributeError in the try except (you don't need it because you initialized the variables
        #     above)

        ...

```

```python
# Make an instance
fib_class = FibonacciClass()

# Do some tests to make sure fibonacci functions are working as expected
#   This is a hard-coded set of correct answers to test against
fib_tests = {0: 0, 1: 1, 2: 1, 3: 2, 4: 3, 5: 5, 6: 8, 7: 13}

for number, value in fib_tests.items():
    # Call the function and compare it to the test value
    if fib_class.fib_sequence(number) != value:
        print('  failed for {0}: expected {1}, got {2}'.format(number, value, fib_class.fib_sequence(number)))

print(f'Number of calls {fib_class.call_count}')
```

```python
grader.check("fib_class")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, week 10 runtime

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for fib_class (points: 4)
```
>>> assert(fib_class.call_count == 21)
```
```
>>> assert(fib_class.cache[1])
```
```
>>> assert(fib_class.cache[7])
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```
