```python
# Initialize Otter
import otter
grader = otter.Notebook("HW_1b_numpy_and_plotting.ipynb")
```

# HW 1b: Statistical analysis of data using numpy and matplotlib

Slides: (https://docs.google.com/presentation/d/1ef0msC9XIT37_Yg94Cf4TL9VBgbIIfkKik8jgshskjE/edit?usp=sharing)

Please do Homework 1a before starting this assignment.

In this homework you will focus on analysis - categorizing the type of a given motion (snap, high five, or clap) and determined its accuracy on the provided dataset.

## Getting started

GUIDE: From HW 1b, copy the library imports, data load, and the functions. 

Everything here should be a direct copy from HW1. Each JN cell has a comment for what should be copied into that cell. Do not just copy everything over into one cell.

```python
# GUIDE: Do the imports here

```

```python
# GUIDE: Load the data set

# Numeric ids to indicate hand motion type from Lab 1
clap_id = 1
high_five_id = 2
snap_id = 3

# Hardcoded column numbers for this data to make the code easier; normally, you would extract this from
# full_dataset_data_description.json.
timestamp_col = 0
subject_number_col = 13
run_number_col = 14
motion_id_col = 15
```

```python
# GUIDE: The function for getting the channel data, given the data, the offset, and the dimensions

```

```python
# GUIDE: The function for getting the descriptor for the data channel, given the name
```

```python
# GUIDE: The function for getting the root mean square error
```

```python
# GUIDE: The function for pulling out just the given subject number, run number, and motion id from the full data set
```

```python
# GUIDE: The function for pulling out the row index that has the given value from the data
```

```python
# Add some tests to make sure your functions/data loading are correct
```

```python
grader.check("copy_code")
```

## Problem 1: Graph RMS for each motion type

You did a basic plot in homework 1a. In this problem, you will adjust the y axis of all of the plots so that it will be easier to compare min/max values across different subjects/motions.

GUIDE
- Copy the plot function over.
- Determine a good y axis maximum value for each motion type by plotting a bunch of them
- Change the plot function so that all motions plot with the same y axis values

This question has a manually-graded portion; if your plots are incorrect you will have points deducted on Gradescope. See slides for what the plots should look like.

```python
# GUIDE: 
# Step 1: Copy the plot function from the previous homework 
# Step 2: Add a parameter that specifies the height of the y axis
# Step 3: Use that parameter to set the y axis (search matplotlib set y axis range)
```

```python
# GUIDE: Copy the code that gets the RMS of the accelerometera data for the left and right hand of subject 6

snap_subject_number = 6
snap_run_number = 3

clap_subject_number = 6
clap_run_number = 3

high_five_subject_number = 6
high_five_run_number = 3

```

```python
# GUIDE: Make the same plot as from homework 1, for subject 6 run number 3 - but set the y limit to something reasonable
#. so that you can see the plots but they all have the same y axis limits
# Plot size
nrows = 3
ncols = 2
fig, axs = plt.subplots(nrows, ncols, figsize=(6, 8))

y_lim = ...
fig.tight_layout()
plt.plot()
```

```python
grader.check("graph_rms")
```

## Determine motion type from data

Observe your graphs from the previous problem. How can you tell, from these graphs alone, what motion type the subject is performing?

Some helpful context:

* Subjects were told to snap with their _right hand_ and high five with their _left hand_.
* The high five was pantomimed; that is, the subject did not make contact with another hand, but instead made a high five motion.

From your observations, write a function that takes in the left hand and right hand accelerometer RMS values and returns the most likely motion type that the subject is performing.

Hint: Look at the peaks in the graphs, the magnitude of the peaks, and the hand(s) that the peaks occur on.

Hint: You could return to the previous problem and graph different subjects/runs for inspiration.

```python
# GUIDES: Put any additional functions you want here
```

```python
# GUIDES: Write a function that analyzes the data and returns the id of the one
#  it guesses
def guess_motion_type(lh_accelerometer_rms, rh_accelerometer_rms):
    """ Guesses the motion that the subject is performing given the RMS of the user's left hand and right hand accelerometers.
    @param lh_accelerometer_rms - RMS of the left hand accelerometer
    @param rh_accelerometer_rms - RMS of the right hand accelerometer
    @return The motion ID of the motion the user was most likely performing.
    """
    pass

```

```python
# GUIDES: Write some tests here to check that you have the "right" answer
#  i.e., if you give it snap data it should return snap (for at least some of the snap data)
```

```python
# GUIDE: get the s3 s4 snap data (from homework 1a)

# # example test code
s3s4_snap_lh_accelerometer_data = get_channel_data(s3s4_snap_data, index_offset=lh_accelerometer_descriptor["index_offset"], n_dims=lh_accelerometer_descriptor["dimensions"])

assert(snap_id == guess_motion_type(rms(s3s4_snap_lh_accelerometer_data), rms(s3s4_snap_rh_accelerometer_data)))
```

```python
grader.check("determine_motion_type")
```

## Calculate accuracy of your heuristic

You now have a function that guesses the motion type for a given motion, and we have ground truth labels for what motions the subjects are performing. Let's calculate the accuracy of your heuristic using this function!

The data set contains 10 subjects (1 through 10) that perform 10 runs (1 through 10) of each motion type (snap, high five, and clap). You can use three nested for loops (one for participant number, one for run number, one for motion ID) to iterate through each of the motion observations and guess the motion type.

We expect to see accuracies above 75%, and accuracies above 90% are possible.

```python
# GUIDES: Produce a numpy array that contains a row with 4 columns containing subject number, run number, actual motion, and guessed motion
accuracy_info = ...

# GUIDES: Count the number of correct and incorrect guesses.
# Hint: np.where is useful here.
correct_count = ...
incorrect_count = ...

# GUIDES: Use correct_count and incorrect_count to calculate accuracy
# Hint: You will need to do conversions to floats to avoid doing integer division.
accuracy = ...

print(f'guess_motion_type is {accuracy * 100}% accurate')
```

```python
grader.check("calculate_accuracy")
```

## Find the "worst" subject

Which subject does your heuristic perform the worst on?

```python
# GUIDES: Populate a numpy array with one entry per subject, counting the number of 
# incorrect motion type predictions.
incorrect_by_subject = ...

# GUIDES: Use incorrect_by_subject to determine the subject number that guess_motion_type
# performed worst on.
# Hint: find_row_with_value might be useful here.
worst_subject_number = ...

print(f'guess_motion_type performed worst on subject {worst_subject_number}')
```

```python
grader.check("worst_subject")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python
import os

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS FA26 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# VS Code stores the path in a special global variable
if '__vsc_ipynb_file__' in globals():
    notebook_path = globals()['__vsc_ipynb_file__']
    notebook_name = os.path.basename(notebook_path)
    json_name = notebook_name[:-6] + "_source.json"
    if not os.path.exists(json_name):
        print(f"Could not find the json file {json_name}; make sure the required VSCode extension is installed and running")

```

```python
grader.check("hours_collaborators")
```

### To submit

- Do a restart then run all to make sure everything runs ok
- Remove excessive print statements - don't make us search through pages of numbers to find your plots
- Save the file
- Submit this .ipynb file, along with the .json file, through gradescope, HWK1 numpy and Plotting
- You do NOT need to submit the data files - we will supply those

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Most likely failure for this assignment is not naming the data directory and files correctly; capitalization matters for the Gradescope grader. 

Tests for calculate_accuracy (points: 3)
```
>>> assert accuracy >= 0.75
```
```
>>> assert np.isclose(accuracy, float(correct_count) / float(correct_count + incorrect_count))
```
```
>>> assert correct_count == len(np.where(accuracy_info[:, -2] == accuracy_info[:, -1])[0])
```
```
>>> assert incorrect_count == len(np.where(accuracy_info[:, -2] != accuracy_info[:, -1])[0])
```


Tests for copy_code (points: 2)
```
>>> assert full_dataset.shape == (31502, 16)
```
```
>>> assert get_descriptor(data_description, 'Left hand accelerometer')['index_offset'] == 7
```
```
>>> assert get_descriptor(data_description, 'Left hand accelerometer')['dimensions'] == 3
```
```
>>> assert np.isclose(rms(get_channel_data(full_dataset, get_descriptor(data_description, 'Left hand accelerometer')['index_offset'], get_descriptor(data_description, 'Left hand accelerometer')['dimensions']))[0], 0.9700343636524773, atol=0.001)
```


Tests for determine_motion_type (points: 5)
```
>>> assert snap_id == guess_motion_type(rms(s3s4_snap_lh_accelerometer_data), rms(s3s4_snap_rh_accelerometer_data))
```


Tests for graph_rms (points: 2)
```
>>> assert 5 < y_lim < 10
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```
```
>>> assert os.path.exists(json_name)
```


Tests for worst_subject (points: 3)
```
>>> assert worst_subject_number > 0
```
```
>>> assert worst_subject_number <= 10
```
```
>>> assert find_row_with_value(incorrect_by_subject, np.max(incorrect_by_subject)) == worst_subject_number - 1
```
