```python
# Initialize Otter
import otter
grader = otter.Notebook("HW_1a_numpy_and_plotting.ipynb")
```

# HW 1: Statistical analysis of data using numpy and matplotlib

Slides: (https://docs.google.com/presentation/d/1ef0msC9XIT37_Yg94Cf4TL9VBgbIIfkKik8jgshskjE/edit?usp=sharing)

Please do lab 1 before starting the first part of this homework and lab 2 before the second part. This homework will continue in Homework 1b.

In this homework the focus is on code *design*. Most of the functionality of the code will be what you did in the labs. I don't expect you to do the code design; I've provided that. Just pay attention to how a couple of code structuring tools (functions, dictionaries) can make your code a little more re-usable, cleaner, and less prone to error. 

Learning to *read* code is also as important as *writing* code. This is (hopefully) a gentle introduction to using/interacting with more advanced concepts/syntax/semantics. 

By the end of the homework assignment, you will have written a method for categorizing the type of a given motion (snap, high five, or clap) and determined its accuracy on the provided dataset.

## Week 1, problem 0: Setting up libraries

GUIDES: Import the libraries that you will need (numpy, json, matplotlib, os)

```python
# GUIDES: put the numpy, os, json, and matplotlib imports here so you can use those libraries (see the tops of Labs 1-3)
#  Parts 1 and 2 only need numpy, os and json. Part 3 needs matplotlib; you can either look at lab 3 and copy the import now or do
#  it later
```

```python
grader.check("libraries")
```

## Week 1, problem 1: Read data set and description

In Lab 1, you read in a few CSV files, where each contained the data from a single motion from a single person. In this problem, you will read in a single CSV file that contains a full dataset containing multiple types of motions from multiple people ("subjects") done multiple times.

The data is stored in `Data/full_dataset.csv`, and is described in `Data/full_dataset_data_description.json`. Note how `full_dataset_data_description.json` is different from `data_description.json`; there are three new data channels:

* Subject number: Identifies the person ("human subject") performing the motion. There are 10 different people in the data set.
* Run number: Each person did each motion 10 times, and this number identifies each unique run of the motion.
* Motion type: Contains the motion ID of the motion performed, corresponding to a high five, snap, or clap.

Almost all of this homework involves copying the code you wrote in Lab 1 but reading these two different files.

A reminder that capitalization matters for file names - data/full_dataset.csv is NOT the same as Data/full_dataset.csv. 

```python
# GUIDES: Read in `Data/full_dataset.csv` into `full_dataset`, and `Data/full_dataset_data_description.json` into `data_description`.

full_dataset = ...
data_description = ...
```

```python
grader.check("read_data")
```

## Week 1, problem 2: Filtering data set

In this homework, you will need to perform operations on individual motions identified by a specific subject, motion type, and run number. For this part, write code to extract the fourth snap performed by subject three. We will later turn this into a function.

```python
# Numeric ids to indicate hand motion type from Lab 1
clap_id = 1
high_five_id = 2
snap_id = 3

# Hardcoded column numbers for this data to make the code easier; normally, you would extract this from
# full_dataset_data_description.json.
timestamp_col = 0
subject_number_col = 13
run_number_col = 14
motion_id_col = 15
```

```python
# GUIDES: Extract a numpy array containing the rows for the fourth snap performed by subject three.
# Hint: You can apply three different boolean filters to whittle down the data set, 
#  use np.all() to combine the filters into one filter, or use logical and.

# Reminder to use the variable snap_id (defined above) when searching for rows that are snaps

s3s4_snap_data = ...
```

```python
grader.check("filter_data_set")
```

## Week 1, problem 3: Calculate RMS

In this homework, you will be writing code to determine which motion a person is performing. One common way to differentiate between different motion types is with well chosen summary statistics of the data. This can include statistics on individual axes, as well as statistics that combine multi-dimensional data (e.g., the x, y, and z data from accelerometers).

We will use the root mean square (RMS) of the accelerometer data as one such summary statistic. RMS is particularly useful to us here since it is an excellent information stream for detecting contact (such as impact between hands, collisions between fingers as you snap), even if hand orientation is not perfectly consistent.

For this problem, you will write a function to calculate the RMS of each row in an array, and use it to calculate the RMS of two specific data channels for the entire data set.

```python
# GUIDES: Copy the function get_channel_data from lab 2.

```

```python
# GUIDES: Copy the function get_descriptor from lab 2.

```

```python
# definition of RMS (Root Mean Squared): sqrt( mean(sum_x x^2) )
def rms(in_data):
   """ Calculate the root mean square of each row in in_data.
   @param in_data - A 2-dimensional numpy array
   @return An n rows by 1 column numpy array containing the RMS of each row in in_data.
   """
   # GUIDES: Fill in this function.
   # Hint: Use numpy functions to operate on the entire array at once rather than
   # iterate through each row.
   pass
```

```python
# Check that the rms function works as expected
test_array = np.array([[1, 2, 3], [4, 5, 6]])

# Calculation for first row (reminder, mean is sum of values divided by number of values)
first_row_rms = np.sqrt((1 * 1 + 2 * 2 + 3 * 3) / 3)

test_array_rms = rms(test_array)
assert(np.isclose(test_array_rms[0], first_row_rms))
assert(np.isclose(test_array_rms[1], 5.06622805))
```

```python
# GUIDES: Use rms, full_dataset, get_descriptor, and get_channel_data to calculate the RMS of left hand accelerometer and right hand accelerometer for the entire data set.

lh_accelerometer_descriptor = ...
lh_accelerometer_rms_check = ...

rh_accelerometer_descriptor = ...
rh_accelerometer_rms_check = ...
```

```python
grader.check("calculate_rms")
```

## Week 2, problem 1: Function for filtering data set

Take the code from Week 1 problem 2 and turn it into a re-usable function.

```python
def filter_data_set(data, subject_number, run_number, motion_id):
    """ Filters a data set to a specific run of a specific motion type from a specific subject.
    @param data - The full dataset, containing multiple motion runs from multiple people.
    @return Just the rows for the given subject number, run number, and motion id.
    """
    pass
```

```python
# GUIDES: Write a test function. Call filter_data_set with the values for s3s4_snap_data,
#   i.e., subject number 3, run number 4, snap data
# And check that the resulting data is the same as s3s4_snap_data

foo = filter_data_set(full_dataset, 3, 4, snap_id)
```

```python
grader.check("filter_function")
```

## Week 2, problem 2: Finding peak and valley of a data channel

Often, the maximum and minimum value of a data channel are informative as further summary statistics, in addition to being helpful for pinpointing the occurrence of key events such as moments of hand contact in the recorded data. The difference between these maximum and minimum values can be helpful, too, as you try to understand the range of your data. 

For this problem, you will write a function that locates the first row containing a specific value, and then use it to find the peak and valley of a data channel. The function will be very similar to the optional section of lab 2.

While data can have multiple peaks and valleys, for simplicity we will only look at the first location with the maximum and minimum value for the data channel.

A reminder that there is a tutorial for using **np.where**.

```python
def find_row_with_value(data, value):
    """ Locates the first index of data containing value.
    @param data - A 1-d array containing data.
    @param value - A specific value to look for in data.
    @return The index of value within data, or -1 if value is not found.
    """
    # GUIDES: Write this function. See the optional section of lab 2 for hints on how to do this.
    pass

```

```python
# Simple test code for your function to help debug issues.
find_row_with_value_test_array = np.array([6, 7, 8, 7, 5])

assert(find_row_with_value(find_row_with_value_test_array, 6) == 0)
assert(find_row_with_value(find_row_with_value_test_array, 7) == 1)
assert(find_row_with_value(find_row_with_value_test_array, 5) == 4)
```

```python
# GUIDES: Find peak and valley of z axis of the right hand accelerometer in s3s4_snap_data
# Hint: You can use get_channel_data to extract just the rh accelerometer channel from s3s4_snap_data.

s3s4_snap_rh_accelerometer_data = ...
s3s4_snap_rh_accelerometer_z_min = ...
s3s4_snap_rh_accelerometer_z_max = ...

s3s4_snap_rh_accelerometer_z_min_index = ...
s3s4_snap_rh_accelerometer_z_max_index = ...
```

```python
grader.check("peak_and_valley")
```

## Week 2, problem 3: Graph RMS for each motion type

We have the full data set, a way to filter it, and the ability to calculate RMS for a given data channel. Let's combine these abilities and graph the left hand and right hand accelerometer RMS for a single clap, snap, and high five. After that, we can take a step back and observe how the graph differs between motion types.

This question has a manually-graded portion; if your plots are incorrect you will have points deducted on Gradescope. See slides for what the plots should look like.

```python
# GUIDES: Copy plot_function from Lab 3.
```

```python
# GUIDES: Create a plot with 3 rows and 2 columns. Plots in the left column contain the RMS of the left hand accelerometer,
# and plots in the right column contain the RMS of the right hand accelerometer. The first row contains graphs for a snap,
# the second row a clap, and the third row a high five. You can plot any single snap, clap, or high five in the data set.

# Use the variables below when filtering the data.
# For the final submitted graph, please use subject 6 run 3.
# But prior to submission, feel free to explore the data
# by looking at graphs for other subjects/runs!

snap_subject_number = 6
snap_run_number = 3

clap_subject_number = 6
clap_run_number = 3

high_five_subject_number = 6
high_five_run_number = 3

# Hint: For calculating and using RMS, you have two options:
#
#   I've copied `full_dataset` into `all_data` for you
# * Add two columns to the end of `all_data` containing the RMS of the left hand accelerometer and the right hand accelerometer
#   for every row in the data set, and then filter `all_data` to the specific subject number, run number, and motion type.
#   You can use `np.append` to add these columns.
#
# OR
#
# * Filter `all_data` to the specific subject number, run number, and motion type, and then calculate the RMS of the left hand
#   accelerometer and right hand accelerometer for that motion in separate numpy arrays.

all_data = np.copy(full_dataset)

```

```python
# Plot size
nrows = 3
ncols = 2
fig, axs = plt.subplots(nrows, ncols, figsize=(6, 8))

# GUIDES: Make the plot
# Put left hand in the first column, right hand in the second column.
# Hint: To use plot_function with your new "rms" data channel, you will need to make a
# dictionary for its descriptor with "name" and "units" entries.
# Note: The time values in the plots in the slides are divided by 1,000

fig.tight_layout()
plt.plot()

```

```python
# Manual grade
print("")
```

```python
grader.check("graph_rms")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python
import os

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS FA26 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# VS Code stores the path in a special global variable
if '__vsc_ipynb_file__' in globals():
    notebook_path = globals()['__vsc_ipynb_file__']
    notebook_name = os.path.basename(notebook_path)
    json_name = notebook_name[:-6] + "_source.json"
    if not os.path.exists(json_name):
        print(f"Could not find the json file {json_name}; make sure the required VSCode extension is installed and running")

```

```python
grader.check("hours_collaborators")
```

### To submit

- Do a restart then run all to make sure everything runs ok
- Remove excessive print statements - don't make us search through pages of numbers to find your plots
- Save the file
- Submit this .ipynb file and .json file through gradescope, HWK1 numpy and Plotting
- You do NOT need to submit the files in the Data directory - we will supply those

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Most likely failure for this assignment is not naming the data directory and files correctly; capitalization matters for the Gradescope grader. 

Tests for calculate_rms (points: 3)
```
>>> assert np.isclose(test_array_rms[0], 2.1602469, atol=0.001)
```
```
>>> assert np.isclose(test_array_rms[1], 5.06622805, atol=0.001)
```
```
>>> assert np.isclose(lh_accelerometer_rms_check[0], 0.9700343636524773, atol=0.001)
```
```
>>> assert np.isclose(lh_accelerometer_rms_check[-1], 0.1439907, atol=0.001)
```
```
>>> assert np.isclose(rh_accelerometer_rms_check[0], 0.8847033, atol=0.001)
```
```
>>> assert np.isclose(rh_accelerometer_rms_check[-1], 0.522844, atol=0.001)
```


Tests for filter_data_set (points: 3)
```
>>> assert np.isclose(s3s4_snap_data[0, 0], 38888.0)
```
```
>>> assert np.isclose(s3s4_snap_data[0, -1], 3.0)
```
```
>>> assert np.isclose(s3s4_snap_data[-1, 0], 39338.0)
```
```
>>> assert np.isclose(s3s4_snap_data[-1, -1], 3.0)
```


Tests for filter_function (points: 1)
```
>>> assert np.allclose(filter_data_set(full_dataset, 3, 4, snap_id), s3s4_snap_data)
```


Tests for graph_rms (points: 3)
```
>>> assert len(axs[1, 1].lines) == 1
```
```
>>> assert len(axs[0, 0].lines) == 1
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```
```
>>> assert os.path.exists(json_name)
```


Tests for libraries (points: 2)
```
>>> assert np.min(np.zeros((2, 2)) == 0.0)
```
```
>>> assert json.loads('10') == 10
```


Tests for peak_and_valley (points: 3)
```
>>> assert find_row_with_value(find_row_with_value_test_array, 6) == 0
```
```
>>> assert find_row_with_value(find_row_with_value_test_array, 7) == 1
```
```
>>> assert find_row_with_value(find_row_with_value_test_array, 5) == 4
```
```
>>> assert np.isclose(s3s4_snap_rh_accelerometer_z_min, -3.56)
```
```
>>> assert np.isclose(s3s4_snap_rh_accelerometer_z_max, 3.24)
```
```
>>> assert s3s4_snap_rh_accelerometer_z_min_index == 47
```
```
>>> assert s3s4_snap_rh_accelerometer_z_max_index == 46
```


Tests for read_data (points: 2)
```
>>> assert full_dataset.shape == (31502, 16)
```
```
>>> assert len(data_description['data_channels']) == 8
```
