```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab_2_functions.ipynb")
```

# Lab 2: Statistical analysis of data using numpy

Lab slides: (https://docs.google.com/presentation/d/1ykwwcQ0onMvAjUxfJmKl9tbo-rJPdB5pRwDEmpDsd-g/edit?usp=sharing)

For this lab the goal is to write functions to pull out one data channel and print out statistics for different types of hand motions. We will write a function to load data from a single file, a second to combine data from multiple files, a third to pull out the data for a selected data channel, then a fourth to calculate statistics.

Written properly, you only need one function to do stats for the entire data channel or for just one type of hand motion. For any data channel. In the homework you'll use this logic to apply statistics to a specific data channel.

```python
# Libraries that we need to import - numpy and json (for loading the description file)
import numpy as np
import json as json
# os is needed for calling os.path.basename
import os
```

### Reading in data

Based on your code from Lab 1, write a function `get_data` that loads the data from a single CSV and adds a column at the end containing the hand motion ID. Then, write a function `get_data_from_files` that loads data from a list of CSV files into a single numpy array using `get_data`.

Both functions will operate on a dictionary containing `"csv_path"` which holds the path to the data file and `"motion_id"` which holds the ID of the motion contained within the CSV file. You will write a third function `get_file_info` that returns this dictionary for a given CSV file, which will extract the motion type from the file name.

Note: The functions initially just contain `pass`, which is just a placeholder that you should remove. In Python, `pass` does nothing at all by design, but removing it and having an empty function is illegal ([docs](https://docs.python.org/3/tutorial/controlflow.html#pass-statements)).

```python
# Numeric ids to indicate hand motion type from Lab 1
clap_id = 1
high_five_id = 2
snap_id = 3
```

```python
# SCRATCH cell
fname_chop_up = "Data/S01C01.csv"
res = os.path.basename(fname_chop_up)

# TODO:  look at res. How would you get out the C F or S character?
```

```python
# TODO: Write the function to take csv_path and return a dictionary
# Step 1: Use `os.path.basename` to get the filename of the CSV file (e.g., S10F01.csv),
# Step 2: Extract the C/F/S character which will appear at a fixed offset in the filename,
# and then 
# Step 3: find the right motion ID based on that character. Remember that C is a
# clap, F is a high five, and S is a snap.
# Step 4: Make a dictionary and return it (see @return for details)
def get_file_info(csv_path):
    """Function that returns a file_info dictionary for a given filepath.
    @param csv_path - path to a CSV file containing a hand motion
    @return A dictionary with key "csv_path", containing csv_path, and "motion_id", containing the type of motion encoded in the file."""
    # The base file path is of the form S##[C|F|S]##.csv.
    pass
```

```python
# Test cell: If you wrote the above function correctly, this should work!
assert(get_file_info("Data/S01C01.csv") == {"csv_path": "Data/S01C01.csv", "motion_id": clap_id})

# TODO... Write another assert test for the high five data 
```

```python
# TODO: Actually read in the data and add the motion id to the end
#   Reminder: file_info will be the dictionary you just created in the above function
def get_data(file_info):
    """ Function that returns the data from the given CSV file.
    @param file_info - a dictionary with keys "csv_path" and "motion_id"
    @return Return array should contain data in file with an extra column at the end containing the motion_id."""
    pass
```

```python
# Test cell: TODO: Write a test that 
#   Calls get_file_info to create a dictionary
#   Passes that dictionary to get_data
#   Checks that the returned data is correct
#     1) the last column is the expected motion id
#     2) the data is correct 
```

```python
# TODO: This function reads in multiple data sets (not just one)
#  Reminder: Each element of file_list is a dictionary 
#   You should be calling your get_data function from this function...
def get_data_from_files(file_list):
    """ Function that returns data from a list of files.
    @param file_list - a list of dictionaries, where each dictionary contains `csv_path` and `motion_id`.
    @return A single return array containing the data from all of the given input files."""
    # Hint: Use np.concatenate to combine multiple numpy arrays.
    pass
```

```python
# SCRATCH cell
# TODO: Use this cell to check that all three functions are working correctly

all_data_check = get_data_from_files([get_file_info("Data/S01C01.csv")])


```

```python
# Read in data from the files in lab 1 using the functions you just wrote.
all_data = get_data_from_files([
    get_file_info("Data/S01C01.csv"),
    get_file_info("Data/S01F01.csv"),
    get_file_info("Data/S01S01.csv"),
])

# The autograder will check that get_file_info is correct, then that
#  the last columns of all_data are set to the correct motion ids,
# then check that the actual data in all_data is correct
# Reminder that you can look in the files (eg, S01F01.csv) to see what the values should be
# Reonder 2: If you have print statements in your functions that will also cause the autograder to fail
```

```python
grader.check("get_data")
```

## Doing the slice

Get the data for one of the channels. 

```python
# This reads in the json data
try:
    with open("Data/data_description.json", "r") as fp:
        data_description = json.load(fp)
except FileNotFoundError:
    print(f"The file was not found; check that the data directory is in the current one and the file is in that directory")
        
```

```python
# TODO: Return the descriptor dictionary for the given name
#  Reminder that data_description is a dictionary with the "data channels" key
#   containing the list of dictionaries...
def get_descriptor(data_description, name):
    """ Search through data_description to find the dictionary with name "name"
    @param name - The name of the data channel to look for. 
    @return the dictionary that has name as the 'name' key"""

    pass
```

```python
# TODO: Write a test that, eg, get_descriptor with "Timestamp" returns the first dictionary
#  in the list in data_description.json
```

```python
# TODO get the channel data that starts at index_offset and has n_dims
def get_channel_data(all_data, index_offset, n_dims):
    """ Get the data for just one channel (e.g., right hand accelerometer)
    @param all_data - numpy array containing data from one (or more) files
    @param index_offset - the column to begin getting data from
    @param n_dims - number of dimensions for the data channel
    @return Returned array should be: number of rows in all_data X n_dims"""

    pass
```

```python
# Test 1 - the right hand accelerometer data
rh_accelerometer_descriptor = get_descriptor(data_description, "Right hand accelerometer")
rh_accelerometer_data = get_channel_data(all_data, index_offset=rh_accelerometer_descriptor["index_offset"], n_dims=rh_accelerometer_descriptor["dimensions"])

```

```python
# SELF TESTS
print(f"Shape of rhs_accelerometer_data is {rh_accelerometer_data.shape}, should be 285 X 3")
print(f"First row, first column value {rh_accelerometer_data[0, 0]:0.2f}, should be 0.70")
print(f"First row, last column value {rh_accelerometer_data[0, -1]:0.2f}, should be -0.41")
print(f"Last row, first column value {rh_accelerometer_data[-1, 0]:0.2f}, should be 0.07")
print(f"Last row, last column value {rh_accelerometer_data[-1, -1]:0.2f}, should be -0.98")

```

```python
# Tests for Left hand gyroscope
lh_gyroscope_descriptor = get_descriptor(data_description, "Left hand gyroscope")
lh_gyroscope_data = get_channel_data(all_data, index_offset=lh_gyroscope_descriptor["index_offset"], n_dims=lh_gyroscope_descriptor["dimensions"])
```

```python
# Check size and first, last element
assert(lh_gyroscope_data.shape == (all_data.shape[0], 3))
assert(np.isclose(lh_gyroscope_data[0, 0], 429.34))
assert(np.isclose(lh_gyroscope_data[-1, -1], -154.36))
```

```python
grader.check("check_slice")
```

## Compute stats: Write a function to calculate the four stats

This is a variation on what you did in lab 1; in this case, we're going to do it with two functions. The first calculates the stats and returns the dictionary (**calc_stats**) the second does the **for** loop to make one dictionary for each dimension in the data.

- Step 1 [this problem] - do the **calc_stats** function
- Step 2 [next problem] - do the loop to calculate the stats for each x,y,z channel

```python
# TODO Fill in the function that calculates the stats for the data
def calc_stats(data):
    """Calculate min, max, mean and standard deviation for the array and put in a dictionary
    @param data a numpy array
    @return a dictionary with the given keys"""

    # Use keys Min, Max, Mean, and SD
    pass
```

```python
# Test the function with known data
test_data = np.linspace(0, 1, 10)
ret_dict = calc_stats(test_data)

assert(np.isclose(ret_dict["Min"], 0.0))
assert(np.isclose(ret_dict["Max"], 1.0))
assert(np.isclose(ret_dict["Mean"], 0.5))
assert(np.isclose(ret_dict["SD"], 0.319, atol=0.01))
```

```python
grader.check("stats_channel")
```

### Now do the second half - 

This function calculates the stats for an entire channel of the data, and stores the result in a list of dictionaries

```python
def calc_stats_for_channel(data, n_dims):
    """ Calculate the stats for a channel
    @param data - an n_timestamps * n_dims size array
    @param n_dims - 1, 2, or 3 (just x; x,y; or x,y,z)
    @return A list of dictionaries. The list is the length of n_dims"""

    stats_list = []
    # TODO Copy in your for loop from the statistics problem in Lab 1
    # If you didn't do a for loop in lab 1, you need to make it a for loop now!
    # - You DO need to slice the data into the x,y,z channels
    # - You need to loop n_dims times
    # - Don't forget to return the array
    pass
```

```python
# SCRATCH CELL
# If you're having trouble, try setting n_dims to 1 and use test_data for the data input
```

```python
# Testing with known data - make a fake data set with 5 time steps and x, y, z data
#  
test_stats = np.zeros((5, 3))
# Set the x data to be ones
test_stats[:, 0] = np.ones(5)
# Set the y data to be twos
test_stats[:, 1] = np.ones(5) * 2
# Set the z data to be threes
test_stats[:, 2] = np.ones(5) * 3

# Now get the actual stats
ret_stats_array = calc_stats_for_channel(test_stats, n_dims=3)

# Check the mean result for x, y, and z - should be 1, 2, and 3 respectively
assert(ret_stats_array[0]["Mean"] == 1.0)
assert(ret_stats_array[1]["Mean"] == 2.0)
assert(ret_stats_array[2]["Mean"] == 3.0)
```

```python
# this should work
ret_stats_rh_accelerometer = calc_stats_for_channel(rh_accelerometer_data, 3)
```

```python
# As should this
res_stats_lh_gyroscope = calc_stats_for_channel(lh_gyroscope_data, 3)
```

```python
grader.check("loop_data_calc_stats")
```

## Boolean slicing to get successful versus unsuccessful statistics out

Use the functions you just wrote to get out the min and max z values for each type of hand motion.

For this problem I have written code that is *incorrect*. You know the functions themselves are correct - you just tested them. The following bits of code have something wrong with either the way the function is called OR with the way the results are gotten back.

```python
# Boolean filters for getting rows for a specific motion type. 
# Motion type should extract the last columnb in the data
#   uncomment and get the last column
# motion_type = all_data[:, ...]

# These should match the specific IDs for each motion type, eg clap_id.
# Uncomment and compare motion type to the appropriate XX_id 
#b_clap = motion_type == ...
#b_snap = motion_type == ...
#b_high_five = motion_type == ...

# Use b_clap to pick out the rows that are for claps. Send all column data for the selected rows.
#   Right hand accelerometer has 3 dimensions (x,y,z)
#   There's two errors here - one that actually will create incorrect results, one that just *happens* to work
#   correctly, although it doesn't do what the first sentance says...
ret_rh_accelerometer_clap = calc_stats_for_channel(rh_accelerometer_data[b_clap], n_dims=1)

# The minimum should be in the third (last) element in the list, the "min" key
z_min_clap = ret_rh_accelerometer_clap["Min"]
z_max_clap = ret_rh_accelerometer_clap["Max"]

# Now, do the same thing above, but for snap and high_five
z_min_snap = ...
z_max_snap = ...
z_min_high_five = ...
z_max_high_five = ...

print(f"Clap: Minimum {z_min_clap} and maximum {z_max_clap} value of right hand accelerometer z channel")
print(f"Snap: Minimum {z_min_snap} and maximum {z_max_snap} value of right hand accelerometer z channel")
print(f"High five: Minimum {z_min_high_five} and maximum {z_max_high_five} value of right hand accelerometer z channel")
```

```python
grader.check("boolean_slicing")
```

## Optional/Extra credit: print out all of the rows where the minimum z value for a clap motion was reached

See the tutorial on **np.where** (c_tutorial_where.ipynb)

TODO: Use **np.where** to pick out the row that has the minimum z value for a clap motion.

```python
# TODO Use np.where to get out the indices. You can use == OR np.isclose() here; either works. In general, use .isclose for 
#  floating point comparisons.
# Append the row number of any matches to this list
all_rows_with_min = []

# Look at JUST the z values in rh_accelerometer_data
all_indices_from_where = ...

# Pseudo code - see tutorial for exact format
# for all row in all_indices_from_where
#    if this is row is from a clap: 
#       print(f"Row: {r}, Time step: {c}")
```

```python
grader.check("optional_where")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS 2S5 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

- Do a restart then run all to make sure everything runs ok
- Save the file
- Submit just this .ipynb file through gradescope, Lab 2, functions
- You do NOT need to submit the data files - we will supply those
- Where there are given variable/file names (eg, foo = ...) DON'T change those, or the autograder will fail

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Most likely failure for this assignment is not naming the data directory and files correctly; capitalization matters for the Gradescope grader. 

Tests for boolean_slicing (points: 3)
```
>>> assert np.isclose(z_min_clap, -1.04)
```
```
>>> assert np.isclose(z_max_clap, 9.06)
```
```
>>> assert np.isclose(z_min_snap, -2.77)
```
```
>>> assert np.isclose(z_max_snap, 1.05)
```
```
>>> assert np.isclose(z_min_high_five, -0.74)
```
```
>>> assert np.isclose(z_max_high_five, -0.32)
```


Tests for check_slice (points: 2)
```
>>> assert rh_accelerometer_data.shape == (285, 3)
```
```
>>> assert np.isclose(rh_accelerometer_data[0, 0], 0.7)
```
```
>>> assert np.isclose(rh_accelerometer_data[0, -1], -0.41)
```
```
>>> assert np.isclose(rh_accelerometer_data[-1, 0], 0.07)
```
```
>>> assert np.isclose(rh_accelerometer_data[-1, -1], -0.98)
```
```
>>> assert lh_gyroscope_data.shape == (all_data.shape[0], 3)
```
```
>>> assert np.isclose(lh_gyroscope_data[0, 0], 429.34)
```
```
>>> assert np.isclose(lh_gyroscope_data[-1, -1], -154.36)
```


Tests for get_data (points: 2)
```
>>> assert get_file_info('Data/S10F10.csv') == {'csv_path': 'Data/S10F10.csv', 'motion_id': high_five_id}
```
```
>>> assert get_file_info('Data/S07S05.csv') == {'csv_path': 'Data/S07S05.csv', 'motion_id': snap_id}
```
```
>>> assert get_file_info('Data/S03C09.csv') == {'csv_path': 'Data/S03C09.csv', 'motion_id': clap_id}
```
```
>>> assert all_data[0, -1] == clap_id
```
```
>>> assert all_data[102, -1] == high_five_id
```
```
>>> assert all_data[-1, -1] == snap_id
```
```
>>> assert all_data[0, 0] == 77093
```
```
>>> assert np.isclose(all_data[0, -2], -311.4)
```
```
>>> assert all_data[101, 0] == 27084
```
```
>>> assert np.isclose(all_data[101, -2], -71.76)
```
```
>>> assert np.isclose(np.mean(all_data[:, 0]), 56082.1684, atol=0.001)
```
```
>>> assert np.isclose(np.mean(all_data[:, -1]), 1.95789, atol=0.001)
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```


Tests for loop_data_calc_stats (points: 4)
```
>>> assert len(ret_stats_array) == 3
```
```
>>> assert ret_stats_array[0]['Mean'] == 1.0
```
```
>>> assert ret_stats_array[1]['Mean'] == 2.0
```
```
>>> assert ret_stats_array[2]['Mean'] == 3.0
```


Tests for optional_where (points: 0)
```
>>> assert all_rows_with_min == [48]
```


Tests for stats_channel (points: 4)
```
>>> assert np.isclose(ret_dict['Min'], 0.0)
```
```
>>> assert np.isclose(ret_dict['Max'], 1.0)
```
```
>>> assert np.isclose(ret_dict['Mean'], 0.5)
```
```
>>> assert np.isclose(ret_dict['SD'], 0.319, atol=0.01)
```
