```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab_2_functions.ipynb")
```

# Lab 2: Statistical analysis of data using numpy

Lab slides: https://docs.google.com/presentation/d/1ykwwcQ0onMvAjUxfJmKl9tbo-rJPdB5pRwDEmpDsd-g/edit?usp=sharing

For this lab the goal is to write functions to pull out one data channel and print out statistics for failed versus successful picks. This will involve using your function from the lecture activity to get just the data you want, then another function to do the statistics (essentially the **calc_stats** function from the lecture activity). 

Written properly, you only need one function to do stats for the entire data channel, just the succesful ones, or just the unsuccessful ones. For any data channel. In the homework you'll use these functions to do this for all of the data and write it back out.

```python
# Libraries that we need to import - numpy and json (for loading the description file)
import numpy as np
import json as json
```

### Reading in data

TODO Copy over code to do the following:
- read in the numerical data and pull out the numerical data and the successful/unsuccessful data
- create a boolean index variable for the successful picks

Note: For this lab I'm going to "hard-wire" all of the numbers (number of time steps, number of data dimensions, etc) to make testing easier. When you move this code over to the homework you'll replace all of the hard-wired numbers with the variables you're calculating in homework 1.


```python
pick_data = ...
pick_channel_data = ...
pick_successful = ...
b_successful = ...

# Hard-wiring these values for the testing code
n_timesteps = 40
n_picks = 660
n_total_dims = 33

```

```python
grader.check("get_data")
```

## Doing the slice

Get the data for one of the channels. 

TODO: Copy over your function **get_channel_data** from lecture activity 2. Note: if your code did not handle doing 1, 2, or 3 dimensions, now you'll need it to. 


```python
# This reads in the json data
try:
    with open("Data/proxy_data_description.json", "r") as fp:
        pick_data_description = json.load(fp)
except FileNotFoundError:
    print(f"The file was not found; check that the data directory is in the current one and the file is in that directory")

```

```python
# TODO: Copy get_channel_data over to here
...
```

```python
# Test 1 - the wrist torque data using hard-wired values
wrist_torque_start_index = 3
n_dims_wrist_torque = 3

wrist_torque_data = get_channel_data(pick_channel_data, 
                                     n_picks=n_picks, 
                                     start_index=wrist_torque_start_index,
                                     n_time_steps=n_timesteps,
                                     n_total_dims=n_total_dims,
                                     n_dims=n_dims_wrist_torque)
```

```python
# SELF TESTS
# Feel free to copy over the asserts from lecture activity 2 to debug the above
```

```python
# Tests for Motor effort finger 1
motor_effort_f1_start_index = 14
n_dims_motor_effort_f1 = 1
motor_effort_f1_data = get_channel_data(pick_channel_data, 
                                        n_picks=n_picks, 
                                        start_index=motor_effort_f1_start_index,
                                        n_time_steps=n_timesteps,
                                        n_total_dims=n_total_dims,
                                        n_dims=n_dims_motor_effort_f1)
```

```python
# Check size and first, last element
assert(motor_effort_f1_data.shape == (n_picks, n_timesteps * n_dims_motor_effort_f1))
assert(np.isclose(motor_effort_f1_data[0, 0], 0.0))
assert(np.isclose(motor_effort_f1_data[-1, -1], 51.11))
```

```python
grader.check("check_slice")
```

## Compute stats: Write a function to calculate the four stats

This is a variation on what you did in lab 1; in this case, we're going to do it with two functions. The first calculates the stats and returns the dictionary (**calc_stats**) the second does the **for** loop to make one dictionary for each dimension in the data.

- Step 1 [this problem] - do the **calc_stats** function
- Step 2 [next problem] - do the loop to calculate the stats for each x,y,z channel

```python
def calc_stats(data):
    """Calculate min, max, mean and standard deviation for the array and put in a dictionary
    @param data a numpy array
    @return a dictionary"""

    # Use keys Min, Max, Mean, and SD
    ...
```

```python
# Test the function with known data
test_data = np.linspace(0, 1, 10)
ret_dict = calc_stats(test_data)

assert(np.isclose(ret_dict["Min"], 0.0))
assert(np.isclose(ret_dict["Max"], 1.0))
assert(np.isclose(ret_dict["Mean"], 0.5))
assert(np.isclose(ret_dict["SD"], 0.319, atol=0.01))
```

```python
grader.check("stats_channel")
```

### Now do the second half - 

This function calculates the stats for an entire channel of the data, and stores the result in a list of dictionaries

```python
def calc_stats_for_channel(data, n_dims):
    """ Calculate the stats for a channel
    @param data - an n_picks X n_timesteps * n_dims size rray
    @param n_dims - 1, 2, or 3 (just x, or x,y, and z)
    @return A list of dictionaries. The list is the lenght of n_dims"""

    stats_list = []
    # TODO Copy in your for loop from the statistics problem in Lab 1
    # - You do NOT need to get the data out from pick data - it's done for you
    # - You DO need to slice the data into the x,y,z channels
    # - You need to loop n_dims times
    # - Don't forget to return the array
    ...

```

```python
# SCRATCH CELL
# If you're having trouble, try setting n_dims to 1 and use test_data for the data input
```

```python
# Testing with known data - make a fake data set with 5 picks, 4 time steps, and x, y, z data
#  
test_stats = np.zeros((5, 4 * 3))
# Set the x data to be ones
test_stats[:, 0::3] = np.ones((5, 4))
# Set the y data to be twos
test_stats[:, 1::3] = np.ones((5, 4)) * 2
# Set the z data to be threes
test_stats[:, 2::3] = np.ones((5, 4)) * 3

# Now get the actual stats
ret_stats_array = calc_stats_for_channel(test_stats, n_dims=3)

# Check the mean result for x, y, and z - should be 1, 2, and 3 respectively
assert(ret_stats_array[0]["Mean"] == 1.0)
assert(ret_stats_array[1]["Mean"] == 2.0)
assert(ret_stats_array[2]["Mean"] == 3.0)
```

```python
# this should work - you can check the result against the values in Data/HW1_check_results.json
ret_stats_wrist_torque = calc_stats_for_channel(wrist_torque_data, n_dims_wrist_torque)
```

```python
# As should this
res_stats_motor_effort_f1 = calc_stats_for_channel(motor_effort_f1_data, n_dims_motor_effort_f1)
```

```python
grader.check("loop_data_calc_stats")
```

## Boolean slicing to get successful versus unsuccessful statistics out

Use the functions you just wrote - plus the boolean index you made at the beginning - to get out the min and max z values for successful versus unsuccessful picks. 

For this problem I have written code that is *incorrect*. You know the functions themselves are correct - you just tested them. The following bits of code have something wrong with either the way the function is called OR with the way the results are gotten back.


```python
# Use b_successful to pick out the rows that are successful. Send all column data for the selected rows.
#   Wrist torque data has 3 dimensions (x,y,z)
#   There's two errors here - one that actually will create incorrect results, one that just *happens* to work
#   correctly, although it doesn't do what the first sentence says...
ret_wrist_torque_successful = calc_stats_for_channel(wrist_torque_data[b_successful], n_dims=1)

# The minimum should be in the third (last) element in the list, the "min" key
z_min_successful = ret_wrist_torque_successful["Min"]
z_max_successful = ret_wrist_torque_successful["Max"]

# Use b_successful NOT true to pick out the picks that are successful.
#  This generates a weird error - it's because not does not work over a numpy array. Try b_successful == False instead
ret_wrist_torque_unsuccessful = calc_stats_for_channel(wrist_torque_data[not b_successful], n_dims=1)

# The minimum should be in the third (last) element in the list, the "min" key
z_min_unsuccessful = ret_wrist_torque_unsuccessful[3]["Min"]
# Why copying and pasting and changing variable names can cause problems...
z_max_unsuccessful = ret_wrist_torque_successful["Max"]



print(f"Successful: Minimum {ret_wrist_torque_successful} and maximum {z_max_successful} value of wrist torque z channel")
print(f"Unsuccessful: Minimum {z_min_unsuccessful} and maximum {z_max_unsuccessful} value of wrist torque z channel")
```

```python
grader.check("boolean_slicing")
```

## Optional/Extra credit: print out all of the indices where the maximum value for the successful pick was reached

See the tutorial on **np.where**

TODO: Use **np.where** to pick out the row, col pair that has the maximum z value of the successful pick

Partial credit for picking out any row, col that has the maximum z value in **wrist_torque_data**, full extra credit for only printing out the row, col indices of successful picks with that *z* value.

```python
# Use np.where to get out the indices. You can use == OR np.isclose() here; either works. In general, use .isclose for 
#  floating point comparisons.
# Append the row number of any matches to this list
all_rows_with_max = []


# Look at JUST the z values in wrist_torque_data
all_indices_from_where = ...
# Pseudo code - see tutorial for exact format
# for all row, column in all_indices_from_where
#.   if this is row is successful 
#.      print(f"Row: {r}, Time step: {c}")

```

```python
grader.check("optional_where")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS TCS3 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

- Do a restart then run all to make sure everything runs ok
- Save the file
- Submit just this .ipynb file through gradescope, Lab 2, functions
- You do NOT need to submit the data files - we will supply those
- Where there are given variable/file names (eg, foo = ...) DON'T change those, or the autograder will fail

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Most likely failure for this assignment is not naming the data directory and files correctly; capitalization matters for the Gradescope grader. 

Tests for boolean_slicing (points: 3)
```
>>> assert(np.isclose(z_min_successful, -0.293665094))
```
```
>>> assert(np.isclose(z_max_successful, 0.340460618))
```
```
>>> assert(np.isclose(z_min_unsuccessful, -0.62552044))
```
```
>>> assert(np.isclose(z_max_unsuccessful, 0.326538637))
```


Tests for check_slice (points: 2)
```
>>> assert(motor_effort_f1_data.shape == (n_picks, n_timesteps * n_dims_motor_effort_f1))
```
```
>>> assert(wrist_torque_data.shape == (n_picks, n_timesteps * n_dims_wrist_torque))
```
```
>>> assert(np.isclose(motor_effort_f1_data[0, 0], 0.0))
```
```
>>> assert(np.isclose(motor_effort_f1_data[-1, -1], 51.11))
```


Tests for get_data (points: 2)
```
>>> assert(pick_channel_data.shape[0] == 660)
```
```
>>> assert(pick_channel_data.shape[1] == n_timesteps * n_total_dims)
```
```
>>> assert(np.count_nonzero(b_successful) == 355)
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for loop_data_calc_stats (points: 4)
```
>>> assert(len(ret_stats_array) == 3)
```
```
>>> assert(ret_stats_array[0]["Mean"] == 1.0)
```
```
>>> assert(ret_stats_array[1]["Mean"] == 2.0)
```
```
>>> assert(ret_stats_array[2]["Mean"] == 3.0)
```


Tests for optional_where (points: 0)
```
>>> assert(all_rows_with_max[0] == 82)
```


Tests for stats_channel (points: 4)
```
>>> assert(np.isclose(ret_dict["Min"], 0.0))
```
```
>>> assert(np.isclose(ret_dict["Max"], 1.0))
```
```
>>> assert(np.isclose(ret_dict["Mean"], 0.5))
```
```
>>> assert(np.isclose(ret_dict["SD"], 0.319, atol=0.01))
```
