```python
# Initialize Otter
import otter
grader = otter.Notebook("HW_4_pinball.ipynb")
```

# HWK 4 pinball simulation

The full pinball simulation: slides https://docs.google.com/presentation/d/1fNTOc_4YXfBkQ_ko9X8eQqAgK3Opnn9c02i72sVDqKU/edit?usp=sharing

Full pinball 
 - position and velocity with Euler simulation (lecture activity, compute_next_step)
 - stop when ball passes below y=0 (lab, simulate_pinball version 1)
 - "Bouncing" off of the top and side walls (lab, outside_xx_wall, handle bounce, simulate_pinball version 1)
 - "Bouncing" off of bumpers (homework)

Week 1: 
- Circular "bumpers" OR a wall at an angle (circular bumpers add extra credit)
- [optional] Some bumpers will impart an acceleration to the pinball
- [optional] Add damping

Week 2: Some simple calculations
- Length of path pinball travels
- Integral of vx and vy


```python
# The usual imports
import numpy as np
import matplotlib.pyplot as plt
```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

<!-- BEGIN QUESTION -->

## Circular bumper or angled wall: How to implement?

TODO: Write code to handle a circular bumper OR an angled wall

- decide which functions/what the format of the functions should be written to handle a bumper or angled wall
- decide how you will represent a bumper or an angled wall

Note: You'll add tests and simulation and plotting in the following questions

This is a manually graded question

**TODO:** 

Use this cell to describe how you are implementing your bumper or angled wall routines
- How are you representing the bumper/angled wall?
- What functions are you creating and why? What are the function inputs and outputs?


```python
# Imports from week 5
# This bit of code tells Jupyter notebooks to also look in the week_6_optimization folder for files
import os
import sys

# If you want them...
sys.path.insert(0, os.path.abspath('../Week_5_matrices'))
```

```python
# TODO: Write additional routines here
# You can write them here or put them in pinball_routines.py and import them here
# Add one cell for each function
...
```

```python
# TODO: Add tests here to test your functions
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Simulation routine

TODO: Copy over the simulation routine from your lab

- Change the function's parameters so that they take in the bumpers/angled wall as well
- Optional: If you implement acceleration or drag then describe it here 

This question is manually graded. Describe what you did, why, and how

TODO: Describe your solution here

```python
# Pinball_routines imports
...
```

```python
# Simulate pinball routine
#. TODO Copy your lab solution here
#       Change the parameters to support your circular bumpers or angled wall
...
```

```python
# Time step
delta_t = 0.01

# Test 1 - do you stop when you hit the floor?
starting_state = np.zeros([3, 2])  # location, velocity, acceleration
starting_state[0, :] = [0, 0] # Start at zero, zero
starting_state[1, :] = [2.4, 8.5]
starting_state[2, :] = [0.0, acceleration_due_to_gravity()]

# Define your bumper/angled wall here
# Default is top wall at y=5, left/right walls at +- 3.0
...
```

```python
# Test the simulate_pinball function, with at least setup
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Plot bumpers or angled wall, walls, and path

TODO: Start with the plotting code from the lab
 - add plotting the bumpers or angled wall
 - pass the data structure you use to define the bumpers and walls to the plot routine
 
TODO: Generate at least 2 plots showing the pinball doing different behaviors (must include reflection off of wall, bumper[s]/angled wall)

This question is manually graded

```python
# TODO: Create a function to plot the pinball track, walls, and bumpers
#  Start with the lab plot
...
```

```python
# TODO: set up different starting positions and/or pinball setups and plot them
...
```

```python
# Manual grade
print("This is a manually-graded question; there is no grader.check() function. See rubric and slides for more information on expected output.")
```

<!-- END QUESTION -->

## Week 2 - Use integration to calculate length of path

TODO: Calculate the length of the pinball path using euclidian distances

```python
def length_of_pinball_path(poses):
    """ Calculate the length of the pinball path using euclidian distances
    @param poses - the x,y poses
    @return length"""
    # TODO: Sum up the distance between all of the pose points
    # Reminder: Distance is sqrt(x^2 + y^2)
    return 0.0
```

```python
# TODO For these tests use this starting state
#  - do NOT place a bumper/wall in the path - it should just go up and down
#  - Use the given starting position and velocity
#  - Use a delta t of 0.01
#  - Do NOT do damping
starting_state = np.zeros([3, 2])  # location, velocity, acceleration
starting_state[0, :] = [0, 0] # Start at zero, zero
# Velocity - mostly up with a bit of x noise
starting_state[1, :] = [0.5, 3.5]
starting_state[2, :] = [0.0, acceleration_due_to_gravity()]

# TODO: call your simulate routine with  delta_t = 0.01 and the above starting state

ret_poses = ...

```

```python
# poses check - 
len_path = (length_of_pinball_path(ret_poses))
print(f"Distance traveled {len_path}, expected 1.386298")
```

```python
grader.check("integrate_path")
```

## Integrate x velocity and y velocity

TODO: Calculate the total relative distance the pinball traveled in x and y by integrating the velocity in x and y
- TODO: Change your simulation code so it returns velocities in addition to poses
- Call trapz with the vx values, using a fixed dx (delta_t) (this is the integral)

```python
# EXAMPLE CODE
# To return the velocities in addition to the poses you want to
#   1) replicate your code for storing the positions, only this time, store the velocities
#   2) return the velocities as well as the positions
#   3) extract the velocities out

# There's two ways to do this, either as a 4xn array (get the velocities out afterwards, or two 2xn arrays)
n = 10
ret_vals = np.zeros((4, n))  # storing x, y, vx, vy in each column
x = 0
y = 0
vx = 0.1
vy = 0.2
state = np.zeros((2, 2))
state[0, :] = [x, y]    # First row is position
state[1, :] = [vx, vy]  # Second row is velocity
for i in range(0, n):
    ret_vals[0:2, i] = state[0, :]   # put the current state position in the first two rows
    ret_vals[2:4, i] = state[1, :]   # put the current state position in the 3rd and fourth rows

    # Increment state
    state[0, :] = state[0, :] + state[1, :] * 0.1
# return ret_vals
    
# Calling function
poses = ret_vals[0:2, :]
vels = ret_vals[2:4, :]

# OR
ret_poses = np.zeros((2, n))  # storing x, y in each column
ret_vels = np.zeros((2, n))  # storing vx, vy in each column
x = 0
y = 0
vx = 0.1
vy = 0.2
state = np.zeros((2, 2))
state[0, :] = [x, y]    # First row is position
state[1, :] = [vx, vy]  # Second row is velocity
for i in range(0, n):
    ret_poses[0:2, i] = state[0, :]   # put the current state position in the first two rows
    ret_vels[0:2, i] = state[1, :]   # put the current state position in the 3rd and fourth rows

    # Increment state
    state[0, :] = state[0, :] + state[1, :] * 0.1

# return ret_poses, ret_vels
# poses, vels = func()

```

```python
def integral_of_path_velocity(velocities, delta_t):
    """ Calculate the x and y distance traveled of the pinball path by integrating the velocities
    @param velocities - the vx,vy velocities at each time step
    @param delta_t - the sampling used to generate the velocities
    @return integral of vx, integral of vy"""
    # TODO: Calculate the integral of vx and vy
    # Note: You'll want trapz because these are samples of the velocity function
    return integral_x, integral_y
```

```python
# This needs to be a 2xn array of the velocities at each time step
ret_velocities = ...

# Note - make sure you're using the starting state given in the path integration
integral_x, integral_y = integral_of_path_velocity(ret_velocities, delta_t)
print(f"Integral of x velocity {integral_x}, expected 0.365")
print(f"Integral of y velocity {integral_y}, expected -0.056209")
```

```python
grader.check("integrate_velocity")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS I25 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

Double check your plots. Don't forget **pinball_routines.py**. We will include **matrix_routines.py**; but it won't hurt if you include it as well.

- Submit this .ipynb file  and pinball_routines.py to homework 4 (pinball)

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Congratulations, and enjoy being done with this class for the term!

Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```


Tests for integrate_path (points: 4)
```
>>> assert np.isclose(len_path, 1.386298, atol=0.001)
```


Tests for integrate_velocity (points: 4)
```
>>> assert np.isclose(integral_x, 0.365, atol=0.001)
```
```
>>> assert np.isclose(integral_y, -0.056209, atol=0.001)
```
