```python
# Initialize Otter
import otter
grader = otter.Notebook("HW_4_skier.ipynb")
```

# HWK 4 skier simulation

The full skier simulation: slides https://docs.google.com/presentation/d/1fNTOc_4YXfBkQ_ko9X8eQqAgK3Opnn9c02i72sVDqKU/edit?usp=sharing

Full skier implementation 
 - position and velocity with Euler simulation (lecture activity/lab)
 - represent ski slope as a connected sequence of linear and quadratic polynomials
 - calculate acceleration from slope of polynomial
 - go airborne
 - stop simulation when dropped off ski slope
 - Optional
 -  Crash into the ground
 -  Bounce when hitting the ground
 -  Friction/drag
 -  Other slopes

Week 1: 
- simulation with connected pieces of ground

Week 2: 
- length of path skier travels
- [optional additions]


There are


```python
# The usual imports
import numpy as np
import matplotlib.pyplot as plt
```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

<!-- BEGIN QUESTION -->

## Documentation

TODO: use this cell to document: data structures, functions, test cases

Note: There are two options for this assignment; you can use the classes/tests supplied in **ski_jump.py** or you can use your own implementation.

Note 2: As in homework 3, you can fully test and debug **ski_jump.py** by executing it.

Although **ski_jump.py** supplies the class structure/function decomposition, it does NOT include how you want to store data. So you should still write that down here.

```python
# TODO If you're using ski_jump.py keep this uncommented; if you're doing your own thing, then replace this with
#   either an import or your functions
# Note: Automating testing in the next section
# Note 2: Please go look at how the testing sets up the ground construction before starting implementation
from ski_jump import Skier, Ground
```

<!-- END QUESTION -->

## Testing with a ground section consisting of one piece of linear ground

For all of the tests the ground is defined as follows:
- a tuple/list of the x breaks between the sections
- a tuple/list of the constraints for the polynomial sections
- if linear then just the slope
- if quadratic then a tuple with the y value at the end (x is given by x breaks) and the slope at the start

This is almost what you did in lab, except the x start and end values are coming from the list of breaks and you'll need to evaluate the previous polynomial to get y_start

```python
# TODO If you are using ski_jump.py then this function is done for you; if you are NOT, then comment out what's in the function
#  and return your own data structure
#  See the prompt at the start of this question and the cell below for what x_breaks and slopes are
def build_ground(x_breaks, y_start, slopes):
    """Create a ground structure
    @param x_breaks - a tuple of monotonically increasing x values that say when to switch polynomials
    @param y_start - the starting y value
    @param slopes - a tuple that has the additional data needed to build the polynomials
    @return your data structure for the ground"""
    return Ground(x_breaks=x_breaks, y_start=y_start, slopes_or_curves=slopes)
```

```python
# TODO If you are using ski_jump.py then this function is done for you; if you are NOT, then comment out what's in the function
#  and return your own data structure
#  Initializes the "skier"
def initialize_skier(x, y, vx, vy):
    """ Initialize the data structure for storing the skier's position/velocity etc
    @param x, y - starting position
    @param vx, vy - starting velocity
    @return your skier data structure"""
    return Skier((x, y), (vx, vy))
```

```python
# TODO If you are using ski_jump.py then this function is done for you; if you are NOT, then comment out what's in the function
#  and run the simulation and return the poses
def simulate_skier(skier, ground, delta_t):
    """ Run the simulation; note STOP at y = 0
    @param skier - the data structure returned from initialize_skier
    @param ground - the data structure returned from build_ground
    @param delta_t - simulation time step
    @param poses - a 2xn numpy array with the locations of the skier"""
    skier.simulate(ground=ground, delta_t=delta_t, eps=0.01)
    return skier.saved_poses()
```

```python
# This is the "ground" used for these tests - a linear polynomial

# Start at x = 0.0, end at x = 10.0
x_breaks_simple = (0.0, 10.0)
# Starting height
y_start_simple = 10.0
# Slope of the line is -0.5
slope_ground_simple = -0.5

ground_simple = build_ground(x_breaks=x_breaks_simple, y_start=y_start_simple, slopes=((slope_ground_simple,), ))
```

```python
# First skier test
x_y_a = (0, 10)
vx_vy_a = (2.0, 2.0 * 4.0)
skier_airborne = initialize_skier(x_y_a[0], x_y_a[1], vx_vy_a[0], vx_vy_a[1])

delta_t = 0.01
poses = simulate_skier(skier=skier_airborne, ground=ground_simple, delta_t=delta_t)

# TODO Test that the simulation starts with the correct pose and ends when crosses below the y = 0 line
```

```python
grader.check("simple_tests")
```

## Make the ground from several pieces
TODO: Once your ground data structure is completed and build_geometry above is filled in, then test

```python
# TODO If you are using ski_jump.py then this function is done for you; if you are NOT, then comment out what's in the function
#  and return the height at the given x value
def height(ground, x):
    """ Return the height at value x
    @param ground - the data structure returned from build_ground
    @return y value (height)"""
    return ground.height(x)

```

```python
x_breaks_ground = (0.0, 2.0, 4.5, 10.0)
y_start_ground = 10.0

# A linear piece with slope -2.0, a flat piece, and a quadratic that starts with a slope of -4.0 and ends at y = 3.5
slopes_ground = ((-2.0,), (0.0,), (3.5, -4.0))

ground = build_ground(x_breaks=x_breaks_ground, y_start=y_start_ground, slopes=slopes_ground)

# TODO Write tests for:
#  Make sure starts at 0.0, 10.0, ends at 10.0, 3.5 and connected

```

```python
grader.check("ground")
```

<!-- BEGIN QUESTION -->

## TODO Create at least four different ski scenarios with 2 different ski slopes

You can change the ski scenario by changing the starting location/velocity of the skier

You must make plot code that shows:

- The ground/sopes
- The skier's path along the slope
- If the skier is airborne than that should be in a different color

Note that you can use the plot code in the lab as a starting point

_Type your answer here, replacing this text._

```python
# FILL IN HERE
```

<!-- END QUESTION -->

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS 2U5 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

Double check your plots. Don't forget **ski_jump.py** if you used it (or take out the import if you are not using). 

- Submit this .ipynb file and any additional .py files to homework 4 (skier)

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Congratulations, and enjoy being done with this class for the term!

Tests for ground (points: 5)
```
>>> assert np.isclose(height(ground=ground, x=x_breaks_ground[0]), y_start_ground, atol=0.01)
```
```
>>> assert np.isclose(height(ground=ground, x=x_breaks_ground[-1]), slopes_ground[2][0], atol=0.01)
```
```
>>> assert np.isclose(height(ground=ground, x=x_breaks_ground[1] - 0.0001), height(ground=ground, x=x_breaks_ground[1] + 0.0001), atol=0.01)
```
```
>>> assert np.isclose(height(ground=ground, x=x_breaks_ground[2] - 0.0001), height(ground=ground, x=x_breaks_ground[2] + 0.0001), atol=0.01)
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```


Tests for simple_tests (points: 5)
```
>>> assert np.isclose(poses[0, 0], x_breaks_simple[0])
```
```
>>> assert np.isclose(poses[1, 0], y_start_simple)
```
```
>>> assert poses[1, -1] < 0.0
```
```
>>> assert poses[0, -1] > x_breaks_simple[-1]
```
