```python
# Initialize Otter
import otter
grader = otter.Notebook("LabHW_Extra_credit_gripper.ipynb")
```

# Using matrices to position parts of a gripper

The goal of this part of the assignment is to use matrices to create a gripper. Note: You can use this code in your homework to add the gripper

Slides: https://docs.google.com/presentation/d/17Tc1uTemxQDT_K2O5bi355wfRjpExtziVCX5Zz5HEqI/edit?usp=sharing

Differences from arm component:
- This time all three parts (palm and fingers) are stored in the class
- Creating the shape matrices and the points will happen in the init method, rather than in a separate set_size call
- The finger rotations are not around the origin, but around the attachment to the palm
- The palm rotation is applied to the palm AND the fingers



## Specific instructions

All of the code you will edit is in gripper.py

GUIDE step 1: Fill in the init function (making the points and the shape matrices). 

GUIDE step 2: Rotation matrices: Set the rotation matrices

GUIDE step 3: Edit your RobotArm class to add the gripper (add_gripper) and set the rotations. Reminder that the final arm link pose matrix should be applied to the gripper as well.

Don't forget to fill in the get_xxx methods

```python
# The usual imports
import numpy as np
import matplotlib.pyplot as plt

# Print all matrices/floating point data with 4 decimal places
np.set_printoptions(formatter={'float': lambda x: "{0:0.4f}".format(x)})

```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

```python
# The matrix routines class
import matrix_routines as mt

# This is the class you'll edit for the extra credt. You can edit (and check) what you're writing in the .py file by executing gripper.py OR 
# edit gripper.py then run the grader checks here
# Note: Do a restart runall in order to load any changes you made to gripper.py
from gripper import Gripper
```

## Step 1: Gripper (fingers plus palm)

The "gripper" is put together from three pieces - the palm, and two fingers. Edit the __init__ and get_xxx methods in gripper.py to create the points and shape matrices. 

GUIDE: Edit the **init** method and **get_** methods in **Gripper**

```python
# The sizes for all of the components
palm_size = 0.1
finger_length = 0.075
```

```python
# SCRATCH CELL - if need be, you can write the code here then check it against the answers given in the next cell
```

```python
# Create the gripper - should set the shape matrices and the points
gripper = Gripper(palm_size=palm_size, finger_length=finger_length)

# Check the shape matrices
mat_palm_check = np.array([[0.005, 0.0, 0.0], [0.0, 0.05, 0.0], [0.0, 0.0, 1.0]])
mat_finger_top_check = np.array([[0.0, 0.0375, 0.0375], [-0.0125, 0.0, 0.05], [0.0, 0.0, 1.0]])
mat_finger_bot_check = np.array([[0.0, 0.0375, 0.0375], [-0.0125, 0.0, -0.05], [0.0, 0.0, 1.0]])

# Feel free to add asserts here
```

```python
# Plot the result - see slides
# Draw the gripper in the plot window - see slides for correct answer

fig, axs = plt.subplots(1, 1, figsize=(2, 2))

mt.plot_axes_and_big_box(axs, box_size=0.25)
gripper.plot(axs, b_do_pose_matrix=False)

```

```python
grader.check("gripper_matrices")
```

## Step 2: Rotate each component

GUIDES: Edit **set_rotations** to rotate the palm and fingers. See slides for some common incorrect matrices. Don't forget to fill in the **get** pose methods.


```python
# SCRATCH CELL
# Use this cell to call the rotation method on one component (say the link1 component) and then check the result.
#  This checks JUST the palm rotation, setting the fingers to zero

rot_palm_amt = -np.pi/4
rot_finger_amnt = 0.0
gripper.set_rotations(palm_rot_amt=rot_palm_amt, finger_rot_amt=rot_finger_amnt)

# If the fingers aren't rotating then the finger rotation matrices are just the palm ones
mat_pose_check_palm = mt.make_rotation_matrix(rot_palm_amt)

assert np.all(np.isclose(gripper.get_palm_pose_matrix(), mat_pose_check_palm, atol=0.001))
assert np.all(np.isclose(gripper.get_finger_pose_matrix(True), mat_pose_check_palm, atol=0.001))
assert np.all(np.isclose(gripper.get_finger_pose_matrix(False), mat_pose_check_palm, atol=0.001))

# Now rotate the fingers - make sure this is the last set_rotations called
rot_finger_amnt2 = np.pi/8
gripper.set_rotations(palm_rot_amt=rot_palm_amt, finger_rot_amt=rot_finger_amnt2)
mat_check_f1 = np.array([[0.9239, 0.383, 0.01622], [-0.3827, 0.9239, -0.01083], [0, 0, 1]])
mat_check_f2 = np.array([[0.383, 0.9239, 0.01083], [-0.9239, 0.3827, -0.0162], [0, 0, 1]])

# Insert your own checks here
```

```python
fig, axs = plt.subplots(1, 3, figsize=(8, 4))

# STEP 2 check: The left plot is the gripper in it's initial configuration; the right one has the palm rotated and the fingers
gripper.set_rotations(palm_rot_amt=0.0, finger_rot_amt=0.0)
mt.plot_axes_and_big_box(axs[0], box_size=0.2)
gripper.plot(axs[0], b_do_pose_matrix=False)
axs[0].set_title(gripper.name)

# Do rotation of just the palm
gripper.set_rotations(palm_rot_amt=-np.pi/4, finger_rot_amt=0.0)
mt.plot_axes_and_big_box(axs[1], box_size=0.2)
gripper.plot(axs[1], b_do_pose_matrix=True)
axs[1].set_title(gripper.name)

# Do rotation of fingers and palm
gripper.set_rotations(palm_rot_amt=-np.pi/4, finger_rot_amt=np.pi/8)
mt.plot_axes_and_big_box(axs[2], box_size=0.2)
gripper.plot(axs[2], b_do_pose_matrix=True)
axs[2].set_title(gripper.name)

fig.tight_layout()
plt.show()

```

```python
grader.check("rotate_pose_matrices")
```

## Draw the gripper at the end of the arm.
Assumes you have completed the first part of the robot arm homework (building the arm).

GUIDE: Edit **set_link_angles** to 1) set the angles of the palms and fingers 2) pass in the end of the arm pose into the gripper's **set_rotation** matrix so that the whole palm/gripper is in the correct pose.

Don't forget to fill in **add_gripper** and **get_gripper** in **RobotArm**


```python
# Import the robot arm
from robot_arm import RobotArm2D
```

```python
# Set up the arm as in the homework, but this time add in the gripper
# These are example inputs for the create_arm_geometry function. 
base_size_param = (0.5, 1.0)   # Height, width
link_sizes_param = [(0.5, 0.25), (0.3, 0.1), (0.2, 0.05)]

# Now add the gripper
palm_size = 0.1
finger_length = 0.075

# Create the robot arm as a class, then add the gripper.
robot_arm = RobotArm2D(base_size_param, link_sizes_param)
robot_arm.add_gripper(palm_size, finger_length)
```

```python
# Check the gripper get method and the shape matrices
# Correct shape matrices
mat_palm_check = np.array([[0.005, 0.0, 0.0], [0.0, 0.05, 0.0], [0.0, 0.0, 1.0]])
mat_finger_top_check = np.array([[0.0, 0.0375, 0.0375], [-0.0125, 0.0, 0.05], [0.0, 0.0, 1.0]])
mat_finger_bot_check = np.array([[0.0, 0.0375, 0.0375], [-0.0125, 0.0, -0.05], [0.0, 0.0, 1.0]])

# Example check 
assert np.all(np.isclose(robot_arm.get_gripper().get_palm_shape_matrix(), mat_palm_check))

```

```python
# Now set the arm angles and the gripper angles
angles_check = [np.pi/2, -np.pi/4, -3.0 * np.pi/4, [np.pi/3.0, np.pi/4.0]]
robot_arm.set_link_angles(link_angles=angles_check[0:-1], palm_angle=angles_check[-1][0], finger_angle=angles_check[-1][1])
```

```python
fig, axs = plt.subplots(1, 2, figsize=(6, 3))
robot_arm.set_link_angles(link_angles=angles_check[0:-1], palm_angle=angles_check[-1][0], finger_angle=angles_check[-1][1])
robot_arm.plot(axs=axs[0])
axs[0].set_title("No fingers")

robot_arm.set_link_angles(link_angles=angles_check[0:-1], palm_angle=angles_check[-1][0], finger_angle=angles_check[-1][1])
robot_arm.plot(axs=axs[1])
axs[1].set_title("Fingers")

fig.tight_layout()
plt.show()

```

```python
mat_check_wrist = np.array([[ 0.5, -0.8660,  -0.5121], \
                            [ 0.8660,  0.5,   0.7121], \
                            [ 0.0,  0.0,  1.0]])
print(robot_arm.get_gripper().get_palm_pose_matrix())
assert np.all(np.isclose(robot_arm.get_gripper().get_palm_pose_matrix(), mat_check_wrist, atol=0.01))
```

```python
mat_check_f1 = np.array([[-0.25881, -0.965926,  -0.507137], \
                         [ 0.965926, -0.258819,   0.750073], \
                         [ 0.0,  0.0,  1.0]])
print(robot_arm.get_gripper().get_finger_pose_matrix(is_top=True))
assert np.all(np.isclose(robot_arm.get_gripper().get_finger_pose_matrix(is_top=True), mat_check_f1, atol=0.01))
```

```python
mat_check_f2 = np.array([[0.965926, -0.258819,  -0.481772], \
                         [0.258819, 0.965926,   0.735428], \
                         [ 0.0,  0.0,  1.0]])
print(robot_arm.get_gripper().get_finger_pose_matrix(is_top=False))
assert np.all(np.isclose(robot_arm.get_gripper().get_finger_pose_matrix(is_top=False), mat_check_f2, atol=0.01))
```

```python
grader.check("arm_with_gripper")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS SP26 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

Read me!!!   **Read me!!!**

- Submit this .ipynb file **AND** arm_component.py. If you don't include arm_component.py Gradescope cannot magically reach out to your computer and find your arm_component.py file.
- We will supply matrix_routines.py for you (it won't break anything if you do include it)

If the Gradescope autograder fails, please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Lots of people forget arm_component.py. 

Tests for arm_with_gripper (points: 2)
```
>>> assert np.all(np.isclose(robot_arm.get_gripper().get_palm_shape_matrix(), mat_palm_check))
```
```
>>> assert np.all(np.isclose(robot_arm.get_gripper().get_finger_pose_matrix(is_top=True), mat_check_f1, atol=0.01))
```
```
>>> assert np.all(np.isclose(robot_arm.get_gripper().get_finger_pose_matrix(is_top=False), mat_check_f2, atol=0.01))
```


Tests for gripper_matrices (points: 2)
```
>>> assert np.all(np.isclose(gripper.get_palm_shape_matrix(), mat_palm_check))
```
```
>>> assert np.all(np.isclose(gripper.get_finger_shape_matrix(True), mat_finger_top_check))
```
```
>>> assert np.all(np.isclose(gripper.get_finger_shape_matrix(False), mat_finger_bot_check))
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```


Tests for rotate_pose_matrices (points: 2)
```
>>> assert np.all(np.isclose(gripper.get_palm_pose_matrix(), mat_pose_check_palm, atol=0.001))
```
```
>>> assert np.all(np.isclose(gripper.get_finger_pose_matrix(True), mat_check_f1, atol=0.001))
```
```
>>> assert np.all(np.isclose(gripper.get_finger_pose_matrix(False), mat_check_f2, atol=0.001))
```
