```python
# Initialize Otter
import otter
grader = otter.Notebook("Intro_assignment.ipynb")
```

# Instructions

This assignment serves several purposes:
- Python refresher (tutorials a-g) in case it's been awhile
- An introduction to the autograder/submitting to gradescope
- An introduction to Jupyter Notebooks and the format of the assignments

For each assignment the top of the Jupyter Notebook will have additional links, either to recommended tutorials and/or slides that further describe the assignment. These slides will have a mix of helpful hints, guidelines, and what the expected outcomes should be (particularly for plots). There is a good chance that any question you have will have an answer in those slides, so check them first. 

## Jupyter notebook structure

For each question there's an autograder check; almost always there will be a cell that has the exact same code as the autograder check so you can see what it is. If you pass the autograder checks here, then (probably) you will pass the Gradescope autograder. **You should always check the result of the Gradescope autograder and make sure you got full credit.** If you don't, the reason you didn't is (probably) here: https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing. There are additional instructions at the end of each Jupyter Notebook that you should read before handing in.

There's usually a lot of supplied code/instructions. Anything you're suppose to change is labeled with "GUIDE". When in doubt, look to see what the autograder is checking for.


```python
# Doing the imports for you
# Note, if this fails try running a_start_here_libraries.ipynb first
#  If you've done that, then make sure that the kernel you've selected for both is the same
#  If that doesn't work, find a TA.
import numpy as np
import json as json
```

## Tutorials A and B: strings

These are minor variations on the practice problems in *c_practice_strings.ipynb*. If you can't answer them, go back to the **a** and **b** tutorials.

```python
# GUIDE: Create a string from the following variables

name_of_file = "data"
file_number = 1
file_ending = "csv"


# GUIDE: replace ... with string creation
data_file_name = ...

print(f"{data_file_name}, {da}")
# This is what the autograder will check
assert data_file_name == "data1.csv"
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[4], line 12
      9 data_file_name = ...
     11 # This is what the autograder will check
---> 12 assert data_file_name == "data1.csv"

AssertionError:
```

```python
file_name = "fname.txt"
which_one = 1

# GUIDE: Change this so that new_file_name is "fname1.csv" by EDITING file_name and using which_one - don't just type the new name in.
#  I.e., it should still work even if we change file_name to be something else and which_one to be a different number
# HINT: take fname.txt apart by using split, with . as the split
#  Build the new name by taking the first part of the split, converting which_one to a string, then gluing it all together with .csv

# Feel free to write more code and do it in pieces - i.e., you might start with 
#     fname_split = file_name.split(".")
# and print out fname_split... OR try clicking on the Variables button above, next to Clear All Outputs

# HINT 2: To get the first element out of a list, remember python starts at 0. So try printing fname_split[0]


# GUIDE: Build new_file_name. You may write as many lines of code as you want - you do not have to do this in one line.
#   Do NOT just put the desired string in. Actually build it from the variables fname and which_one
new_file_name = file_name

# This is what the autograder will check
assert new_file_name == "fname1.csv"
```

```python
grader.check("Strings")
```

Output:
```
Strings results:
    Strings - 1 result:
        ❌ Test case failed
        Trying:
            assert data_file_name == 'data1.csv'
        Expecting nothing
        **********************************************************************
        Line 1, in Strings 0
        Failed example:
            assert data_file_name == 'data1.csv'
        Exception raised:
            Traceback (most recent call last):
              File "/opt/anaconda3/lib/python3.12/doctest.py", line 1368, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest Strings 0[0]>", line 1, in <module>
                assert data_file_name == 'data1.csv'
                       ^^^^^^^^^^^^^^^^^^^^^^^^^^^^^
            AssertionError

    Strings - 2 result:
        ❌ Test case failed
        Trying:
            assert new_file_name == 'fname1.csv'
        Expecting nothing
        **********************************************************************
        Line 1, in 
```

## Tutorial D: Lists

These are minor variations on the practice problems in *e_practice_lists.ipynb*. If you can't answer them, go back to the **d** tutorial. 

If you have never used Python lists before, there are a couple on-line tutorials recommended at the top of *d_tutorial_list.ipynb*. 

```python
# GUIDE Create a list of filenames, fname1.csv, fname2.csv, ... fname3.csv
# You must use the variables listed below to set the name of the file, the number of file names to create, and the file ending

file_name = "fname"
file_ending = ".csv"

n_to_do = 3

my_list_of_file_names = []

# Do NOT re-use the variable name new_file_name here, btw. Pick a different name

# HINT 1: You'll want the function range() for the for loop
# HINT 2: You can either have range go from 1 to 3 OR have range go from 0 to 2 and just add one when you create the string
# HINT 3: I recommend building the string and then appending it (use two lines) rather than doing it in one fell swoop 
#    (easier to debug)
...

# The autograder is going to check the following:
#  That the list has n_to_do elements in it
#  That the first element in the list is fname1.csv
#  That the last element in the list is fname3.csv

# GUIDE Write your own tests here using assert to check the above three things

```

```python
grader.check("Lists")
```

## Tutorials F and G: dictionaries

These are minor variations on the practice problems in *g_practice_dictionary.ipynb*. If you can't answer them, go back to **f_tutorial_dictionary** tutorial (there's a link to an on-line tutorial in there that is also helpful).

```python
# GUIDE: Create a dictionary that has 3 key-pairs in it
#     key: "name", value "Me"
#     key: "age", value 10
#     key: "grades", value a list with three elements,  "A", "A", "A"
# NOTE: The autograder will not pass until you do both this bit and the next


# GUIDE: Create my_simple_dict
#   You can either create the dictionary all at once or just add each key-value pair in turn
my_simple_dict = {}

# Remember that string comparison is case-sensitive
assert my_simple_dict["name"] == "Me"
assert my_simple_dict["age"] == 10
assert len(my_simple_dict["grades"]) == 3
for grade in my_simple_dict["grades"]:
    assert grade == "A"
```

```python
# GUIDE: 
#    - Add one to the age
#    - Add another "A" grade to the list
# Do these by editing the existing key-value pairs

# replace the ... with your code
...

assert my_simple_dict["age"] == 11
assert len(my_simple_dict["grades"]) == 4
```

```python
grader.check("Dictionaries")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS W26 (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5
```

```python
grader.check("hours_collaborators")
```

### To submit

Read me!!! This last cell will always have hand-in instructions. They're pretty boring (most of the time) but around week 5 you'll need to do more than just hand in this .ipynb. It is your responsibility to check if the Gradescope autograder ran correctly.

- Hit clear all outputs (this removes all the outputs)
- Hit Restart (say yes when it asks if it's ok to blow everything away)
- Hit 'Run All' to make sure that your code executes from top to bottom from scratch. 

If you hand in a file without doing this we will deduct points (the little numbers to the left of the first cell should be 1 and all cells executed)

- Make sure the file has saved (there is no little black dot to the right of the file name in the top bar)
- Submit this .ipynb file to Gradescope. Do NOT submit the tutorial files - Gradescope expects 1, and only 1, file with an .ipynb extension.
- Go get a cup of coffee (or tea or whatever). Come back and check that your autograder score is 5/5

If the Gradescope autograder fails (you did not get 5/5 and you know everything worked), please check here first for common reasons for it to fail
    https://docs.google.com/presentation/d/1tYa5oycUiG4YhXUq5vHvPOpWJ4k_xUPp2rUNIL7Q9RI/edit?usp=sharing

Make sure you remove all the print statements you put in that print out lots of stuff. You will get deducted for this, too, if the TAs have to wade through pages of print out stuff. One or two lines is ok.

Tests for Dictionaries (points: 1)
```
>>> assert my_simple_dict['name'] == 'Me'
```
```
>>> assert my_simple_dict['age'] == 11
```
```
>>> assert len(my_simple_dict['grades']) == 4
```
```
>>> for grade in my_simple_dict['grades']:
...     assert grade == 'A'
```


Tests for Lists (points: 1)
```
>>> assert len(my_list_of_file_names) == n_to_do
```
```
>>> assert my_list_of_file_names[0] == 'fname1.csv'
```
```
>>> assert my_list_of_file_names[-1] == 'fname' + str(n_to_do) + '.csv'
```


Tests for Strings (points: 1)
```
>>> assert data_file_name == 'data1.csv'
```
```
>>> assert new_file_name == 'fname1.csv'
```


Tests for hours_collaborators (points: 1)
```
>>> assert not 'not filled out' in worked_with_names
```
```
>>> assert not 'not filled out' in websites
```
```
>>> assert hours > 0
```
