```python
# Initialize Otter
import otter
grader = otter.Notebook("HW_5_pinball_classes.ipynb")
```

# Instructions



```python
# Doing the imports for you
import numpy as np
import json as json
from scipy.stats import linregress
import matplotlib.pyplot as plt
```

# Question 1

```python
# Code question 1
```

```python
grader.check("Question 1")
```

# Instructions


```python
# Code question 2
```

```python
grader.check("Question 2")
```

# Question 3 instructions


```python
# code question 3
```

```python
grader.check("Question 3")
```

# Question 4 instructions


```python
# code question 3
```

```python
grader.check("Question 3")
```

# Question 5 instructions


```python
# code question 3
```

```python
grader.check("Question 3")
```

# Question 6 instructions


```python
# code question 3
```

```python
grader.check("Question 3")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, hwk 5 pinball classes

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for Question 1 (points: 1)
```
>>> # Code test 1
```


Tests for Question 2 (points: 2)
```
>>> # Code test 2
```


Tests for Question 3 (points: 2)
```
>>> # Code test 3
```
```
>>> # Code test 3
```
```
>>> # Code test 3
```
```
>>> # Code test 3
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```
