```python
# Initialize Otter
import otter
grader = otter.Notebook("HW_5_pinball_classes.ipynb")
```

# Pinball Classs 

One more class - the pinball class that does the simulation

In this homework you'll take the classes you created in the lecture activity and lab (**PinballWall** and **PinballBumper**) and use them in a class that wraps up the entire simulation (**Pinball**)

Slides: https://docs.google.com/presentation/d/1DCSykIMBSlzCNNnmUcxPqfGxJmeGYd-7LF_W_-_7fUU/edit?usp=sharing

```python
# The usual imports
import numpy as np
import json as json
```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

# Import PinballWall and PinballBumper and make sure they work

TODO import the class definitions and run the test cases for each of those classes

```python

```

```python
# Assumes you copied your wall class into lec_act_9_class_pinball_wall.py
from lec_act_9_class_pinball_wall import PinballWall
from Lab_9_class_pinball_bumper import PinballBumper
```

```python
# TODO: Create at least one instance of each class and run their test functions to make sure they work
#. (you're free to create more than one, but at least create one wall and one bumper with the given names)

check_wall_instance = ...
check_bumper_instance = ...
"""
```

```python
grader.check("wall_and_bumper")
```

# Pinball class

TODO this question: Copy your **compute_next_step** code into the method, and check that it works

```python
class Pinball:
    """ Walls and bumpers for a rather boring pinball game"""
    def __init__(self, left, right, height, delta_t=0.01):
        """ Just set size of pinball box - assumes floor is y = 0
        Note: We'll add the walls and bumpers later - but we're going to create the variables we need here
        (It's very bad form to add new variables later) """
        self.left = left
        self.right = right
        self.height = height

        # What delta t to use
        self.delta_t = delta_t

        # If you're doing drag...
        self.do_drag = False

        # We're going to add any walls/bumpers here
        self.obstacles = []

        # For storing the path
        self.poses = None
        self.velocities = None

    @staticmethod
    def acceleration_due_to_gravity():
        """Somewhat silly - but if we need to change it, then we can  change it just here
        Note the static method tag - this says we don't need a self pointer"""
        gravity = -9.8     # m/s
        return gravity

    def add_obstacle(self, obstacle):
        """ Add in a bumper or wall
        @param obstacle - an instance of a bumper or a wall"""
        # Putting this as a separate method because it works better to make and add then try to pass everything
        #  into the init function. More sophisticated versions would also allow removing/changing obstacles
        self.obstacles.append(obstacle)

    def _compute_next_step(self, current_state):
        """ How to compute the next position and velocity from this one
        Note that I'm passing in current state and returning next here, rather than storing and updating an internal
        variable. This makes it clearer what the method does
        The _ in front of the method signals to someone using this class that this method is "private" and probably
        shouldn't be called
        @param current_state - the pose (x, y) and velocity (vx, vy) and acceleration (ax, ay) as a numpy array
        @return the new position, velocity as a tuple"""

        # TODO: Copy in your compute next step from pinball_routines. The only thing you should have to change is that
        #  you're now going to get delta t from self instead of passing it in
        result = np.zeros(current_state.shape)
        ...

        return result

    def simulate_pinball(self, starting_state):
        """ Call compute one time step multiple times and store it in a numpy array
        This should be your simulate code, again with delta_t and do_drag replaced with self.delta_t, as well
        as compute_next_step
        @param starting_state - the starting position, velocity, acceleration
        @return position & velocity values as two 2xtimesteps numpy array
        """

        # TODO - put your simulate function here. Only pass in the starting state - the remainder of the
        #  data you use to pass in should be in the self. pointer

        # The returned array.We do not know the size, so do not pre-allocate
        self.poses = []
        self.velocities = []

        # Use a while loop instead of the for loop
        # Set the stopping criteria based on current state y value
        # We know the first pose is the initial one
        # Notice that the poses are being stored in the self variable
        # Use this code to do the inside/outside & reflect for each wall
            # This should work, and will replace your if/then code for walls and bumpers, if your classes are
            #   implemented correctly
            # for o in self.obstacles:
            #     if o.outside(current_state[0, :]):
            #         pt_back, vel_back = o.reflect(current_state[0, :], current_state[1, :])
            #         current_state[0, :] = np.array(pt_back).transpose()
            #         current_state[1, :] = np.array(vel_back).transpose()

        ...
        # It's ok to return the self results
        return self.poses, self.velocities

    def plot_pinball_hw(self):
        """ plot the results of running the system AND the "correct" closed form result
        Note that everything we need is already in the class
        """
        nrows = 1
        ncols = 1
        fig, axs = plt.subplots(nrows, ncols, figsize=(4, 4))

        # The values we calculated in calculate_n_time_steps
        axs.plot([self.left, self.right], [self.height, self.height], color='gray', linestyle='dotted')
        axs.plot([self.left, self.right], [0, 0], color='gray', linestyle='dotted')
        axs.plot([self.left, self.left], [0, self.height], color='gray', linestyle='dotted')
        axs.plot([self.right, self.right], [0, self.height], color='gray', linestyle='dotted')

        # Again, this should work if your wall and bumper plots are correct
        for o in self.obstacles:
            o.plot(axs, self.left, self.right, self.height)

        axs.plot(self.poses[0, 0], self.poses[1, 0], 'xr', label="Start")
        axs.plot(self.poses[0, :], self.poses[1, :], '.-k', label="Poses")

        axs.axis('equal')
        axs.set_title(f"Not so boring pinball, 0-{self.delta_t * self.poses.shape[0]} s")

```

```python
# TODO Create the pinball class
#   Use: left -5.0, right 5.0, height 10.0, delta_t 0.1

pinball = ...
# END PROMPT"""
```

```python
# Check one step with the following start parameters
starting_state = np.zeros([3, 2])  # meters
starting_state[0, :] = [0, 0] # Start at zero, zero
# Velocity - mostly up with a bit of x
starting_state[1, :] = [-0.25, 5.0]
# Acceleration is really boring
starting_state[2, :] = [0.0, Pinball.acceleration_due_to_gravity()]

# TODO: Check first time step works
first_time_step = ...
print(f"Checking first time step {first_time_step}")
"""

# Notice my_pinball.delta_t - this is the delta t that you save and store in the pinball class
assert(np.all(np.isclose(first_time_step[0, :], starting_state[0, :] + my_pinball.delta_t * starting_state[1, :])))
assert(np.all(np.isclose(first_time_step[1, :], starting_state[1, :] + my_pinball.delta_t * starting_state[2, :])))
```

```python
grader.check("pinball_class")
```

# Simulate pinball method

TODO: Copy your simulate code into the **simulate_pinball** method in the class
- step 1: leave out the walls and the bumpers and check that you still stop correctly and simulate the poses correctly
- step 2: Use the code given in the todo section to add in looping over all of the obstacles, calling the **inside** method followed by the **reflect** method
- step 3: Check that calling simulate_pinball still works
- step 4: Add in a wall and make the ball bounce off of the wall
- step 5: Repeat for other walls and bumpers

Note: The plot code should work

```python
# TODO Now create and add any obstacles
# Required: Side walls at -4.5, 4.5 and top at 9.5
# Required: At least one bumper


# This creates an instance of a wall and adds it to the pinball
my_wall = PinballWall("Vertial", -4.5)
pinball.add_obstacle(my_wall)
# END PROMPT"""
```

```python
# Now run a simulation
#  setup
starting_state = np.zeros([3, 2])  # location, velocity, acceleration
starting_state[0, :] = [0, 0] # Start at zero, zero
# Velocity - mostly up with a bit of x noise
starting_state[1, :] = [2.4, 8.5]
starting_state[2, :] = [0.0, pinball.acceleration_due_to_gravity()]

# TODO: Run the simulation
...

```

```python
# ... and plot the results
pinball.plot_pinball_hw()
```

```python
grader.check("add_simulation")
```

# Question 4 instructions


```python
# code question 3
```

```python
grader.check("Question 3")
```

# Question 5 instructions


```python
# code question 3
```

```python
grader.check("Question 3")
```

# Question 6 instructions


```python
# code question 3
```

```python
grader.check("Question 3")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, hwk 5 pinball classes

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for Question 3 (points: 2)
```
>>> # Code test 3
```
```
>>> # Code test 3
```
```
>>> # Code test 3
```


Tests for add_simulation (points: 10)
```
>>> # Code test 3
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for pinball_class (points: 5)
```
>>> assert(np.all(np.isclose(first_time_step[0, :], starting_state[0, :] + my_pinball.delta_t * starting_state[1, :])))
```
```
>>> assert(np.all(np.isclose(first_time_step[1, :], starting_state[1, :] + my_pinball.delta_t * starting_state[2, :])))
```


Tests for wall_and_bumper (points: 5)
```
>>> assert(check_wall_instance.inside([0.0, 0.0]))
```
```
>>> assert(not check_bumper_instance.inside([-100, -100]))
```
