```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_9_class_pinball_wall.ipynb")
```

# Creating a pinball wall class

Extend the wall class from the tutorial to do two new things
- Reflect off of the wall [vertical/horizontal, general is optional]
- Plot self

The code is in **pinball_wall.py**

```python
# Doing the imports for you
import numpy as np
import matplotlib.pyplot as plt
```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

# Add reflect to wall code

In this problem, fill in **reflect** method in **PinballWall** in **pinball_wall.py**.

You're welcome to create more class variables in the **__init__** function or add more methods

TODO: 
- reflect (this question)
- test_reflect (next question)
- plot (last question - optional, but you'll need it for the homework

```python

```

```python
# You can import classes just like you import functions
from pinball_wall import PinballWall
```

```python
# Create instances and call test functions
# Now create two instances of the class
# Notice that we do NOT pass in a self variable
my_vert_wall = PinballWall(wall_type="Vertical", intercept_value=.2)
my_horiz_wall = PinballWall(wall_type="Horizontal", intercept_value=.2)

# Print them both - this will call the string function
print(f"Vertical wall: {my_vert_wall}")
print(f"Horizontal wall: {my_horiz_wall}")

# TODO check your reflect function here
```

```python
grader.check("class_defn")
```

# Write a test reflect function

TODO: Create some tests in the **test_reflect** method. 

```python
# Call the test function
assert(my_vert_wall.test_reflect())
```

```python
grader.check("test_reflect")
```

# Fill in the plot function 

TODO: Fill in the plot function

This is optional for the lecture activity, but you will need to do it for the homework.

Slanted lines are also optional for the entire assignment

```python
def plot_check():
    nrows = 1
    ncols = 1
    fig, axs = plt.subplots(nrows, ncols, figsize=(4, 4))

    # The values we calculated in calculate_n_time_steps
    left = -4.0
    right = 4.0
    height = 10.0

    axs.plot([left, right], [height, height], color='grey', linestyle="dotted")
    axs.plot([left, right], [0, 0], color='grey', linestyle="dotted")
    axs.plot([left, left], [0, height], color='grey', linestyle="dotted")
    axs.plot([right, right], [0, height], color='grey', linestyle="dotted")

    walls = []
    walls.append(PinballWall("Vertical", -3.0))
    walls.append(PinballWall("Vertical", 3.0))
    walls.append(PinballWall("Horizontal", 8.0))
    walls.append(PinballWall("General", a_b_c=[0.5, 0.5, 1.0]))
    walls.append(PinballWall("General", a_b_c=[-0.5, 0.5, 1.0]))
    for w in walls:
        w.plot(axs, left, right, height)

    axs.axis('equal')
    axs.set_title(f"Pinball walls")
```

```python
plot_check()
```

```python
grader.check("Plot_wall")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, week 9 lecture activity

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for Plot_wall (points: 0)
```
>>> plot_check()
```


Tests for class_defn (points: 2)
```
>>> assert(my_vert_wall.inside(my_vert_wall.reflect([0.25, 0.0], [1, 0])[0]))
```
```
>>> assert(my_vert_wall.reflect([0.25, 0.0], [1, 0])[1][0] < 0)
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for test_reflect (points: 2)
```
>>> assert(my_vert_wall.test_reflect())
```
