```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_9_class_pinball_wall.ipynb")
```

# Creating classes Example (walls) 

Extend the wall class from the tutorial to do two new things
- Reflect off of the wall [vertical/horizontal, general is optional]
- Plot self

Notice the method names are **inside** and **reflect** - we'll see why when you do the homework

```python
# Doing the imports for you
import numpy as np
import matplotlib.pyplot as plt
```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

# Add reflect to wall code

I've copied over the class code from the tutorial. In this problem, fill in **reflect**.

You're welcome to create more class variables in the init function or add more methods

TODO: 
- reflect (this question)
- test_reflect (next question)
- plot (last question - optional, but you'll need it for the homework

```python

```

```python
class PinballWall:
    def __init__(self, wall_type="Vertical", intercept_value=None, a_b_c=None):
        """Create a horizontal, vertical, or general half-plane wall
         if vertical/horizontal, use intercept value
         if general, use a_b_c
        @param wall_type - Vertical, Horizontal or General (ax + by + c = 0)
        @param intercept_value - number, where to place the wall on the x or y axis (Vertical or Horizontal)
        @param a_b_c - for General wall, list of 3 values"""

        # Save what kind of wall it is, then save as a general ax + by + c equation
        # What this does:
        #  - checks that the input is at least somewhat correct - eg, if they asked for a vertical wall they
        #    need to specify the intercept value
        # Lets the user "special case" a vertical or horizontal wall, rather than having to convert to the
        #   ax + by + c format
        if wall_type == "Vertical":
            self.wall_type = "Vertical"
            if intercept_value is None:
                raise ValueError("If vertical wall, need to specify intercept value")
            if intercept_value > 0:
                self.abc = [1.0, 0.0, -intercept_value]
            else:
                self.abc = [-1.0, 0.0, intercept_value]
        elif wall_type == "Horizontal":
            self.wall_type = "Horizontal"
            if intercept_value is None:
                raise ValueError("If horizontal wall, need to specify intercept value")
            if intercept_value > 0:
                self.abc = [0.0, 1.0, -intercept_value]
            else:
                self.abc = [0.0, -1.0, intercept_value]
        elif wall_type == "General":
            self.wall_type = "General"
            if a_b_c is None:
                raise ValueError("If general wall, need to specify a, b, and c a_b_c")
            if a_b_c[0] * 0 + a_b_c[1] * 0 + a_b_c[2] < 0:
                self.abc = [a_b_c[0], a_b_c[1], a_b_c[2]]
            else:
                self.abc = [-a_b_c[0], -a_b_c[1], -a_b_c[2]]
        else:
            raise ValueError("Wall type is not one of Vertical, Horizontal, or General")

    def evaluate_halfplane(self, x_y):
        """ Evaluate ax + by + c
        @param x_y numpy array/tuple for the current location
        @return ax + by + c"""
        # self is the same self we set in the __init__ function
        return self.abc[0] * x_y[0] + self.abc[1] * x_y[1] + self.abc[2]

    def outside(self, x_y):
        """ Is the point outside of this half/plane representing the wall?
        @param x_y numpy array/tuple for the current location
        @return True or False"""
        # We can call any of the methods defined in the class
        return self.evaluate_halfplane(x_y) >= 0.0

    def inside(self, x_y):
        """ We can call class methods from within class methods
        @param x_y numpy array/tuple for the current location
        @return True or False"""
        return not self.outside(x_y)

    def reflect(self, x_y, vx_vy):
        """ Reflect off of the wall
        @param x_y - the point
        @param vx_vy - the current velocity
        @return pt, vec - reflected point and vector"""
        # TODO - put your wall reflection code here
        #  Remember that you have what type of wall (vertical/horizontal) stored
        #  This should just be moving the code you did last week into this method
        #    pinball_routines inside/outside
        ...
        # TODO: Fix this...
        return x_y, vx_vy

    def __str__(self):
        """
        The string representation of this class.
        :return: a string
        """
        # f"" creates a string; in this case, instead of printing it, just return it
        return f"{self.abc[0]:0.3f}x + {self.abc[1]:0.3f}y + {self.abc[2]:0.3f}"

    def test_on_wall(self, x_y):
        """ Pass in a point that should be on the wall; should return 0
         @param x_y - xy point
         @preturn True or False"""
        return np.isclose(self.evaluate_halfplane(x_y), 0.0)

    def test_origin_inside(self):
        """ Test that we correctly oriented the wall so that (0,0) is inside"""
        # Notice call to self
        assert(self.inside([0, 0]))
        return True

    def test_reflect(self):
        """ Put any tests for reflection here
        @return True if they all pass"""
        print("Testing reflection by...")
        ...
        return False

    def plot(self, axs, left, right, height):
        """ Plot the wall inside the box defined by left, right, 0-height
        @param axs - the plot axes
        @param left left side of box
        @param right right side of box
        @param height height of box"""
        # TODO: Add plot code
        #   Note - you're free to save more information in the class init function to make this easier (eg, the y
        #    intercept)
        # Plotting the general wall is optional
        #   Hint: get x values between the left and right walls, solve for y = (-ax -c)/b and keep only the points
        #    inside the box
        print("Doing plot...")
        ...
```

```python
# Create instances and call test functions
# Now create two instances of the class
# Notice that we do NOT pass in a self variable
my_vert_wall = PinballWall(wall_type="Vertical", intercept_value=.2)
my_horiz_wall = PinballWall(wall_type="Horizontal", intercept_value=.2)

# Print them both - this will call the string function
print(f"Vertical wall: {my_vert_wall}")
print(f"Horizontal wall: {my_horiz_wall}")

# TODO check your reflect function here
```

```python
grader.check("class_defn")
```

# Write a test reflect function

TODO: Create some tests in the **test_reflect** method. 

```python
# Call the test function
assert(my_vert_wall.test_reflect())
```

```python
grader.check("test_reflect")
```

# Fill in the plot function 

TODO: Fill in the plot function

This is optional for the lecture activity, but you will need to do it for the homework.

Slanted lines are also optional for the entire assignment

```python
def plot_check():
    nrows = 1
    ncols = 1
    fig, axs = plt.subplots(nrows, ncols, figsize=(4, 4))

    # The values we calculated in calculate_n_time_steps
    left = -4.0
    right = 4.0
    height = 10.0

    axs.plot([left, right], [height, height], color='grey', linestyle="dotted")
    axs.plot([left, right], [0, 0], color='grey', linestyle="dotted")
    axs.plot([left, left], [0, height], color='grey', linestyle="dotted")
    axs.plot([right, right], [0, height], color='grey', linestyle="dotted")

    walls = []
    walls.append(PinballWall("Vertical", -3.0))
    walls.append(PinballWall("Vertical", 3.0))
    walls.append(PinballWall("Horizontal", 8.0))
    walls.append(PinballWall("General", a_b_c=[0.5, 0.5, 1.0]))
    walls.append(PinballWall("General", a_b_c=[-0.5, 0.5, 1.0]))
    for w in walls:
        w.plot(axs, left, right, height)

    axs.axis('equal')
    axs.set_title(f"Pinball walls")
```

```python
plot_check()
```

```python
grader.check("Plot_wall")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, week 9 lecture activity

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for Plot_wall (points: 0)
```
>>> plot_check()
```


Tests for class_defn (points: 2)
```
>>> assert(my_vert_wall.inside(my_vert_wall.reflect([0.25, 0.0], [1, 0])[0]))
```
```
>>> assert(my_vert_wall.reflect([0.25, 0.0], [1, 0])[1][0] < 0)
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```


Tests for test_reflect (points: 2)
```
>>> assert(my_vert_wall.test_reflect())
```
