```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab_6_optimization.ipynb")
```

# Setting up a function for fmin

Set up the function that we will be passing to fmin

Slides: https://docs.google.com/presentation/d/1m-8UphLgGLVaDKKcZ-NA_ejgOFlkD1Rlz3_sW0fHBM8/edit?usp=sharing

```python
# The usual imports
import numpy as np
import matplotlib.pyplot as plt

```

```python
# Imports from week 5
# This bit of code tells Jupyter notebooks to also look in the week_6_optimization folder for files
import os
import sys
sys.path.insert(0, os.path.abspath('../Week_5_matrices'))

# Routines from HW 3, Lab 5
# These are the routines used in the lab and/or you need to fill in for the homework (lab 5, week 5)
from arm_routines import create_arm_geometry, set_angles_of_arm_geometry, get_gripper_location
```

```python
# These are the routines used in the lab and/or you need to fill in for the homework (lab 6, week 6)
from arm_optimization_routines import vector_to_goal, distance_to_goal, func_for_fmin, angles_list_to_numpy
```

```python
# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

## Distance from gripper locatin to target

TODO: Edit vector_to_goal and distance_to_goal in arm_optimization_routines.py

```python
# Create the arm geometry
base_size_param = (1.0, 0.5)
link_sizes_param = [(0.5, 0.25), (0.3, 0.1), (0.2, 0.05)]
palm_width_param = 0.1
finger_size_param = (0.075, 0.025)

arm_geometry = create_arm_geometry(base_size_param, link_sizes_param, palm_width_param, finger_size_param)

# Set the angles of the arm
angles_gripper_check = [np.pi/6.0, -np.pi/4, 1.5 * np.pi/4, [np.pi/3.0, -np.pi/8.0, np.pi/6.0]]
set_angles_of_arm_geometry(arm_geometry, angles_gripper_check)
```

```python
# Get the gripper location
x, y = get_gripper_location(arm_geometry)
print(f"Gripper location ({x}, {y}), should be (-0.383, 1.323)")
```

```python
target = np.array([0.75, 1.25])
vx, vy = vector_to_goal(arm_geometry, target)
print(f"Vector to goal ({vx}, {vy}), should be (1.133, -0.073)")
```

```python
d_to_target = distance_to_goal(arm_geometry, target)
print(f"Distance to goal {d_to_target}, should be 1.1353")
```

```python
grader.check("dist_func")
```

# Edit the function that will be passed to fmin

TODO: edit **func_for_fmin** to take in a numpy array of angles and return the distance

```python
# This is a helper function to convert the list of angles in angles_gripper into a numpy array, which is
#  what fmin works with
angles_numpy = angles_list_to_numpy(angles_gripper_check)
check_dist = func_for_fmin(angles_numpy, arm_geometry, target)
print(f"Check func for min {check_dist}, should be 1.1353")

```

```python
grader.check("Func")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, week 6 lab. Don't forget arm_routines.py and arm_optimization_routines.py

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for Func (points: 5)
```
>>> assert(np.isclose(check_dist, 1.1353, atol=0.1))
```


Tests for dist_func (points: 4)
```
>>> assert(np.isclose(vx, 1.133, atol=0.01))
```
```
>>> assert(np.isclose(vy, -0.073, atol=0.01))
```
```
>>> assert(np.isclose(d_to_target, 1.1353, atol=0.01))
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```
