```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_6_optimization.ipynb")
```

```python
import numpy as np
import matplotlib.pyplot as plt
# Slightly different import format here - this gets just those routines by name (no mr.)
from matrix_routines import make_scale_matrix, make_translation_matrix, make_rotation_matrix
from object_routines import read_object, make_object_by_clicking, plot_object_in_world_coord_system

# These commands will force JN to actually re-load the external file when you re-execute the import command
%load_ext autoreload
%autoreload 2
```

<!-- BEGIN QUESTION -->

# Make scene

TODO: Edit the code to do "more" (see requirements)

```python
def make_scene():
    # initial "scene"
    # TODO do something else
    sq = read_object("Square")
    star = read_object("Star")

    fig, axs = plt.subplots(1, 1)

    plot_object_in_world_coord_system(axs, sq)
    # Option 1: Store the matrix in the object before calling plot
    cols = ["azure", "aqua", "cyan", "skyblue", "dodgerblue"]
    for col, dt in zip(cols, np.linspace(0, 2 * np.pi, 5)):
        star["Color"] = col
        star["Matrix"] = make_translation_matrix(np.cos(dt), np.sin(dt))
        plot_object_in_world_coord_system(axs, star)

    # Option 2: Use the global matrix to move the object around
    star["Color"] = "red"
    star["Matrix"] = np.identity(3)
    for ds in [0.25, 0.5, 1.0, 1.25]:
        plot_object_in_world_coord_system(axs, star, make_scale_matrix(ds, 1.0))

```

```python
make_scene()
```

```python
%matplotlib notebook
use_name = "name"
make_object_by_clicking(use_name)
```

<!-- END QUESTION -->

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, week 6 pre-lecture

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```
