```python
# Initialize Otter
import otter
grader = otter.Notebook("lec_act_6_fmin.ipynb")
```

```python
# Usual imports
from scipy.optimize import fmin
import matplotlib.pyplot as plt
import numpy as np
```

# Change the starting location to get the other local minima

Slides: https://docs.google.com/presentation/d/1xKhWyquiP-_wM9AnJJ57BBtF0eitEqSdODq34ttMZgk/edit?usp=sharing

TODO: Set up fmin so that you get the minima at x = 0

```python
# This function has several local minima - how to find the best one?
def my_multiple_minima_func(x):
    """A function with multiple minima
    @param x the input x value
    @returns ax^2 + bx + c"""
    return (3.0 * x - 4.0) * (0.5 * x + 0.2) * (x - 3.0) * (x - 3.0) + 2.0 * x + 0.5
```

```python
# Change this so you get the global minima at x = 0
x_at_min, f_at_x_min, _, _, _ = fmin(my_multiple_minima_func, 4.0, full_output=True)
print(f"Minimum of f is {f_at_x_min}, happens at x={x_at_min[0]}")
```

```python
# Plot the function and the minima
fig, axs = plt.subplots(1, 1, figsize=(4, 4))
ts = np.linspace(-2.0, 5.0)
axs.plot(ts, my_multiple_minima_func(ts), linestyle='-', label="Local minima")
axs.plot(x_at_min, f_at_x_min, 'Xr')

axs.set_xlabel('x')
axs.set_ylabel('y')

axs.set_title("Minimum of quartic: Local minima")
```

```python
grader.check("fmin_start_location")
```

# Be less accurate (but faster)

TODO: Change xtol and/or ftol until the total number of iterations is under 10

```python
# TODO: Change xtol and ftol until the number of iterations is under 10
# (Also play around with the values to see how LONG you can make it run...
x_not_as_accurate_min, f_not_as_accurate_min, iters, _, _ = fmin(my_multiple_minima_func, 4.0, xtol=0.0001, ftol=0.0001, full_output=True)
...

```

```python
fig2, axs2 = plt.subplots(1, 1, figsize=(4, 4))
ts = np.linspace(2.6, 2.9)
axs2.plot(ts, my_multiple_minima_func(ts), linestyle='-', label="Zoom in on curve")

x_at_min, f_at_x_min, _, _, _ = fmin(my_multiple_minima_func, 4.0, full_output=True)
axs2.plot(x_at_min, f_at_x_min, 'Xr', label="Accurate minima")
print(f"Minimum of f is {f_at_x_min}, happens at x={x_at_min[0]}")

axs2.plot(x_not_as_accurate_min, f_not_as_accurate_min, 'Xg', label="Not Accurate minima")

axs2.set_xlabel('x')
axs2.set_ylabel('y')
axs2.legend()

axs2.set_title("Minimum of quartic: Zoom in on Local minima")
```

```python
grader.check("fmin_iter")
```

## Hours and collaborators
Required for every assignment - fill out before you hand-in.

Listing names and websites helps you to document who you worked with and what internet help you received in the case of any plagiarism issues. You should list names of anyone (in class or not) who has substantially helped you with an assignment - or anyone you have *helped*. You do not need to list TAs.

Listing hours helps us track if the assignments are too long.

```python

# List of names (creates a set)
worked_with_names = {"not filled out"}
# List of URLS (creates a set)
websites = {"not filled out"}
# Approximate number of hours, including lab/in-class time
hours = -1.5

# for all row, column in all_indices_from_where
#.   if this is the column for wrist torque 
#.      print(f"Row: {r}, Time step: {c // n_time_steps} Successful y/n: {pick_data[r, -1] == 1}, value: {pick_data[r, c]}")
```

```python
grader.check("hours_collaborators")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Submit through gradescope, week 6 lecture activity

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for fmin_iter (points: 2)
```
>>> assert(iters < 10)
```
```
>>> assert np.isclose(f_at_x_min, f_not_as_accurate_min, atol=0.1)
```


Tests for fmin_start_location (points: 1)
```
>>> assert(np.isclose(x_at_min, 0.1545625, atol=0.0001))
```
```
>>> assert(np.isclose(f_at_x_min, -7.1299378, atol=0.01))
```


Tests for hours_collaborators (points: 1)
```
>>> assert(not "not filled out" in worked_with_names)
```
```
>>> assert(not "not filled out" in websites)
```
```
>>> assert(hours > 0)
```
